[PP] final: 26. s(X0) != s(X1) | X0 = X1 (0:8) T [term algebras injectivity]
[PP] final: 25. zero != s(X0) (0:4) T [term algebras distinctness]
[PP] final: 52. zero = X0 | s(p(X0)) = X0 (0:8) [cnf transformation 24]
[PP] final: 19. $less(X0,X1) | $less(X1,X0) | X0 = X1 (0:9) T [theory axiom]
[PP] final: 18. ~$less(X0,X1) | ~$less(X1,X2) | $less(X0,X2) (0:9) T [theory axiom]
[PP] final: 17. ~$less(X0,X0) (0:3) T [theory axiom]
[PP] final: 56. v != a(main_end,sK0) (0:5) [cnf transformation 46]
[PP] final: 55. ~$less(sK0,0) (0:3) [cnf transformation 46]
[PP] final: 54. $less(sK0,alength) (0:3) [cnf transformation 46]
[PP] final: 53. ~$less(alength,0) (0:3) [cnf transformation 46]
[PP] final: 80. ~'Sub'(X1,X1) (0:3) [equality resolution 57]
[PP] final: 58. X0 = X1 | i(l8(X0)) != i(l8(X1)) (0:10) [cnf transformation 37]
[PP] final: 59. 'Sub'(s(X0),s(X1)) | ~'Sub'(X0,X1) (0:8) [cnf transformation 38]
[PP] final: 60. alength = i(l8(nl8)) (0:5) [cnf transformation 6]
[PP] final: 62. a(l8(nl8),X1) = a(l8(X0),X1) | ~'Sub'(X0,s(nl8)) | a(l8(sK1(X0,X1)),X1) != a(l8(s(sK1(X0,X1))),X1) (0:27) [cnf transformation 48]
[PP] final: 61. a(l8(nl8),X1) = a(l8(X0),X1) | ~'Sub'(X0,s(nl8)) | 'Sub'(X0,sK1(X0,X1)) (0:18) [cnf transformation 48]
[PP] final: 64. i(l8(sK2(X0))) = X0 | $less(X0,0) | ~$less(X0,alength) (0:12) [cnf transformation 51]
[PP] final: 63. 'Sub'(sK2(X0),nl8) | $less(X0,0) | ~$less(X0,alength) (0:10) [cnf transformation 51]
[PP] final: 77. i(l8(zero)) = 0 (0:5) [cnf transformation 43]
[PP] final: 76. a(l7,X0) = a(l8(zero),X0) (0:8) [cnf transformation 43]
[PP] final: 75. i(l10(X1)) = i(l8(X1)) (0:7) [cnf transformation 43]
[PP] final: 74. a(l10(X1),X2) = a(l8(X1),X2) (0:9) [cnf transformation 43]
[PP] final: 73. v = a(l11(X3),i(l10(X3))) (0:8) [cnf transformation 43]
[PP] final: 72. a(l11(X3),X4) = a(l10(X3),X4) | i(l10(X3)) = X4 (0:14) [cnf transformation 43]
[PP] final: 71. i(l11(X3)) = i(l10(X3)) (0:7) [cnf transformation 43]
[PP] final: 70. i(l8(s(X3))) = $sum(i(l11(X3)),1) (0:10) [cnf transformation 43]
[PP] final: 69. a(l11(X3),X5) = a(l8(s(X3)),X5) (0:10) [cnf transformation 43]
[PP] final: 68. $less(i(l8(X6)),alength) | ~'Sub'(X6,nl8) (0:8) [cnf transformation 43]
[PP] final: 67. ~$less(i(l8(nl8)),alength) (0:5) [cnf transformation 43]
[PP] final: 66. i(main_end) = i(l8(nl8)) (0:6) [cnf transformation 43]
[PP] final: 65. a(main_end,X7) = a(l8(nl8),X7) (0:8) [cnf transformation 43]
[PP] final: 78. 'Sub'(X0,s(X1)) | ~'Sub'(X0,X1) (0:7) [cnf transformation 44]
[PP] final: 79. 'Sub'(X0,s(X0)) (0:4) [cnf transformation 1]
preprocessing finished
[PP] onPreprocessingEnd(), Proving Helper
[SA] passive: 26. s(X0) != s(X1) | X0 = X1 (0:8) T [term algebras injectivity]
[SA] passive: 25. zero != s(X0) (0:4) T [term algebras distinctness]
[SA] passive: 52. zero = X0 | s(p(X0)) = X0 (0:8) [cnf transformation 24]
[SA] passive: 19. $less(X0,X1) | $less(X1,X0) | X0 = X1 (0:9) T [theory axiom]
[SA] passive: 18. ~$less(X0,X1) | ~$less(X1,X2) | $less(X0,X2) (0:9) T [theory axiom]
[SA] passive: 17. ~$less(X0,X0) (0:3) T [theory axiom]
[SA] passive: 56. v != a(main_end,sK0) (0:5) [cnf transformation 46]
[SA] passive: 55. ~$less(sK0,0) (0:3) [cnf transformation 46]
[SA] passive: 54. $less(sK0,alength) (0:3) [cnf transformation 46]
[SA] passive: 53. ~$less(alength,0) (0:3) [cnf transformation 46]
[SA] passive: 80. ~'Sub'(X1,X1) (0:3) [equality resolution 57]
[SA] passive: 58. X0 = X1 | i(l8(X0)) != i(l8(X1)) (0:10) [cnf transformation 37]
[SA] passive: 59. 'Sub'(s(X0),s(X1)) | ~'Sub'(X0,X1) (0:8) [cnf transformation 38]
[SA] passive: 60. alength = i(l8(nl8)) (0:5) [cnf transformation 6]
[SA] passive: 62. a(l8(nl8),X1) = a(l8(X0),X1) | ~'Sub'(X0,s(nl8)) | a(l8(sK1(X0,X1)),X1) != a(l8(s(sK1(X0,X1))),X1) (0:27) [cnf transformation 48]
[SA] passive: 61. a(l8(nl8),X1) = a(l8(X0),X1) | ~'Sub'(X0,s(nl8)) | 'Sub'(X0,sK1(X0,X1)) (0:18) [cnf transformation 48]
[SA] passive: 64. i(l8(sK2(X0))) = X0 | $less(X0,0) | ~$less(X0,alength) (0:12) [cnf transformation 51]
[SA] passive: 63. 'Sub'(sK2(X0),nl8) | $less(X0,0) | ~$less(X0,alength) (0:10) [cnf transformation 51]
[SA] passive: 77. i(l8(zero)) = 0 (0:5) [cnf transformation 43]
[SA] passive: 76. a(l7,X0) = a(l8(zero),X0) (0:8) [cnf transformation 43]
[SA] passive: 75. i(l10(X1)) = i(l8(X1)) (0:7) [cnf transformation 43]
[SA] passive: 81. i(l10(sK2(X0))) = X0 | $less(X0,0) | ~$less(X0,alength) (0:12) [backward demodulation 75,64]
[SA] passive: 74. a(l10(X1),X2) = a(l8(X1),X2) (0:9) [cnf transformation 43]
[SA] passive: 84. a(l10(nl8),X1) = a(l10(X0),X1) | ~'Sub'(X0,s(nl8)) | 'Sub'(X0,sK1(X0,X1)) (0:18) [forward demodulation 82,74]
[SA] passive: 87. a(l10(sK1(X0,X1)),X1) != a(l10(s(sK1(X0,X1))),X1) | a(l10(nl8),X1) = a(l10(X0),X1) | ~'Sub'(X0,s(nl8)) (0:27) [forward demodulation 86,74]
[SA] passive: 73. v = a(l11(X3),i(l10(X3))) (0:8) [cnf transformation 43]
[SA] passive: 72. a(l11(X3),X4) = a(l10(X3),X4) | i(l10(X3)) = X4 (0:14) [cnf transformation 43]
[SA] passive: 71. i(l11(X3)) = i(l10(X3)) (0:7) [cnf transformation 43]
[SA] passive: 88. i(l11(X3)) = X4 | a(l11(X3),X4) = a(l10(X3),X4) (0:14) [backward demodulation 71,72]
[SA] passive: 89. v = a(l11(X3),i(l11(X3))) (0:8) [backward demodulation 71,73]
[SA] passive: 90. i(l11(X1)) = i(l8(X1)) (0:7) [backward demodulation 71,75]
[SA] passive: 91. i(l11(sK2(X0))) = X0 | $less(X0,0) | ~$less(X0,alength) (0:12) [backward demodulation 71,81]
[SA] passive: 92. i(l11(s(X3))) = $sum(i(l11(X3)),1) (0:10) [forward demodulation 70,90]
[SA] passive: 69. a(l11(X3),X5) = a(l8(s(X3)),X5) (0:10) [cnf transformation 43]
[SA] passive: 93. $less(i(l11(X6)),alength) | ~'Sub'(X6,nl8) (0:8) [forward demodulation 68,90]
[SA] passive: 95. alength = i(main_end) (0:4) [forward demodulation 66,60]
[SA] passive: 65. a(main_end,X7) = a(l8(nl8),X7) (0:8) [cnf transformation 43]
[SA] passive: 78. 'Sub'(X0,s(X1)) | ~'Sub'(X0,X1) (0:7) [cnf transformation 44]
[SA] passive: 79. 'Sub'(X0,s(X0)) (0:4) [cnf transformation 1]
[SA] active: 53. ~$less(alength,0) (0:3:1) [cnf transformation 46]
[SA] active: 54. $less(sK0,alength) (0:3:1) [cnf transformation 46]
[SA] active: 55. ~$less(sK0,0) (0:3:1) [cnf transformation 46]
[SA] active: 80. ~'Sub'(X1,X1) (0:3:1) [equality resolution 57]
[SA] active: 17. ~$less(X0,X0) (0:3:1) T [theory axiom]
[SA] active: 79. 'Sub'(X0,s(X0)) (0:4:1) [cnf transformation 1]
[SA] active: 95. alength = i(main_end) (0:4:1) [forward demodulation 66,60]
[SA] active: 25. zero != s(X0) (0:4:1) T [term algebras distinctness]
[SA] active: 56. v != a(main_end,sK0) (0:5:1) [cnf transformation 46]
[SA] active: 60. alength = i(l8(nl8)) (0:5:1) [cnf transformation 6]
[SA] active: 77. i(l8(zero)) = 0 (0:5:1) [cnf transformation 43]
[SA] active: 71. i(l11(X3)) = i(l10(X3)) (0:7:1) [cnf transformation 43]
[SA] active: 78. 'Sub'(X0,s(X1)) | ~'Sub'(X0,X1) (0:7:1) [cnf transformation 44]
[SA] passive: 96. ~'Sub'(s(X0),X0) (1:4) [resolution 78,80]
[SA] active: 90. i(l11(X1)) = i(l8(X1)) (0:7:1) [backward demodulation 71,75]
[SA] passive: 97. alength = i(l11(nl8)) (1:5) [superposition 90,60]
[SA] passive: 98. i(l11(zero)) = 0 (1:5) [superposition 90,77]
[SA] active: 96. ~'Sub'(s(X0),X0) (1:4:1) [resolution 78,80]
[SA] passive: 101. ~'Sub'(s(s(X0)),X0) (2:5) [resolution 96,78]
[SA] active: 59. 'Sub'(s(X0),s(X1)) | ~'Sub'(X0,X1) (0:8:1) [cnf transformation 38]
[SA] active: 97. alength = i(l11(nl8)) (1:5:1) [superposition 90,60]
[SA] active: 65. a(main_end,X7) = a(l8(nl8),X7) (0:8:1) [cnf transformation 43]
[SA] active: 98. i(l11(zero)) = 0 (1:5:1) [superposition 90,77]
[SA] active: 76. a(l7,X0) = a(l8(zero),X0) (0:8:1) [cnf transformation 43]
[SA] active: 101. ~'Sub'(s(s(X0)),X0) (2:5:1) [resolution 96,78]
[SA] passive: 105. ~'Sub'(s(s(s(X1))),X1) (3:6) [resolution 101,78]
[SA] active: 89. v = a(l11(X3),i(l11(X3))) (0:8:1) [backward demodulation 71,73]
[SA] passive: 106. v = a(l11(nl8),alength) (2:6) [superposition 89,97]
[SA] passive: 107. v = a(l11(zero),0) (2:6) [superposition 89,98]
[SA] active: 106. v = a(l11(nl8),alength) (2:6:1) [superposition 89,97]
[SA] active: 93. ~'Sub'(X6,nl8) | $less(i(l11(X6)),alength) (0:8:1) [forward demodulation 68,90]
[SA] active: 107. v = a(l11(zero),0) (2:6:1) [superposition 89,98]
[SA] active: 26. s(X0) != s(X1) | X0 = X1 (0:8:1) T [term algebras injectivity]
[SA] active: 105. ~'Sub'(s(s(s(X1))),X1) (3:6:1) [resolution 101,78]
[SA] passive: 110. ~'Sub'(s(s(s(s(X1)))),X1) (4:7) [resolution 105,78]
[SA] active: 52. s(p(X0)) = X0 | zero = X0 (0:8:1) [cnf transformation 24]
[SA] passive: 112. s(X2) != X1 | p(X1) = X2 | zero = X1 (1:11) [superposition 26,52]
[SA] passive: 114. 'Sub'(X5,s(X6)) | ~'Sub'(p(X5),X6) | zero = X5 (1:11) [superposition 59,52]
[SA] passive: 115. 'Sub'(s(X8),X7) | ~'Sub'(X8,p(X7)) | zero = X7 (1:11) [superposition 59,52]
[SA] passive: 116. 'Sub'(X10,X9) | ~'Sub'(X10,p(X9)) | zero = X9 (1:10) [superposition 78,52]
[SA] passive: 117. 'Sub'(p(X11),X11) | zero = X11 (1:7) [superposition 79,52]
[SA] passive: 118. ~'Sub'(X12,p(X12)) | zero = X12 (2:7) [superposition 96,52]
[SA] passive: 119. ~'Sub'(s(X13),p(X13)) | zero = X13 (3:8) [superposition 101,52]
[SA] passive: 120. ~'Sub'(s(s(X14)),p(X14)) | zero = X14 (4:9) [superposition 105,52]
[SA] active: 117. 'Sub'(p(X11),X11) | zero = X11 (1:7:1) [superposition 79,52]
[SA] passive: 121. zero = nl8 | $less(i(l11(p(nl8))),alength) (2:9) [resolution 117,93]
[SA] active: 74. a(l10(X1),X2) = a(l8(X1),X2) (0:9:1) [cnf transformation 43]
[SA] passive: 122. a(main_end,X0) = a(l10(nl8),X0) (1:8) [superposition 74,65]
[SA] passive: 126. a(main_end,X1) = a(l10(X0),X1) | a(l10(sK1(X0,X1)),X1) != a(l10(s(sK1(X0,X1))),X1) | ~'Sub'(X0,s(nl8)) (0:26) [backward demodulation 122,87]
[SA] passive: 127. a(main_end,X1) = a(l10(X0),X1) | ~'Sub'(X0,s(nl8)) | 'Sub'(X0,sK1(X0,X1)) (0:17) [backward demodulation 122,84]
[SA] passive: 123. a(l7,X1) = a(l10(zero),X1) (1:8) [superposition 74,76]
[SA] active: 118. ~'Sub'(X12,p(X12)) | zero = X12 (2:7:1) [superposition 96,52]
[SA] active: 18. ~$less(X1,X2) | ~$less(X0,X1) | $less(X0,X2) (0:9:1) T [theory axiom]
[SA] passive: 128. ~$less(X0,sK0) | $less(X0,alength) (1:6) [resolution 18,54]
[SA] active: 128. ~$less(X0,sK0) | $less(X0,alength) (1:6:1) [resolution 18,54]
[SA] active: 19. $less(X1,X0) | $less(X0,X1) | X0 = X1 (0:9:2) T [theory axiom]
[SA] passive: 130. $less(X1,X2) | X1 = X2 | ~$less(X3,X2) | $less(X3,X1) (1:12) T [resolution 19,18]
[SA] passive: 131. $less(0,alength) | alength = 0 (1:6) [resolution 19,53]
[SA] passive: 132. $less(0,sK0) | 0 = sK0 (1:6) [resolution 19,55]
[SA] passive: 133. $less(sK0,X4) | sK0 = X4 | $less(X4,alength) (2:9) [resolution 19,128]
[SA] active: 131. $less(0,alength) | alength = 0 (1:6:1) [resolution 19,53]
[SA] passive: 140. alength = 0 | ~$less(X0,0) | $less(X0,alength) (2:9) [resolution 131,18]
[SA] active: 58. i(l8(X0)) != i(l8(X1)) | X0 = X1 (0:10:1) [cnf transformation 37]
[SA] passive: 141. alength != i(l8(X0)) | nl8 = X0 (1:8) [superposition 58,60]
[SA] passive: 142. i(l8(X1)) != 0 | zero = X1 (1:8) [superposition 58,77]
[SA] passive: 143. i(l11(X2)) != i(l8(X3)) | X2 = X3 (1:10) [superposition 58,90]
[SA] active: 132. $less(0,sK0) | 0 = sK0 (1:6:1) [resolution 19,55]
[SA] passive: 148. 0 = sK0 | $less(0,alength) (2:6) [resolution 132,128]
[SA] passive: 149. 0 = sK0 | ~$less(X0,0) | $less(X0,sK0) (2:9) [resolution 132,18]
[SA] active: 63. 'Sub'(sK2(X0),nl8) | $less(X0,0) | ~$less(X0,alength) (0:10:1) [cnf transformation 51]
[SA] passive: 150. $less(X0,0) | ~$less(X0,alength) | $less(i(l11(sK2(X0))),alength) (1:12) [resolution 63,93]
[SA] active: 148. $less(0,alength) | 0 = sK0 (2:6:1) [resolution 132,128]
[SA] passive: 151. 0 = sK0 | ~$less(X0,0) | $less(X0,alength) (3:9) [resolution 148,18]
[SA] active: 69. a(l11(X3),X5) = a(l8(s(X3)),X5) (0:10:1) [cnf transformation 43]
[SA] passive: 155. a(l10(X0),X1) = a(l11(p(X0)),X1) | zero = X0 (1:13) [forward demodulation 152,74]
[SA] passive: 153. a(l11(X2),X3) = a(l10(s(X2)),X3) (1:10) [superposition 69,74]
[SA] passive: 156. a(l11(sK1(X0,X1)),X1) != a(l10(sK1(X0,X1)),X1) | a(main_end,X1) = a(l10(X0),X1) | ~'Sub'(X0,s(nl8)) (0:25) [backward demodulation 153,126]
[SA] active: 110. ~'Sub'(s(s(s(s(X1)))),X1) (4:7:1) [resolution 105,78]
[SA] passive: 158. ~'Sub'(s(s(s(s(s(X1))))),X1) (5:8) [resolution 110,78]
[SA] passive: 159. ~'Sub'(s(s(s(X0))),p(X0)) | zero = X0 (5:10) [superposition 110,52]
[SA] active: 92. i(l11(s(X3))) = $sum(i(l11(X3)),1) (0:10:1) [forward demodulation 70,90]
[SA] passive: 160. i(l11(s(nl8))) = $sum(alength,1) (2:8) [superposition 92,97]
[SA] passive: 162. i(l11(s(zero))) = 1 (2:6) [evaluation 161]
[SA] active: 162. i(l11(s(zero))) = 1 (2:6:1) [evaluation 161]
[SA] passive: 165. i(l11(s(s(zero)))) = 2 (3:7) [evaluation 163]
[SA] passive: 164. v = a(l11(s(zero)),1) (3:7) [superposition 89,162]
[SA] active: 91. ~$less(X0,alength) | $less(X0,0) | i(l11(sK2(X0))) = X0 (0:12:1) [backward demodulation 71,81]
[SA] passive: 173. i(l11(sK2(sK0))) = sK0 (1:6) [subsumption resolution 166,55]
[SA] passive: 167. $less(X0,0) | i(l11(sK2(X0))) = X0 | $less(alength,X0) | alength = X0 (1:15) [resolution 91,19]
[SA] passive: 172. i(l11(sK2(0))) = 0 | 0 = sK0 (3:9) [evaluation 169]
[SA] passive: 171. i(l11(sK2(0))) = 0 | alength = 0 (2:9) [evaluation 170]
[SA] active: 173. i(l11(sK2(sK0))) = sK0 (1:6:1) [subsumption resolution 166,55]
[SA] passive: 174. i(l11(s(sK2(sK0)))) = $sum(sK0,1) (2:9) [superposition 92,173]
[SA] passive: 175. v = a(l11(sK2(sK0)),sK0) (2:7) [superposition 89,173]
[SA] active: 88. a(l11(X3),X4) = a(l10(X3),X4) | i(l11(X3)) = X4 (0:14:1) [backward demodulation 71,72]
[SA] active: 175. v = a(l11(sK2(sK0)),sK0) (2:7:1) [superposition 89,173]
[SA] active: 127. 'Sub'(X0,sK1(X0,X1)) | ~'Sub'(X0,s(nl8)) | a(main_end,X1) = a(l10(X0),X1) (0:17:1) [backward demodulation 122,84]
[SA] active: 164. v = a(l11(s(zero)),1) (3:7:1) [superposition 89,162]
[SA] active: 156. a(l11(sK1(X0,X1)),X1) != a(l10(sK1(X0,X1)),X1) | a(main_end,X1) = a(l10(X0),X1) | ~'Sub'(X0,s(nl8)) (0:25:1) [backward demodulation 153,126]
[SA] passive: 177. a(main_end,X1) = a(l10(X0),X1) | ~'Sub'(X0,s(nl8)) | i(l11(sK1(X0,X1))) = X1 (1:19) [trivial inequality removal 176]
[SA] active: 165. i(l11(s(s(zero)))) = 2 (3:7:1) [evaluation 163]
[SA] passive: 180. i(l11(s(s(s(zero))))) = 3 (4:8) [evaluation 178]
[SA] passive: 179. v = a(l11(s(s(zero))),2) (4:8) [superposition 89,165]
[SA] active: 122. a(main_end,X0) = a(l10(nl8),X0) (1:8:1) [superposition 74,65]
[SA] passive: 183. alength = X0 | a(main_end,X0) = a(l11(nl8),X0) (2:11) [forward demodulation 181,97]
[SA] active: 123. a(l7,X1) = a(l10(zero),X1) (1:8:1) [superposition 74,76]
[SA] passive: 187. 0 = X0 | a(l7,X0) = a(l11(zero),X0) (2:11) [forward demodulation 185,98]
[SA] active: 141. alength != i(l8(X0)) | nl8 = X0 (1:8:1) [superposition 58,60]
[SA] passive: 190. alength != 0 | zero = nl8 (2:6) [superposition 141,77]
[SA] passive: 191. alength != i(l11(X0)) | nl8 = X0 (2:8) [superposition 141,90]
[SA] active: 190. alength != 0 | zero = nl8 (2:6:1) [superposition 141,77]
[SA] active: 142. i(l8(X1)) != 0 | zero = X1 (1:8:1) [superposition 58,77]
[SA] passive: 194. i(l11(X0)) != 0 | zero = X0 (2:8) [superposition 142,90]
[SA] active: 160. i(l11(s(nl8))) = $sum(alength,1) (2:8:1) [superposition 92,97]
[SA] passive: 195. i(l11(s(s(nl8)))) = $sum($sum(alength,1),1) (3:11) [superposition 92,160]
[SA] passive: 196. v = a(l11(s(nl8)),$sum(alength,1)) (3:9) [superposition 89,160]
[SA] active: 116. ~'Sub'(X10,p(X9)) | 'Sub'(X10,X9) | zero = X9 (1:10:1) [superposition 78,52]
[SA] passive: 197. 'Sub'(p(p(X0)),X0) | zero = X0 | zero = p(X0) (2:12) [resolution 116,117]
[SA] active: 191. alength != i(l11(X0)) | nl8 = X0 (2:8:1) [superposition 141,90]
[SA] passive: 199. alength != 1 | s(zero) = nl8 (3:7) [superposition 191,162]
[SA] passive: 200. alength != 2 | s(s(zero)) = nl8 (4:8) [superposition 191,165]
[SA] passive: 201. alength != $sum(alength,1) | s(nl8) = nl8 (3:9) [superposition 191,160]
[SA] passive: 203. alength != sK0 | nl8 = sK2(sK0) (3:7) [superposition 191,173]
[SA] active: 143. i(l11(X2)) != i(l8(X3)) | X2 = X3 (1:10:1) [superposition 58,90]
[SA] passive: 204. i(l8(X1)) != 1 | s(zero) = X1 (3:9) [superposition 143,162]
[SA] passive: 205. i(l8(X2)) != 2 | s(s(zero)) = X2 (4:10) [superposition 143,165]
[SA] passive: 206. i(l8(X3)) != $sum(alength,1) | s(nl8) = X3 (3:11) [superposition 143,160]
[SA] passive: 207. i(l8(X5)) != sK0 | sK2(sK0) = X5 (2:9) [superposition 143,173]
[SA] passive: 210. i(l11(X2)) != i(l11(X3)) | X2 = X3 (2:10) [superposition 143,90]
[SA] active: 203. alength != sK0 | nl8 = sK2(sK0) (3:7:1) [superposition 191,173]
[SA] active: 153. a(l11(X2),X3) = a(l10(s(X2)),X3) (1:10:1) [superposition 69,74]
[SA] passive: 212. a(l11(X2),X3) = a(l11(s(X2)),X3) | i(l11(s(X2))) = X3 (2:16) [superposition 153,88]
[SA] active: 199. alength != 1 | s(zero) = nl8 (3:7:1) [superposition 191,162]
[SA] active: 114. ~'Sub'(p(X5),X6) | 'Sub'(X5,s(X6)) | zero = X5 (1:11:1) [superposition 59,52]
[SA] passive: 215. 'Sub'(X1,s(s(p(X1)))) | zero = X1 (2:9) [resolution 114,79]
[SA] passive: 216. 'Sub'(X2,s(s(X3))) | zero = X2 | ~'Sub'(p(X2),X3) (2:12) [resolution 114,78]
[SA] passive: 217. 'Sub'(X4,s(sK1(p(X4),X5))) | zero = X4 | ~'Sub'(p(X4),s(nl8)) | a(main_end,X5) = a(l10(p(X4)),X5) (2:24) [resolution 114,127]
[SA] active: 194. i(l11(X0)) != 0 | zero = X0 (2:8:1) [superposition 142,90]
[SA] passive: 227. alength != -1 (3:3) [term algebras distinctness 226]
[SA] passive: 224. 0 != sK0 | zero = sK2(sK0) (3:7) [superposition 194,173]
[SA] active: 115. ~'Sub'(X8,p(X7)) | 'Sub'(s(X8),X7) | zero = X7 (1:11:1) [superposition 59,52]
[SA] passive: 228. 'Sub'(s(p(p(X0))),X0) | zero = X0 | zero = p(X0) (2:13) [resolution 115,117]
[SA] active: 227. alength != -1 (3:3:1) [term algebras distinctness 226]
[SA] active: 112. s(X2) != X1 | p(X1) = X2 | zero = X1 (1:11:1) [superposition 26,52]
[SA] passive: 229. X0 != X1 | p(X0) = p(X1) | zero = X1 | zero = X0 (2:14) [superposition 112,52]
[SA] passive: 231. p(s(X0)) = X0 (2:5) [term algebras distinctness 230]
[SA] active: 231. p(s(X0)) = X0 (2:5:1) [term algebras distinctness 230]
[SA] active: 150. $less(i(l11(sK2(X0))),alength) | ~$less(X0,alength) | $less(X0,0) (1:12:1) [resolution 63,93]
[SA] passive: 242. ~$less(X0,alength) | $less(X0,0) | $less(i(l11(sK2(X0))),0) | i(l11(sK2(X0))) = i(l11(sK2(i(l11(sK2(X0)))))) (2:24) [resolution 150,91]
[SA] passive: 243. ~$less(X1,alength) | $less(X1,0) | ~$less(X2,i(l11(sK2(X1)))) | $less(X2,alength) (2:15) [resolution 150,18]
[SA] active: 224. 0 != sK0 | zero = sK2(sK0) (3:7:1) [superposition 194,173]
[SA] active: 130. ~$less(X3,X2) | X1 = X2 | $less(X1,X2) | $less(X3,X1) (1:12:1) T [resolution 19,18]
[SA] passive: 245. X0 = X1 | $less(X0,X1) | $less(X2,X0) | $less(X1,X2) | X1 = X2 (2:15) T [resolution 130,19]
[SA] passive: 247. alength = X6 | $less(X6,alength) | $less(i(l11(sK2(X7))),X6) | ~$less(X7,alength) | $less(X7,0) (2:18) [resolution 130,150]
[SA] passive: 248. alength = X8 | $less(X8,alength) | $less(0,X8) | 0 = sK0 (3:12) [resolution 130,148]
[SA] passive: 249. alength = X9 | $less(X9,alength) | $less(0,X9) | alength = 0 (2:12) [resolution 130,131]
[SA] passive: 250. sK0 = X10 | $less(X10,sK0) | $less(0,X10) | 0 = sK0 (2:12) [resolution 130,132]
[SA] passive: 251. alength = X11 | $less(X11,alength) | $less(sK0,X11) (2:9) [resolution 130,54]
[SA] active: 119. ~'Sub'(s(X13),p(X13)) | zero = X13 (3:8:1) [superposition 101,52]
[SA] passive: 252. ~'Sub'(X0,p(p(X0))) | zero = p(X0) | zero = X0 (4:12) [superposition 119,52]
[SA] active: 155. a(l10(X0),X1) = a(l11(p(X0)),X1) | zero = X0 (1:13:1) [forward demodulation 152,74]
[SA] passive: 256. v = a(l10(X2),i(l11(p(X2)))) | zero = X2 (2:12) [superposition 155,89]
[SA] active: 179. v = a(l11(s(s(zero))),2) (4:8:1) [superposition 89,165]
[SA] active: 167. $less(alength,X0) | $less(X0,0) | i(l11(sK2(X0))) = X0 | alength = X0 (1:15:2) [resolution 91,19]
[SA] passive: 260. $less(X0,0) | i(l11(sK2(X0))) = X0 | alength = X0 | X0 = X1 | $less(X1,X0) | $less(alength,X1) (2:21) [resolution 167,130]
[SA] passive: 261. $less(X2,0) | i(l11(sK2(X2))) = X2 | alength = X2 | ~$less(X3,alength) | $less(X3,X2) (2:18) [resolution 167,18]
[SA] passive: 268. $less(alength,X0) | i(l11(sK2(X0))) = X0 | alength = X0 | 0 = X1 | $less(X1,0) | $less(X0,X1) (2:21) [resolution 167,130]
[SA] passive: 269. $less(alength,X2) | i(l11(sK2(X2))) = X2 | alength = X2 | ~$less(X3,X2) | $less(X3,0) (2:18) [resolution 167,18]
[SA] active: 180. i(l11(s(s(s(zero))))) = 3 (4:8:1) [evaluation 178]
[SA] passive: 273. i(l8(X0)) != 3 | s(s(s(zero))) = X0 (5:11) [superposition 143,180]
[SA] passive: 274. alength != 3 | s(s(s(zero))) = nl8 (5:9) [superposition 191,180]
[SA] passive: 277. i(l11(s(s(s(s(zero)))))) = 4 (5:9) [evaluation 275]
[SA] passive: 276. v = a(l11(s(s(s(zero)))),3) (5:9) [superposition 89,180]
[SA] active: 177. ~'Sub'(X0,s(nl8)) | a(main_end,X1) = a(l10(X0),X1) | i(l11(sK1(X0,X1))) = X1 (1:19:1) [trivial inequality removal 176]
[SA] passive: 279. a(main_end,X1) = a(l10(X2),X1) | i(l11(sK1(X2,X1))) = X1 | ~'Sub'(X2,nl8) (2:18) [resolution 177,78]
[SA] passive: 283. a(main_end,X3) = a(l11(X4),X3) | i(l11(sK1(s(X4),X3))) = X3 | ~'Sub'(X4,nl8) (2:19) [forward demodulation 280,153]
[SA] active: 200. alength != 2 | s(s(zero)) = nl8 (4:8:1) [superposition 191,165]
[SA] active: 133. $less(sK0,X4) | $less(X4,alength) | sK0 = X4 (2:9:2) [resolution 19,128]
[SA] passive: 286. $less(X0,alength) | sK0 = X0 | X0 = X1 | $less(X1,X0) | $less(sK0,X1) (3:15) [resolution 133,130]
[SA] passive: 287. $less(X2,alength) | sK0 = X2 | ~$less(X3,sK0) | $less(X3,X2) (3:12) [resolution 133,18]
[SA] passive: 291. $less(sK0,X0) | sK0 = X0 | $less(X0,0) | i(l11(sK2(X0))) = X0 (3:15) [resolution 133,91]
[SA] passive: 293. $less(sK0,X1) | sK0 = X1 | alength = X2 | $less(X2,alength) | $less(X1,X2) (3:15) [resolution 133,130]
[SA] passive: 294. $less(sK0,X3) | sK0 = X3 | ~$less(X4,X3) | $less(X4,alength) (3:12) [resolution 133,18]
[SA] active: 158. ~'Sub'(s(s(s(s(s(X1))))),X1) (5:8:1) [resolution 110,78]
[SA] passive: 297. ~'Sub'(s(s(s(s(s(s(X1)))))),X1) (6:9) [resolution 158,78]
[SA] passive: 298. ~'Sub'(s(s(s(s(X0)))),p(X0)) | zero = X0 (6:11) [superposition 158,52]
[SA] active: 140. ~$less(X0,0) | alength = 0 | $less(X0,alength) (2:9:1) [resolution 131,18]
[SA] passive: 300. alength = 0 | $less(X1,alength) | $less(0,X1) | 0 = X1 (3:12) [resolution 140,19]
[SA] active: 149. ~$less(X0,0) | 0 = sK0 | $less(X0,sK0) (2:9:1) [resolution 132,18]
[SA] passive: 306. 0 = sK0 | $less(X0,sK0) | $less(alength,X0) | i(l11(sK2(X0))) = X0 | alength = X0 (3:18) [resolution 149,167]
[SA] passive: 307. 0 = sK0 | $less(X1,sK0) | $less(0,X1) | 0 = X1 (3:12) [resolution 149,19]
[SA] active: 171. i(l11(sK2(0))) = 0 | alength = 0 (2:9:1) [evaluation 170]
[SA] passive: 320. zero = sK2(0) | alength = 0 (3:7) [trivial inequality removal 314]
[SA] passive: 315. i(l8(X0)) != 0 | sK2(0) = X0 | alength = 0 (3:12) [superposition 143,171]
[SA] passive: 319. i(l11(s(sK2(0)))) = 1 | alength = 0 (3:10) [evaluation 317]
[SA] passive: 318. v = a(l11(sK2(0)),0) | alength = 0 (3:10) [superposition 89,171]
[SA] active: 320. zero = sK2(0) | alength = 0 (3:7:1) [trivial inequality removal 314]
[SA] passive: 329. 'Sub'(zero,nl8) | alength = 0 (4:6) [subsumption resolution 324,131]
[SA] active: 174. i(l11(s(sK2(sK0)))) = $sum(sK0,1) (2:9:1) [superposition 92,173]
[SA] passive: 337. sK0 != -1 (3:3) [term algebras distinctness 336]
[SA] passive: 331. i(l8(X0)) != $sum(sK0,1) | s(sK2(sK0)) = X0 (3:12) [superposition 143,174]
[SA] passive: 332. alength != $sum(sK0,1) | s(sK2(sK0)) = nl8 (3:10) [superposition 191,174]
[SA] passive: 333. i(l11(s(s(sK2(sK0))))) = $sum($sum(sK0,1),1) (3:12) [superposition 92,174]
[SA] passive: 334. v = a(l11(s(sK2(sK0))),$sum(sK0,1)) (3:10) [superposition 89,174]
[SA] active: 337. sK0 != -1 (3:3:1) [term algebras distinctness 336]
[SA] active: 207. i(l8(X5)) != sK0 | sK2(sK0) = X5 (2:9:1) [superposition 143,173]
[SA] passive: 340. i(l11(X0)) != sK0 | sK2(sK0) = X0 (3:9) [superposition 207,90]
[SA] active: 329. 'Sub'(zero,nl8) | alength = 0 (4:6:1) [subsumption resolution 324,131]
[SA] active: 251. $less(sK0,X11) | $less(X11,alength) | alength = X11 (2:9:2) [resolution 130,54]
[SA] passive: 344. $less(X0,alength) | alength = X0 | X0 = X1 | $less(X1,X0) | $less(sK0,X1) (3:15) [resolution 251,130]
[SA] passive: 345. $less(X2,alength) | alength = X2 | ~$less(X3,sK0) | $less(X3,X2) (3:12) [resolution 251,18]
[SA] passive: 351. $less(sK0,X0) | alength = X0 | $less(X0,0) | i(l11(sK2(X0))) = X0 (3:15) [resolution 251,91]
[SA] passive: 353. $less(sK0,X1) | alength = X1 | alength = X2 | $less(X2,alength) | $less(X1,X2) (3:15) [resolution 251,130]
[SA] passive: 354. $less(sK0,X3) | alength = X3 | ~$less(X4,X3) | $less(X4,alength) (3:12) [resolution 251,18]
[SA] active: 121. $less(i(l11(p(nl8))),alength) | zero = nl8 (2:9:1) [resolution 117,93]
[SA] passive: 357. zero = nl8 | $less(i(l11(p(nl8))),0) | i(l11(p(nl8))) = i(l11(sK2(i(l11(p(nl8)))))) (3:21) [resolution 121,91]
[SA] passive: 358. zero = nl8 | alength = X0 | $less(X0,alength) | $less(i(l11(p(nl8))),X0) (3:15) [resolution 121,130]
[SA] passive: 359. zero = nl8 | ~$less(X1,i(l11(p(nl8)))) | $less(X1,alength) (3:12) [resolution 121,18]
[SA] active: 215. 'Sub'(X1,s(s(p(X1)))) | zero = X1 (2:9:1) [resolution 114,79]
[SA] passive: 360. zero = p(X0) | 'Sub'(X0,s(s(s(p(p(X0)))))) | zero = X0 (3:15) [resolution 215,114]
[SA] active: 151. ~$less(X0,0) | 0 = sK0 | $less(X0,alength) (3:9:1) [resolution 148,18]
[SA] passive: 366. 0 = sK0 | $less(X1,alength) | $less(0,X1) | 0 = X1 (4:12) [resolution 151,19]
[SA] active: 210. i(l11(X2)) != i(l11(X3)) | X2 = X3 (2:10:1) [superposition 143,90]
[SA] passive: 374. i(l11(X1)) != 1 | s(zero) = X1 (3:9) [superposition 210,162]
[SA] passive: 375. i(l11(X2)) != 2 | s(s(zero)) = X2 (4:10) [superposition 210,165]
[SA] passive: 376. i(l11(X3)) != 3 | s(s(s(zero))) = X3 (5:11) [superposition 210,180]
[SA] passive: 377. i(l11(X4)) != $sum(alength,1) | s(nl8) = X4 (3:11) [superposition 210,160]
[SA] passive: 378. i(l11(X5)) != $sum(sK0,1) | s(sK2(sK0)) = X5 (3:12) [superposition 210,174]
[SA] passive: 381. i(l11(X8)) != 0 | sK2(0) = X8 | alength = 0 (3:12) [superposition 210,171]
[SA] active: 172. i(l11(sK2(0))) = 0 | 0 = sK0 (3:9:1) [evaluation 169]
[SA] passive: 394. v = a(l11(sK2(0)),0) | 0 = sK0 (4:10) [superposition 89,172]
[SA] passive: 402. i(l11(s(sK2(0)))) = 1 | 0 = sK0 (4:10) [evaluation 395]
[SA] passive: 396. i(l8(X0)) != 0 | sK2(0) = X0 | 0 = sK0 (4:12) [superposition 143,172]
[SA] passive: 397. alength != 0 | nl8 = sK2(0) | 0 = sK0 (4:10) [superposition 191,172]
[SA] passive: 401. zero = sK2(0) | 0 = sK0 (4:7) [trivial inequality removal 398]
[SA] passive: 399. i(l11(X1)) != 0 | sK2(0) = X1 | 0 = sK0 (4:12) [superposition 210,172]
[SA] active: 183. a(main_end,X0) = a(l11(nl8),X0) | alength = X0 (2:11:1) [forward demodulation 181,97]
[SA] active: 401. zero = sK2(0) | 0 = sK0 (4:7:1) [trivial inequality removal 398]
[SA] passive: 416. 'Sub'(zero,nl8) | 0 = sK0 (5:6) [subsumption resolution 411,148]
[SA] active: 187. a(l7,X0) = a(l11(zero),X0) | 0 = X0 (2:11:1) [forward demodulation 185,98]
[SA] active: 416. 'Sub'(zero,nl8) | 0 = sK0 (5:6:1) [subsumption resolution 411,148]
[SA] active: 249. $less(0,X9) | $less(X9,alength) | alength = X9 | alength = 0 (2:12:2) [resolution 130,131]
[SA] passive: 424. $less(X0,alength) | alength = X0 | alength = 0 | X0 = X1 | $less(X1,X0) | $less(0,X1) (3:18) [resolution 249,130]
[SA] passive: 425. $less(X2,alength) | alength = X2 | alength = 0 | ~$less(X3,0) | $less(X3,X2) (3:15) [resolution 249,18]
[SA] passive: 431. $less(0,X0) | alength = X0 | alength = 0 | $less(X0,0) | i(l11(sK2(X0))) = X0 (3:18) [resolution 249,91]
[SA] passive: 433. $less(0,X1) | alength = X1 | alength = 0 | alength = X2 | $less(X2,alength) | $less(X1,X2) (3:18) [resolution 249,130]
[SA] passive: 434. $less(0,X3) | alength = X3 | alength = 0 | ~$less(X4,X3) | $less(X4,alength) (3:15) [resolution 249,18]
[SA] active: 340. i(l11(X0)) != sK0 | sK2(sK0) = X0 (3:9:1) [superposition 207,90]
[SA] passive: 440. 1 != sK0 | s(zero) = sK2(sK0) (4:8) [superposition 340,162]
[SA] passive: 441. sK0 != 2 | s(s(zero)) = sK2(sK0) (4:9) [superposition 340,165]
[SA] passive: 442. sK0 != 3 | s(s(s(zero))) = sK2(sK0) (5:10) [superposition 340,180]
[SA] passive: 443. $sum(alength,1) != sK0 | s(nl8) = sK2(sK0) (4:10) [superposition 340,160]
[SA] passive: 444. $sum(sK0,1) != sK0 | s(sK2(sK0)) = sK2(sK0) (4:11) [superposition 340,174]
[SA] passive: 448. 0 != sK0 | sK2(0) = sK2(sK0) | alength = 0 (4:11) [superposition 340,171]
[SA] active: 250. $less(0,X10) | $less(X10,sK0) | sK0 = X10 | 0 = sK0 (2:12:2) [resolution 130,132]
[SA] passive: 450. $less(X0,sK0) | sK0 = X0 | 0 = sK0 | X0 = X1 | $less(X1,X0) | $less(0,X1) (3:18) [resolution 250,130]
[SA] passive: 451. $less(X2,sK0) | sK0 = X2 | 0 = sK0 | ~$less(X3,0) | $less(X3,X2) (3:15) [resolution 250,18]
[SA] passive: 457. $less(0,X0) | sK0 = X0 | 0 = sK0 | $less(X0,alength) (3:12) [resolution 250,128]
[SA] passive: 459. $less(0,X1) | sK0 = X1 | 0 = sK0 | sK0 = X2 | $less(X2,sK0) | $less(X1,X2) (3:18) [resolution 250,130]
[SA] passive: 460. $less(0,X3) | sK0 = X3 | 0 = sK0 | ~$less(X4,X3) | $less(X4,sK0) (3:15) [resolution 250,18]
[SA] active: 440. 1 != sK0 | s(zero) = sK2(sK0) (4:8:1) [superposition 340,162]
[SA] active: 197. 'Sub'(p(p(X0)),X0) | zero = X0 | zero = p(X0) (2:12:1) [resolution 116,117]
[SA] passive: 474. zero = X0 | zero = p(X0) | 'Sub'(p(X0),s(X0)) (3:12) [duplicate literal removal 466]
[SA] passive: 477. i(l11(sK1(p(nl8),X1))) = X1 | a(main_end,X1) = a(l10(p(nl8)),X1) | zero = nl8 (3:20) [forward demodulation 476,231]
[SA] passive: 468. zero = p(X2) | zero = p(p(X2)) | 'Sub'(s(p(p(p(X2)))),X2) | zero = X2 (3:19) [resolution 197,115]
[SA] passive: 469. zero = p(X3) | zero = p(p(X3)) | 'Sub'(p(p(p(X3))),X3) | zero = X3 (3:18) [resolution 197,116]
[SA] passive: 470. zero = nl8 | zero = p(nl8) | $less(i(l11(p(p(nl8)))),alength) (3:14) [resolution 197,93]
[SA] passive: 478. zero = X0 | 'Sub'(p(X0),s(X0)) (3:8) [forward demodulation 472,231]
[SA] active: 478. 'Sub'(p(X0),s(X0)) | zero = X0 (3:8:1) [forward demodulation 472,231]
[SA] passive: 484. zero = X1 | 'Sub'(X1,s(s(X1))) (4:8) [duplicate literal removal 480]
[SA] passive: 483. 'Sub'(X0,s(s(X0))) (4:5) [term algebras distinctness 481]
[SA] active: 216. 'Sub'(X2,s(s(X3))) | zero = X2 | ~'Sub'(p(X2),X3) (2:12:1) [resolution 114,78]
[SA] passive: 491. zero = p(X6) | ~'Sub'(p(p(X6)),X7) | 'Sub'(X6,s(s(s(X7)))) | zero = X6 (3:18) [resolution 216,114]
[SA] passive: 492. 'Sub'(X1,s(X0)) | zero = X1 | ~'Sub'(p(X1),p(X0)) | zero = X0 (3:15) [superposition 216,52]
[SA] active: 483. 'Sub'(X0,s(s(X0))) (4:5:1) [term algebras distinctness 481]
[SA] passive: 505. 'Sub'(X0,s(s(s(p(X0))))) | zero = X0 (5:10) [resolution 483,114]
[SA] active: 256. v = a(l10(X2),i(l11(p(X2)))) | zero = X2 (2:12:1) [superposition 155,89]
[SA] passive: 518. v = a(l10(s(X0)),i(l11(X0))) (3:9) [term algebras distinctness 507]
[SA] passive: 508. v = a(l11(X1),i(l11(p(X1)))) | zero = X1 | i(l11(X1)) = i(l11(p(X1))) (3:20) [superposition 256,88]
[SA] passive: 510. v = a(main_end,i(l11(p(nl8)))) | zero = nl8 (3:11) [superposition 256,122]
[SA] active: 484. 'Sub'(X1,s(s(X1))) | zero = X1 (4:8:1) [duplicate literal removal 480]
[SA] active: 228. 'Sub'(s(p(p(X0))),X0) | zero = X0 | zero = p(X0) (2:13:1) [resolution 115,117]
[SA] passive: 535. i(l11(sK1(s(p(nl8)),X0))) = X0 | a(main_end,X0) = a(l11(p(nl8)),X0) | zero = nl8 (3:21) [forward demodulation 534,231]
[SA] passive: 524. zero = p(X1) | zero = p(p(X1)) | 'Sub'(s(s(p(p(p(X1))))),X1) | zero = X1 (3:20) [resolution 228,115]
[SA] passive: 526. zero = nl8 | zero = p(nl8) | $less(i(l11(s(p(p(nl8))))),alength) (3:15) [resolution 228,93]
[SA] passive: 536. zero = X0 | 'Sub'(s(p(X0)),s(X0)) (3:9) [forward demodulation 530,231]
[SA] active: 196. v = a(l11(s(nl8)),$sum(alength,1)) (3:9:1) [superposition 89,160]
[SA] active: 229. X0 != X1 | p(X0) = p(X1) | zero = X1 | zero = X0 (2:14:1) [superposition 112,52]
[SA] active: 201. alength != $sum(alength,1) | s(nl8) = nl8 (3:9:1) [superposition 191,160]
[SA] active: 243. ~$less(X2,i(l11(sK2(X1)))) | $less(X1,0) | ~$less(X1,alength) | $less(X2,alength) (2:15:1) [resolution 150,18]
[SA] passive: 539. $less(X0,0) | ~$less(X0,alength) | $less(X1,alength) | $less(i(l11(sK2(X0))),X1) | i(l11(sK2(X0))) = X1 (3:21) [resolution 243,19]
[SA] active: 204. i(l8(X1)) != 1 | s(zero) = X1 (3:9:1) [superposition 143,162]
[SA] active: 245. $less(X2,X0) | $less(X1,X2) | $less(X0,X1) | X0 = X1 | X1 = X2 (2:15:3) T [resolution 130,19]
[SA] passive: 563. $less(X2,X3) | $less(X4,X2) | X2 = X4 | X2 = X3 | X4 = X5 | $less(X5,X4) | $less(X3,X5) (3:21) T [resolution 245,130]
[SA] passive: 564. $less(X6,X7) | $less(X8,X6) | X6 = X8 | X6 = X7 | ~$less(X9,X7) | $less(X9,X8) (3:18) T [resolution 245,18]
[SA] passive: 565. $less(X10,X11) | $less(alength,X10) | alength = X10 | X10 = X11 | $less(X11,0) | i(l11(sK2(X11))) = X11 (3:21) [resolution 245,91]
[SA] passive: 566. $less(X12,X13) | $less(i(l11(sK2(X14))),X12) | i(l11(sK2(X14))) = X12 | X12 = X13 | $less(X14,0) | ~$less(X14,alength) | $less(X13,alength) (3:27) [resolution 245,243]
[SA] passive: 567. $less(X15,alength) | $less(0,X15) | 0 = X15 | alength = X15 (3:12) [resolution 245,53]
[SA] passive: 568. $less(X16,sK0) | $less(0,X16) | 0 = X16 | sK0 = X16 (3:12) [resolution 245,55]
[SA] passive: 569. $less(X17,X18) | $less(0,X17) | 0 = X17 | X17 = X18 | 0 = sK0 | $less(X18,alength) (4:18) [resolution 245,151]
[SA] passive: 570. $less(X19,X20) | $less(0,X19) | 0 = X19 | X19 = X20 | 0 = sK0 | $less(X20,sK0) (3:18) [resolution 245,149]
[SA] passive: 571. $less(X21,X22) | $less(0,X21) | 0 = X21 | X21 = X22 | alength = 0 | $less(X22,alength) (3:18) [resolution 245,140]
[SA] passive: 572. $less(X23,X24) | $less(sK0,X23) | sK0 = X23 | X23 = X24 | $less(X24,alength) (3:15) [resolution 245,128]
[SA] passive: 574. $less(X2,X3) | $less(X3,X4) | X3 = X4 | X2 = X4 | X2 = X5 | $less(X5,X2) | $less(X4,X5) (3:21) T [resolution 245,130]
[SA] passive: 575. $less(X6,X7) | $less(X7,X8) | X7 = X8 | X6 = X8 | ~$less(X9,X8) | $less(X9,X6) (3:18) T [resolution 245,18]
[SA] passive: 577. $less(i(l11(sK2(X12))),X13) | $less(X13,X14) | X13 = X14 | i(l11(sK2(X12))) = X14 | $less(X12,0) | ~$less(X12,alength) | $less(X14,alength) (3:27) [resolution 245,243]
[SA] passive: 580. $less(0,X17) | $less(X17,X18) | X17 = X18 | 0 = X18 | 0 = sK0 | $less(X18,alength) (4:18) [resolution 245,151]
[SA] passive: 581. $less(0,X19) | $less(X19,X20) | X19 = X20 | 0 = X20 | 0 = sK0 | $less(X20,sK0) (3:18) [resolution 245,149]
[SA] passive: 582. $less(0,X21) | $less(X21,X22) | X21 = X22 | 0 = X22 | alength = 0 | $less(X22,alength) (3:18) [resolution 245,140]
[SA] passive: 585. $less(X2,X3) | $less(X4,X2) | X3 = X4 | X2 = X4 | X4 = X5 | $less(X5,X4) | $less(X3,X5) (3:21) T [resolution 245,130]
[SA] passive: 586. $less(X6,X7) | $less(X8,X6) | X7 = X8 | X6 = X8 | ~$less(X9,X7) | $less(X9,X8) (3:18) T [resolution 245,18]
[SA] passive: 587. $less(X10,X11) | $less(alength,X10) | alength = X11 | alength = X10 | $less(X11,0) | i(l11(sK2(X11))) = X11 (3:21) [resolution 245,91]
[SA] passive: 588. $less(X12,X13) | $less(i(l11(sK2(X14))),X12) | i(l11(sK2(X14))) = X13 | i(l11(sK2(X14))) = X12 | $less(X14,0) | ~$less(X14,alength) | $less(X13,alength) (3:30) [resolution 245,243]
[SA] passive: 591. $less(X17,X18) | $less(0,X17) | 0 = X18 | 0 = X17 | 0 = sK0 | $less(X18,alength) (4:18) [resolution 245,151]
[SA] passive: 592. $less(X19,X20) | $less(0,X19) | 0 = X20 | 0 = X19 | 0 = sK0 | $less(X20,sK0) (3:18) [resolution 245,149]
[SA] passive: 593. $less(X21,X22) | $less(0,X21) | 0 = X22 | 0 = X21 | alength = 0 | $less(X22,alength) (3:18) [resolution 245,140]
[SA] passive: 594. $less(X23,X24) | $less(sK0,X23) | sK0 = X24 | sK0 = X23 | $less(X24,alength) (3:15) [resolution 245,128]
[SA] active: 374. i(l11(X1)) != 1 | s(zero) = X1 (3:9:1) [superposition 210,162]
[SA] active: 212. a(l11(X2),X3) = a(l11(s(X2)),X3) | i(l11(s(X2))) = X3 (2:16:1) [superposition 153,88]
[SA] passive: 618. a(l11(X0),X1) = a(l11(p(X0)),X1) | i(l11(s(p(X0)))) = X1 | zero = X0 (3:20) [superposition 212,52]
[SA] active: 518. v = a(l10(s(X0)),i(l11(X0))) (3:9:1) [term algebras distinctness 507]
[SA] passive: 628. v = a(l10(s(zero)),0) (4:7) [superposition 518,98]
[SA] passive: 629. v = a(l10(s(s(zero))),1) (4:8) [superposition 518,162]
[SA] passive: 630. v = a(l10(s(s(s(zero)))),2) (4:9) [superposition 518,165]
[SA] passive: 631. v = a(l10(s(s(s(s(zero))))),3) (5:10) [superposition 518,180]
[SA] passive: 632. v = a(l10(s(s(nl8))),$sum(alength,1)) (4:10) [superposition 518,160]
[SA] passive: 633. v = a(l10(s(s(sK2(sK0)))),$sum(sK0,1)) (4:11) [superposition 518,174]
[SA] passive: 634. v = a(l10(s(nl8)),alength) (4:7) [superposition 518,97]
[SA] passive: 635. v = a(l10(s(sK2(sK0))),sK0) (4:8) [superposition 518,173]
[SA] passive: 636. v = a(l10(s(sK2(0))),0) | 0 = sK0 (4:11) [superposition 518,172]
[SA] passive: 637. v = a(l10(s(sK2(0))),0) | alength = 0 (4:11) [superposition 518,171]
[SA] passive: 639. v = a(l11(s(X2)),i(l11(X2))) | i(l11(X2)) = i(l11(s(X2))) (4:17) [superposition 518,88]
[SA] active: 247. $less(i(l11(sK2(X7))),X6) | $less(X6,alength) | alength = X6 | ~$less(X7,alength) | $less(X7,0) (2:18:1) [resolution 130,150]
[SA] passive: 643. $less(X1,alength) | alength = X1 | ~$less(X2,alength) | $less(X2,0) | X1 = X3 | $less(X3,X1) | $less(i(l11(sK2(X2))),X3) (3:24) [resolution 247,130]
[SA] passive: 644. $less(X4,alength) | alength = X4 | ~$less(X5,alength) | $less(X5,0) | ~$less(X6,i(l11(sK2(X5)))) | $less(X6,X4) (3:21) [resolution 247,18]
[SA] active: 628. v = a(l10(s(zero)),0) (4:7:1) [superposition 518,98]
[SA] passive: 670. v = a(l11(s(zero)),0) (5:7) [evaluation 669]
[SA] active: 261. ~$less(X3,alength) | i(l11(sK2(X2))) = X2 | alength = X2 | $less(X2,0) | $less(X3,X2) (2:18:1) [resolution 167,18]
[SA] passive: 671. i(l11(sK2(X0))) = X0 | alength = X0 | $less(X0,0) | $less(X1,X0) | $less(0,X1) | alength = X1 | alength = 0 (3:24) [resolution 261,249]
[SA] passive: 672. i(l11(sK2(X2))) = X2 | alength = X2 | $less(X2,0) | $less(X3,X2) | $less(sK0,X3) | alength = X3 (3:21) [resolution 261,251]
[SA] passive: 673. i(l11(sK2(X4))) = X4 | alength = X4 | $less(X4,0) | $less(X5,X4) | $less(sK0,X5) | sK0 = X5 (3:21) [resolution 261,133]
[SA] passive: 674. i(l11(sK2(X6))) = X6 | alength = X6 | $less(X6,0) | $less(X7,X6) | $less(X8,X7) | $less(alength,X8) | alength = X7 | alength = X8 (3:27) [resolution 261,245]
[SA] passive: 675. i(l11(sK2(X9))) = X9 | alength = X9 | $less(X9,0) | $less(X10,X9) | $less(alength,X11) | $less(X11,X10) | X10 = X11 | alength = X10 (3:27) [resolution 261,245]
[SA] passive: 676. i(l11(sK2(X12))) = X12 | alength = X12 | $less(X12,0) | $less(X13,X12) | $less(X14,X13) | $less(alength,X14) | alength = X14 | X13 = X14 (3:27) [resolution 261,245]
[SA] passive: 680. i(l11(sK2(X20))) = X20 | alength = X20 | $less(X20,0) | $less(i(l11(sK2(X21))),X20) | ~$less(X21,alength) | $less(X21,0) (3:24) [resolution 261,150]
[SA] passive: 681. i(l11(sK2(X22))) = X22 | alength = X22 | $less(X22,0) | $less(i(l11(p(nl8))),X22) | zero = nl8 (3:21) [resolution 261,121]
[SA] passive: 683. i(l11(sK2(X25))) = X25 | alength = X25 | $less(X25,0) | $less(0,X25) | 0 = sK0 (3:18) [resolution 261,148]
[SA] active: 634. v = a(l10(s(nl8)),alength) (4:7:1) [superposition 518,97]
[SA] passive: 694. alength = $sum(alength,1) | v = a(l11(s(nl8)),alength) (5:12) [forward demodulation 691,160]
[SA] active: 269. ~$less(X3,X2) | i(l11(sK2(X2))) = X2 | alength = X2 | $less(alength,X2) | $less(X3,0) (2:18:1) [resolution 167,18]
[SA] passive: 696. i(l11(sK2(X0))) = X0 | alength = X0 | $less(alength,X0) | $less(X1,0) | $less(X2,X1) | $less(X0,X2) | X0 = X1 | X0 = X2 (3:27) [resolution 269,245]
[SA] passive: 697. i(l11(sK2(X3))) = X3 | alength = X3 | $less(alength,X3) | $less(X4,0) | $less(X3,X5) | $less(X5,X4) | X4 = X5 | X3 = X4 (3:27) [resolution 269,245]
[SA] passive: 698. i(l11(sK2(X6))) = X6 | alength = X6 | $less(alength,X6) | $less(X7,0) | $less(X8,X7) | $less(X6,X8) | X6 = X8 | X7 = X8 (3:27) [resolution 269,245]
[SA] passive: 699. i(l11(sK2(X9))) = X9 | alength = X9 | $less(alength,X9) | $less(X10,0) | $less(X9,X10) | X9 = X10 (3:21) [resolution 269,19]
[SA] active: 670. v = a(l11(s(zero)),0) (5:7:1) [evaluation 669]
[SA] active: 279. ~'Sub'(X2,nl8) | i(l11(sK1(X2,X1))) = X1 | a(main_end,X1) = a(l10(X2),X1) (2:18:1) [resolution 177,78]
[SA] passive: 727. i(l11(sK1(sK2(X0),X1))) = X1 | a(main_end,X1) = a(l10(sK2(X0)),X1) | $less(X0,0) | ~$less(X0,alength) (3:23) [resolution 279,63]
[SA] passive: 733. a(main_end,X2) = a(l7,X2) | i(l11(sK1(zero,X2))) = X2 | 0 = sK0 (6:17) [forward demodulation 728,123]
[SA] passive: 734. a(main_end,X3) = a(l7,X3) | i(l11(sK1(zero,X3))) = X3 | alength = 0 (5:17) [forward demodulation 729,123]
[SA] passive: 731. i(l11(sK1(p(p(nl8)),X5))) = X5 | a(main_end,X5) = a(l10(p(p(nl8))),X5) | zero = nl8 | zero = p(nl8) (3:26) [resolution 279,197]
[SA] passive: 735. a(main_end,X6) = a(l11(p(p(nl8))),X6) | i(l11(sK1(s(p(p(nl8))),X6))) = X6 | zero = nl8 | zero = p(nl8) (3:27) [forward demodulation 732,153]
[SA] active: 635. v = a(l10(s(sK2(sK0))),sK0) (4:8:1) [superposition 518,173]
[SA] passive: 740. $sum(sK0,1) = sK0 | v = a(l11(s(sK2(sK0))),sK0) (5:13) [forward demodulation 737,174]
[SA] active: 283. ~'Sub'(X4,nl8) | i(l11(sK1(s(X4),X3))) = X3 | a(main_end,X3) = a(l11(X4),X3) (2:19:1) [forward demodulation 280,153]
[SA] passive: 742. i(l11(sK1(s(sK2(X0)),X1))) = X1 | a(main_end,X1) = a(l11(sK2(X0)),X1) | $less(X0,0) | ~$less(X0,alength) (3:24) [resolution 283,63]
[SA] passive: 743. i(l11(sK1(s(zero),X2))) = X2 | a(main_end,X2) = a(l11(zero),X2) | 0 = sK0 (6:19) [resolution 283,416]
[SA] passive: 744. i(l11(sK1(s(zero),X3))) = X3 | a(main_end,X3) = a(l11(zero),X3) | alength = 0 (5:19) [resolution 283,329]
[SA] passive: 747. i(l11(sK1(s(s(p(p(nl8)))),X6))) = X6 | a(main_end,X6) = a(l11(s(p(p(nl8)))),X6) | zero = nl8 | zero = p(nl8) (3:29) [resolution 283,228]
[SA] active: 629. v = a(l10(s(s(zero))),1) (4:8:1) [superposition 518,162]
[SA] passive: 753. v = a(l11(s(s(zero))),1) (5:8) [evaluation 752]
[SA] active: 260. $less(alength,X1) | $less(X1,X0) | $less(X0,0) | X0 = X1 | alength = X0 | i(l11(sK2(X0))) = X0 (2:21:3) [resolution 167,130]
[SA] passive: 756. $less(X3,X4) | $less(X4,0) | X3 = X4 | alength = X4 | i(l11(sK2(X4))) = X4 | X3 = X5 | $less(X5,X3) | $less(alength,X5) (3:27) [resolution 260,130]
[SA] passive: 757. $less(X6,X7) | $less(X7,0) | X6 = X7 | alength = X7 | i(l11(sK2(X7))) = X7 | ~$less(X8,alength) | $less(X8,X6) (3:24) [resolution 260,18]
[SA] passive: 768. $less(alength,X3) | $less(X4,0) | X3 = X4 | alength = X4 | i(l11(sK2(X4))) = X4 | X4 = X5 | $less(X5,X4) | $less(X3,X5) (3:27) [resolution 260,130]
[SA] passive: 769. $less(alength,X6) | $less(X7,0) | X6 = X7 | alength = X7 | i(l11(sK2(X7))) = X7 | ~$less(X8,X6) | $less(X8,X7) (3:24) [resolution 260,18]
[SA] passive: 782. $less(alength,X4) | $less(X4,X5) | X4 = X5 | alength = X5 | i(l11(sK2(X5))) = X5 | 0 = sK0 | $less(X5,sK0) (3:24) [resolution 260,149]
[SA] passive: 786. $less(alength,X11) | $less(X11,X12) | X11 = X12 | alength = X12 | i(l11(sK2(X12))) = X12 | 0 = X13 | $less(X13,0) | $less(X12,X13) (3:27) [resolution 260,130]
[SA] passive: 787. $less(alength,X14) | $less(X14,X15) | X14 = X15 | alength = X15 | i(l11(sK2(X15))) = X15 | ~$less(X16,X15) | $less(X16,0) (3:24) [resolution 260,18]
[SA] active: 753. v = a(l11(s(s(zero))),1) (5:8:1) [evaluation 752]
[SA] active: 268. $less(alength,X0) | $less(X1,0) | $less(X0,X1) | 0 = X1 | alength = X0 | i(l11(sK2(X0))) = X0 (2:21:3) [resolution 167,130]
[SA] passive: 808. $less(X6,0) | $less(X7,X6) | 0 = X6 | alength = X7 | i(l11(sK2(X7))) = X7 | ~$less(X8,alength) | $less(X8,X7) (3:24) [resolution 268,18]
[SA] passive: 847. $less(X12,0) | $less(i(l11(sK2(X13))),X12) | 0 = X12 | alength = i(l11(sK2(X13))) | i(l11(sK2(X13))) = i(l11(sK2(i(l11(sK2(X13)))))) | $less(X13,0) | ~$less(X13,alength) (3:36) [subsumption resolution 811,17]
[SA] passive: 819. $less(alength,X2) | $less(X2,X3) | 0 = X3 | alength = X2 | i(l11(sK2(X2))) = X2 | 0 = sK0 | $less(X3,alength) (4:24) [resolution 268,151]
[SA] passive: 820. $less(alength,X4) | $less(X4,X5) | 0 = X5 | alength = X4 | i(l11(sK2(X4))) = X4 | 0 = sK0 | $less(X5,sK0) (3:24) [resolution 268,149]
[SA] passive: 821. $less(alength,X6) | $less(X6,X7) | 0 = X7 | alength = X6 | i(l11(sK2(X6))) = X6 | alength = 0 | $less(X7,alength) (3:24) [resolution 268,140]
[SA] passive: 824. $less(alength,X11) | $less(X11,X12) | 0 = X12 | alength = X11 | i(l11(sK2(X11))) = X11 | 0 = X13 | $less(X13,0) | $less(X12,X13) (3:27) [resolution 268,130]
[SA] passive: 825. $less(alength,X14) | $less(X14,X15) | 0 = X15 | alength = X14 | i(l11(sK2(X14))) = X14 | ~$less(X16,X15) | $less(X16,0) (3:24) [resolution 268,18]
[SA] passive: 828. $less(alength,X3) | $less(X4,0) | 0 = X4 | alength = X3 | i(l11(sK2(X3))) = X3 | X4 = X5 | $less(X5,X4) | $less(X3,X5) (3:27) [resolution 268,130]
[SA] passive: 829. $less(alength,X6) | $less(X7,0) | 0 = X7 | alength = X6 | i(l11(sK2(X6))) = X6 | ~$less(X8,X6) | $less(X8,X7) (3:24) [resolution 268,18]
[SA] active: 536. 'Sub'(s(p(X0)),s(X0)) | zero = X0 (3:9:1) [forward demodulation 530,231]
[SA] active: 217. 'Sub'(X4,s(sK1(p(X4),X5))) | zero = X4 | ~'Sub'(p(X4),s(nl8)) | a(main_end,X5) = a(l10(p(X4)),X5) (2:24:1) [resolution 114,127]
[SA] passive: 855. zero = p(X0) | ~'Sub'(p(p(X0)),s(nl8)) | a(main_end,X1) = a(l10(p(p(X0))),X1) | 'Sub'(X0,s(s(sK1(p(p(X0)),X1)))) | zero = X0 (3:32) [resolution 217,114]
[SA] passive: 859. a(main_end,X1) = a(l10(X0),X1) | ~'Sub'(X0,s(nl8)) | 'Sub'(s(X0),s(sK1(X0,X1))) (3:19) [forward demodulation 858,231]
[SA] active: 441. sK0 != 2 | s(s(zero)) = sK2(sK0) (4:9:1) [superposition 340,165]
[SA] active: 242. $less(i(l11(sK2(X0))),0) | $less(X0,0) | ~$less(X0,alength) | i(l11(sK2(X0))) = i(l11(sK2(i(l11(sK2(X0)))))) (2:24:1) [resolution 150,91]
[SA] passive: 861. $less(X1,0) | ~$less(X1,alength) | i(l11(sK2(X1))) = i(l11(sK2(i(l11(sK2(X1)))))) | 0 = sK0 | $less(i(l11(sK2(X1))),sK0) (3:27) [resolution 242,149]
[SA] passive: 864. $less(X4,0) | ~$less(X4,alength) | i(l11(sK2(X4))) = i(l11(sK2(i(l11(sK2(X4)))))) | 0 = X5 | $less(X5,0) | $less(i(l11(sK2(X4))),X5) (3:30) [resolution 242,130]
[SA] passive: 865. $less(X6,0) | ~$less(X6,alength) | i(l11(sK2(X6))) = i(l11(sK2(i(l11(sK2(X6)))))) | ~$less(X7,i(l11(sK2(X6)))) | $less(X7,0) (3:27) [resolution 242,18]
[SA] passive: 880. i(l11(sK2(0))) = i(l11(sK2(i(l11(sK2(0)))))) | 0 = sK0 (5:15) [evaluation 879]
[SA] passive: 883. i(l11(sK2(0))) = i(l11(sK2(i(l11(sK2(0)))))) | alength = 0 (4:15) [evaluation 882]
[SA] active: 120. ~'Sub'(s(s(X14)),p(X14)) | zero = X14 (4:9:1) [superposition 105,52]
[SA] passive: 887. ~'Sub'(s(X0),p(p(X0))) | zero = p(X0) | zero = X0 (5:13) [superposition 120,52]
[SA] active: 318. v = a(l11(sK2(0)),0) | alength = 0 (3:10:1) [superposition 89,171]
[SA] active: 630. v = a(l10(s(s(s(zero)))),2) (4:9:1) [superposition 518,165]
[SA] passive: 898. v = a(l11(s(s(s(zero)))),2) (5:9) [evaluation 897]
[SA] active: 319. i(l11(s(sK2(0)))) = 1 | alength = 0 (3:10:1) [evaluation 317]
[SA] passive: 901. v = a(l11(s(sK2(0))),1) | alength = 0 (4:11) [superposition 89,319]
[SA] passive: 913. i(l11(s(s(sK2(0))))) = 2 | alength = 0 (4:11) [evaluation 902]
[SA] passive: 903. i(l8(X0)) != 1 | s(sK2(0)) = X0 | alength = 0 (4:13) [superposition 143,319]
[SA] passive: 904. alength != 1 | s(sK2(0)) = nl8 | alength = 0 (4:11) [superposition 191,319]
[SA] passive: 906. i(l11(X1)) != 1 | s(sK2(0)) = X1 | alength = 0 (4:13) [superposition 210,319]
[SA] passive: 908. 1 != sK0 | s(sK2(0)) = sK2(sK0) | alength = 0 (4:12) [superposition 340,319]
[SA] passive: 910. v = a(l10(s(s(sK2(0)))),1) | alength = 0 (4:12) [superposition 518,319]
[SA] active: 274. alength != 3 | s(s(s(zero))) = nl8 (5:9:1) [superposition 191,180]
[SA] active: 332. alength != $sum(sK0,1) | s(sK2(sK0)) = nl8 (3:10:1) [superposition 191,174]
[SA] active: 276. v = a(l11(s(s(s(zero)))),3) (5:9:1) [superposition 89,180]
[SA] active: 334. v = a(l11(s(sK2(sK0))),$sum(sK0,1)) (3:10:1) [superposition 89,174]
[SA] active: 277. i(l11(s(s(s(s(zero)))))) = 4 (5:9:1) [evaluation 275]
[SA] passive: 919. v = a(l11(s(s(s(s(zero))))),4) (6:10) [superposition 89,277]
[SA] passive: 929. i(l11(s(s(s(s(s(zero))))))) = 5 (6:10) [evaluation 920]
[SA] passive: 921. i(l8(X0)) != 4 | s(s(s(s(zero)))) = X0 (6:12) [superposition 143,277]
[SA] passive: 922. alength != 4 | s(s(s(s(zero)))) = nl8 (6:10) [superposition 191,277]
[SA] passive: 924. i(l11(X1)) != 4 | s(s(s(s(zero)))) = X1 (6:12) [superposition 210,277]
[SA] passive: 926. sK0 != 4 | s(s(s(s(zero)))) = sK2(sK0) (6:11) [superposition 340,277]
[SA] passive: 928. v = a(l10(s(s(s(s(s(zero)))))),4) (6:11) [superposition 518,277]
[SA] active: 195. i(l11(s(s(nl8)))) = $sum($sum(alength,1),1) (3:11:1) [superposition 92,160]
[SA] passive: 930. v = a(l11(s(s(nl8))),$sum($sum(alength,1),1)) (4:12) [superposition 89,195]
[SA] passive: 931. i(l11(s(s(s(nl8))))) = $sum($sum($sum(alength,1),1),1) (4:14) [superposition 92,195]
[SA] passive: 932. i(l8(X0)) != $sum($sum(alength,1),1) | s(s(nl8)) = X0 (4:14) [superposition 143,195]
[SA] passive: 933. alength != $sum($sum(alength,1),1) | s(s(nl8)) = nl8 (4:12) [superposition 191,195]
[SA] passive: 948. alength != -2 (4:3) [term algebras distinctness 947]
[SA] passive: 935. i(l11(X1)) != $sum($sum(alength,1),1) | s(s(nl8)) = X1 (4:14) [superposition 210,195]
[SA] passive: 937. $sum($sum(alength,1),1) != sK0 | s(s(nl8)) = sK2(sK0) (4:13) [superposition 340,195]
[SA] passive: 939. v = a(l10(s(s(s(nl8)))),$sum($sum(alength,1),1)) (4:13) [superposition 518,195]
[SA] active: 948. alength != -2 (4:3:1) [term algebras distinctness 947]
[SA] active: 206. i(l8(X3)) != $sum(alength,1) | s(nl8) = X3 (3:11:1) [superposition 143,160]
[SA] active: 898. v = a(l11(s(s(s(zero)))),2) (5:9:1) [evaluation 897]
[SA] active: 377. i(l11(X4)) != $sum(alength,1) | s(nl8) = X4 (3:11:1) [superposition 210,160]
[SA] passive: 977. $sum(alength,1) != $sum($sum(alength,1),1) | s(nl8) = nl8 (4:13) [term algebras injectivity 962]
[SA] passive: 976. $sum(alength,1) != $sum(sK0,1) | nl8 = sK2(sK0) (4:11) [term algebras injectivity 964]
[SA] active: 297. ~'Sub'(s(s(s(s(s(s(X1)))))),X1) (6:9:1) [resolution 158,78]
[SA] passive: 994. ~'Sub'(s(s(s(s(s(s(s(X1))))))),X1) (7:10) [resolution 297,78]
[SA] passive: 996. ~'Sub'(s(s(s(s(s(X0))))),p(X0)) | zero = X0 (7:12) [superposition 297,52]
[SA] active: 510. v = a(main_end,i(l11(p(nl8)))) | zero = nl8 (3:11:1) [superposition 256,122]
[SA] active: 394. v = a(l11(sK2(0)),0) | 0 = sK0 (4:10:1) [superposition 89,172]
[SA] active: 248. $less(0,X8) | $less(X8,alength) | alength = X8 | 0 = sK0 (3:12:2) [resolution 130,148]
[SA] passive: 1004. $less(X1,alength) | alength = X1 | 0 = sK0 | X1 = X2 | $less(X2,X1) | $less(0,X2) (4:18) [resolution 248,130]
[SA] passive: 1005. $less(X3,alength) | alength = X3 | 0 = sK0 | ~$less(X4,0) | $less(X4,X3) (4:15) [resolution 248,18]
[SA] passive: 1013. $less(0,X0) | alength = X0 | 0 = sK0 | i(l11(sK2(X1))) = X1 | alength = X1 | $less(X1,0) | $less(X0,X1) (4:24) [resolution 248,261]
[SA] passive: 1017. $less(0,X4) | alength = X4 | 0 = sK0 | alength = X5 | $less(X5,alength) | $less(X4,X5) (4:18) [resolution 248,130]
[SA] passive: 1018. $less(0,X6) | alength = X6 | 0 = sK0 | ~$less(X7,X6) | $less(X7,alength) (4:15) [resolution 248,18]
[SA] active: 397. alength != 0 | nl8 = sK2(0) | 0 = sK0 (4:10:1) [superposition 191,172]
[SA] active: 287. ~$less(X3,sK0) | sK0 = X2 | $less(X2,alength) | $less(X3,X2) (3:12:1) [resolution 133,18]
[SA] passive: 1024. sK0 = X0 | $less(X0,alength) | $less(X1,X0) | $less(0,X1) | sK0 = X1 | 0 = sK0 (4:18) [resolution 287,250]
[SA] passive: 1027. sK0 = X6 | $less(X6,alength) | $less(X7,X6) | $less(X8,X7) | $less(sK0,X8) | sK0 = X7 | sK0 = X8 (4:21) [resolution 287,245]
[SA] passive: 1028. sK0 = X9 | $less(X9,alength) | $less(X10,X9) | $less(sK0,X11) | $less(X11,X10) | X10 = X11 | sK0 = X10 (4:21) [resolution 287,245]
[SA] passive: 1029. sK0 = X12 | $less(X12,alength) | $less(X13,X12) | $less(X14,X13) | $less(sK0,X14) | sK0 = X14 | X13 = X14 (4:21) [resolution 287,245]
[SA] passive: 1042. sK0 = X17 | $less(X17,alength) | $less(X18,X17) | $less(alength,X18) | 0 = sK0 | alength = X18 | i(l11(sK2(X18))) = X18 (4:24) [subsumption resolution 1031,55]
[SA] active: 402. i(l11(s(sK2(0)))) = 1 | 0 = sK0 (4:10:1) [evaluation 395]
[SA] passive: 1045. v = a(l11(s(sK2(0))),1) | 0 = sK0 (5:11) [superposition 89,402]
[SA] passive: 1061. i(l11(s(s(sK2(0))))) = 2 | 0 = sK0 (5:11) [evaluation 1046]
[SA] passive: 1047. i(l8(X0)) != 1 | s(sK2(0)) = X0 | 0 = sK0 (5:13) [superposition 143,402]
[SA] passive: 1048. alength != 1 | s(sK2(0)) = nl8 | 0 = sK0 (5:11) [superposition 191,402]
[SA] passive: 1050. i(l11(X1)) != 1 | s(sK2(0)) = X1 | 0 = sK0 (5:13) [superposition 210,402]
[SA] passive: 1052. 1 != sK0 | s(sK2(0)) = sK2(sK0) | 0 = sK0 (5:12) [superposition 340,402]
[SA] passive: 1055. v = a(l10(s(s(sK2(0)))),1) | 0 = sK0 (5:12) [superposition 518,402]
[SA] active: 294. ~$less(X4,X3) | sK0 = X3 | $less(sK0,X3) | $less(X4,alength) (3:12:1) [resolution 133,18]
[SA] passive: 1065. sK0 = X4 | $less(sK0,X4) | $less(X5,alength) | $less(X6,X5) | $less(X4,X6) | X4 = X5 | X4 = X6 (4:21) [resolution 294,245]
[SA] passive: 1066. sK0 = X7 | $less(sK0,X7) | $less(X8,alength) | $less(X7,X9) | $less(X9,X8) | X8 = X9 | X7 = X8 (4:21) [resolution 294,245]
[SA] passive: 1067. sK0 = X10 | $less(sK0,X10) | $less(X11,alength) | $less(X12,X11) | $less(X10,X12) | X10 = X12 | X11 = X12 (4:21) [resolution 294,245]
[SA] passive: 1097. sK0 = X27 | $less(sK0,X27) | $less(X28,0) | $less(X27,X28) | 0 = X28 | alength = X27 | i(l11(sK2(X27))) = X27 (4:24) [subsumption resolution 1078,17]
[SA] active: 443. $sum(alength,1) != sK0 | s(nl8) = sK2(sK0) (4:10:1) [superposition 340,160]
[SA] active: 300. $less(0,X1) | $less(X1,alength) | alength = 0 | 0 = X1 (3:12:2) [resolution 140,19]
[SA] passive: 1102. $less(X4,alength) | alength = 0 | 0 = X4 | ~$less(X5,0) | $less(X5,X4) (4:15) [resolution 300,18]
[SA] passive: 1111. $less(0,X0) | alength = 0 | 0 = X0 | i(l11(sK2(X1))) = X1 | alength = X1 | $less(X1,0) | $less(X0,X1) (4:24) [resolution 300,261]
[SA] passive: 1116. $less(0,X5) | alength = 0 | 0 = X5 | alength = X6 | $less(X6,alength) | $less(X5,X6) (4:18) [resolution 300,130]
[SA] passive: 1117. $less(0,X7) | alength = 0 | 0 = X7 | ~$less(X8,X7) | $less(X8,alength) (4:15) [resolution 300,18]
[SA] active: 205. i(l8(X2)) != 2 | s(s(zero)) = X2 (4:10:1) [superposition 143,165]
[SA] active: 307. $less(0,X1) | $less(X1,sK0) | 0 = sK0 | 0 = X1 (3:12:2) [resolution 149,19]
[SA] passive: 1132. $less(X4,sK0) | 0 = sK0 | 0 = X4 | ~$less(X5,0) | $less(X5,X4) (4:15) [resolution 307,18]
[SA] passive: 1141. $less(0,X0) | 0 = sK0 | 0 = X0 | sK0 = X1 | $less(X1,alength) | $less(X0,X1) (4:18) [resolution 307,287]
[SA] passive: 1146. $less(0,X5) | 0 = sK0 | 0 = X5 | sK0 = X6 | $less(X6,sK0) | $less(X5,X6) (4:18) [resolution 307,130]
[SA] passive: 1147. $less(0,X7) | 0 = sK0 | 0 = X7 | ~$less(X8,X7) | $less(X8,sK0) (4:15) [resolution 307,18]
[SA] active: 375. i(l11(X2)) != 2 | s(s(zero)) = X2 (4:10:1) [superposition 210,165]
[SA] active: 315. i(l8(X0)) != 0 | sK2(0) = X0 | alength = 0 (3:12:1) [superposition 143,171]
[SA] active: 632. v = a(l10(s(s(nl8))),$sum(alength,1)) (4:10:1) [superposition 518,160]
[SA] passive: 1188. $sum(alength,1) = $sum($sum(alength,1),1) | v = a(l11(s(s(nl8))),$sum(alength,1)) (5:19) [forward demodulation 1185,195]
[SA] active: 331. i(l8(X0)) != $sum(sK0,1) | s(sK2(sK0)) = X0 (3:12:1) [superposition 143,174]
[SA] active: 442. sK0 != 3 | s(s(s(zero))) = sK2(sK0) (5:10:1) [superposition 340,180]
[SA] active: 333. i(l11(s(s(sK2(sK0))))) = $sum($sum(sK0,1),1) (3:12:1) [superposition 92,174]
[SA] passive: 1196. v = a(l11(s(s(sK2(sK0)))),$sum($sum(sK0,1),1)) (4:13) [superposition 89,333]
[SA] passive: 1197. i(l11(s(s(s(sK2(sK0)))))) = $sum($sum($sum(sK0,1),1),1) (4:15) [superposition 92,333]
[SA] passive: 1198. i(l8(X0)) != $sum($sum(sK0,1),1) | s(s(sK2(sK0))) = X0 (4:15) [superposition 143,333]
[SA] passive: 1199. alength != $sum($sum(sK0,1),1) | s(s(sK2(sK0))) = nl8 (4:13) [superposition 191,333]
[SA] passive: 1222. sK0 != -2 (4:3) [term algebras distinctness 1221]
[SA] passive: 1201. i(l11(X1)) != $sum($sum(sK0,1),1) | s(s(sK2(sK0))) = X1 (4:15) [superposition 210,333]
[SA] passive: 1203. $sum($sum(sK0,1),1) != sK0 | s(s(sK2(sK0))) = sK2(sK0) (4:14) [superposition 340,333]
[SA] passive: 1208. $sum(alength,1) != $sum($sum(sK0,1),1) | s(sK2(sK0)) = nl8 (4:14) [term algebras injectivity 1206]
[SA] passive: 1207. v = a(l10(s(s(s(sK2(sK0))))),$sum($sum(sK0,1),1)) (4:14) [superposition 518,333]
[SA] active: 1222. sK0 != -2 (4:3:1) [term algebras distinctness 1221]
[SA] active: 345. ~$less(X3,sK0) | alength = X2 | $less(X2,alength) | $less(X3,X2) (3:12:1) [resolution 251,18]
[SA] passive: 1223. alength = X0 | $less(X0,alength) | $less(X1,X0) | $less(0,X1) | 0 = sK0 | 0 = X1 (4:18) [resolution 345,307]
[SA] passive: 1224. alength = X2 | $less(X2,alength) | $less(X3,X2) | $less(0,X3) | sK0 = X3 | 0 = sK0 (4:18) [resolution 345,250]
[SA] passive: 1227. alength = X8 | $less(X8,alength) | $less(X9,X8) | $less(X10,X9) | $less(sK0,X10) | sK0 = X9 | sK0 = X10 (4:21) [resolution 345,245]
[SA] passive: 1228. alength = X11 | $less(X11,alength) | $less(X12,X11) | $less(sK0,X13) | $less(X13,X12) | X12 = X13 | sK0 = X12 (4:21) [resolution 345,245]
[SA] passive: 1229. alength = X14 | $less(X14,alength) | $less(X15,X14) | $less(X16,X15) | $less(sK0,X16) | sK0 = X16 | X15 = X16 (4:21) [resolution 345,245]
[SA] active: 159. ~'Sub'(s(s(s(X0))),p(X0)) | zero = X0 (5:10:1) [superposition 110,52]
[SA] passive: 1245. ~'Sub'(s(s(X0)),p(p(X0))) | zero = p(X0) | zero = X0 (6:14) [superposition 159,52]
[SA] active: 354. ~$less(X4,X3) | alength = X3 | $less(sK0,X3) | $less(X4,alength) (3:12:1) [resolution 251,18]
[SA] passive: 1248. alength = X0 | $less(sK0,X0) | $less(X1,alength) | $less(X0,X1) | X0 = X1 (4:15) [resolution 354,19]
[SA] passive: 1250. alength = X4 | $less(sK0,X4) | $less(X5,alength) | $less(X6,X5) | $less(X4,X6) | X4 = X5 | X4 = X6 (4:21) [resolution 354,245]
[SA] passive: 1251. alength = X7 | $less(sK0,X7) | $less(X8,alength) | $less(X7,X9) | $less(X9,X8) | X8 = X9 | X7 = X8 (4:21) [resolution 354,245]
[SA] passive: 1252. alength = X10 | $less(sK0,X10) | $less(X11,alength) | $less(X12,X11) | $less(X10,X12) | X10 = X12 | X11 = X12 (4:21) [resolution 354,245]
[SA] passive: 1291. alength = X29 | $less(sK0,X29) | $less(X30,0) | $less(X29,X30) | 0 = X30 | i(l11(sK2(X29))) = X29 (4:21) [subsumption resolution 1289,17]
[SA] active: 505. 'Sub'(X0,s(s(s(p(X0))))) | zero = X0 (5:10:1) [resolution 483,114]
[SA] passive: 1292. zero = p(X0) | 'Sub'(X0,s(s(s(s(p(p(X0))))))) | zero = X0 (6:16) [resolution 505,114]
[SA] active: 378. i(l11(X5)) != $sum(sK0,1) | s(sK2(sK0)) = X5 (3:12:1) [superposition 210,174]
[SA] passive: 1323. $sum(sK0,1) != $sum($sum(alength,1),1) | s(nl8) = sK2(sK0) (4:14) [term algebras injectivity 1302]
[SA] passive: 1322. $sum(sK0,1) != $sum($sum(sK0,1),1) | s(sK2(sK0)) = sK2(sK0) (4:15) [term algebras injectivity 1303]
[SA] active: 631. v = a(l10(s(s(s(s(zero))))),3) (5:10:1) [superposition 518,180]
[SA] passive: 1344. v = a(l11(s(s(s(s(zero))))),3) (6:10) [evaluation 1343]
[SA] active: 381. i(l11(X8)) != 0 | sK2(0) = X8 | alength = 0 (3:12:1) [superposition 210,171]
[SA] active: 919. v = a(l11(s(s(s(s(zero))))),4) (6:10:1) [superposition 89,277]
[SA] active: 457. $less(0,X0) | $less(X0,alength) | 0 = sK0 | sK0 = X0 (3:12:2) [resolution 250,128]
[SA] passive: 1378. $less(X3,alength) | 0 = sK0 | sK0 = X3 | X3 = X4 | $less(X4,X3) | $less(0,X4) (4:18) [resolution 457,130]
[SA] passive: 1379. $less(X5,alength) | 0 = sK0 | sK0 = X5 | ~$less(X6,0) | $less(X6,X5) (4:15) [resolution 457,18]
[SA] passive: 1389. $less(0,X0) | 0 = sK0 | sK0 = X0 | i(l11(sK2(X1))) = X1 | alength = X1 | $less(X1,0) | $less(X0,X1) (4:24) [resolution 457,261]
[SA] passive: 1390. $less(0,X2) | 0 = sK0 | sK0 = X2 | $less(X2,0) | i(l11(sK2(X2))) = X2 (4:18) [resolution 457,91]
[SA] passive: 1396. $less(0,X8) | 0 = sK0 | sK0 = X8 | ~$less(X9,X8) | $less(X9,alength) (4:15) [resolution 457,18]
[SA] active: 922. alength != 4 | s(s(s(s(zero)))) = nl8 (6:10:1) [superposition 191,277]
[SA] active: 567. $less(0,X15) | $less(X15,alength) | 0 = X15 | alength = X15 (3:12:2) [resolution 245,53]
[SA] passive: 1408. $less(X3,alength) | 0 = X3 | alength = X3 | X3 = X4 | $less(X4,X3) | $less(0,X4) (4:18) [resolution 567,130]
[SA] passive: 1409. $less(X5,alength) | 0 = X5 | alength = X5 | ~$less(X6,0) | $less(X6,X5) (4:15) [resolution 567,18]
[SA] passive: 1419. $less(0,X0) | 0 = X0 | alength = X0 | i(l11(sK2(X1))) = X1 | alength = X1 | $less(X1,0) | $less(X0,X1) (4:24) [resolution 567,261]
[SA] passive: 1425. $less(0,X6) | 0 = X6 | alength = X6 | alength = X7 | $less(X7,alength) | $less(X6,X7) (4:18) [resolution 567,130]
[SA] passive: 1426. $less(0,X8) | 0 = X8 | alength = X8 | ~$less(X9,X8) | $less(X9,alength) (4:15) [resolution 567,18]
[SA] active: 929. i(l11(s(s(s(s(s(zero))))))) = 5 (6:10:1) [evaluation 920]
[SA] passive: 1432. v = a(l11(s(s(s(s(s(zero)))))),5) (7:11) [superposition 89,929]
[SA] passive: 1452. i(l11(s(s(s(s(s(s(zero)))))))) = 6 (7:11) [evaluation 1433]
[SA] passive: 1434. i(l8(X0)) != 5 | s(s(s(s(s(zero))))) = X0 (7:13) [superposition 143,929]
[SA] passive: 1435. alength != 5 | s(s(s(s(s(zero))))) = nl8 (7:11) [superposition 191,929]
[SA] passive: 1437. i(l11(X1)) != 5 | s(s(s(s(s(zero))))) = X1 (7:13) [superposition 210,929]
[SA] passive: 1439. sK0 != 5 | s(s(s(s(s(zero))))) = sK2(sK0) (7:12) [superposition 340,929]
[SA] passive: 1445. v = a(l10(s(s(s(s(s(s(zero))))))),5) (7:12) [superposition 518,929]
[SA] active: 568. $less(0,X16) | $less(X16,sK0) | 0 = X16 | sK0 = X16 (3:12:2) [resolution 245,55]
[SA] passive: 1480. $less(X2,sK0) | 0 = X2 | sK0 = X2 | i(l11(sK2(X2))) = X2 | alength = X2 | $less(alength,X2) (4:21) [evaluation 1456]
[SA] passive: 1457. $less(X3,sK0) | 0 = X3 | sK0 = X3 | X3 = X4 | $less(X4,X3) | $less(0,X4) (4:18) [resolution 568,130]
[SA] passive: 1458. $less(X5,sK0) | 0 = X5 | sK0 = X5 | ~$less(X6,0) | $less(X6,X5) (4:15) [resolution 568,18]
[SA] passive: 1461. $less(i(l11(sK2(X8))),sK0) | i(l11(sK2(X8))) = 0 | i(l11(sK2(X8))) = sK0 | $less(X8,0) | ~$less(X8,alength) | $less(0,alength) (4:27) [resolution 568,243]
[SA] passive: 1468. $less(0,X0) | 0 = X0 | sK0 = X0 | alength = X1 | $less(X1,alength) | $less(X0,X1) (4:18) [resolution 568,345]
[SA] passive: 1469. $less(0,X2) | 0 = X2 | sK0 = X2 | sK0 = X3 | $less(X3,alength) | $less(X2,X3) (4:18) [resolution 568,287]
[SA] passive: 1470. $less(0,X4) | 0 = X4 | sK0 = X4 | $less(X4,alength) (4:12) [resolution 568,128]
[SA] passive: 1475. $less(0,X8) | 0 = X8 | sK0 = X8 | sK0 = X9 | $less(X9,sK0) | $less(X8,X9) (4:18) [resolution 568,130]
[SA] passive: 1476. $less(0,X10) | 0 = X10 | sK0 = X10 | ~$less(X11,X10) | $less(X11,sK0) (4:15) [resolution 568,18]
[SA] active: 1344. v = a(l11(s(s(s(s(zero))))),3) (6:10:1) [evaluation 1343]
[SA] active: 359. ~$less(X1,i(l11(p(nl8)))) | zero = nl8 | $less(X1,alength) (3:12:1) [resolution 121,18]
[SA] passive: 1484. zero = nl8 | $less(X0,alength) | $less(i(l11(p(nl8))),X0) | i(l11(p(nl8))) = X0 (4:18) [resolution 359,19]
[SA] passive: 1486. zero = nl8 | $less(X2,alength) | $less(X3,X2) | $less(i(l11(p(nl8))),X3) | i(l11(p(nl8))) = X2 | i(l11(p(nl8))) = X3 (4:27) [resolution 359,245]
[SA] passive: 1487. zero = nl8 | $less(X4,alength) | $less(i(l11(p(nl8))),X5) | $less(X5,X4) | X4 = X5 | i(l11(p(nl8))) = X4 (4:24) [resolution 359,245]
[SA] passive: 1488. zero = nl8 | $less(X6,alength) | $less(X7,X6) | $less(i(l11(p(nl8))),X7) | i(l11(p(nl8))) = X7 | X6 = X7 (4:24) [resolution 359,245]
[SA] passive: 1505. zero = nl8 | $less(X10,0) | $less(i(l11(p(nl8))),X10) | 0 = X10 | alength = i(l11(p(nl8))) | i(l11(p(nl8))) = i(l11(sK2(i(l11(p(nl8)))))) (4:33) [subsumption resolution 1491,17]
[SA] passive: 1502. zero = nl8 | $less(0,alength) | $less(i(l11(p(nl8))),sK0) | i(l11(p(nl8))) = 0 | i(l11(p(nl8))) = sK0 (4:24) [resolution 359,568]
[SA] active: 994. ~'Sub'(s(s(s(s(s(s(s(X1))))))),X1) (7:10:1) [resolution 297,78]
[SA] passive: 1507. ~'Sub'(s(s(s(s(s(s(s(s(X1)))))))),X1) (8:11) [resolution 994,78]
[SA] passive: 1509. ~'Sub'(s(s(s(s(s(s(X0)))))),p(X0)) | zero = X0 (8:13) [superposition 994,52]
[SA] active: 474. 'Sub'(p(X0),s(X0)) | zero = p(X0) | zero = X0 (3:12:1) [duplicate literal removal 466]
[SA] active: 444. $sum(sK0,1) != sK0 | s(sK2(sK0)) = sK2(sK0) (4:11:1) [superposition 340,174]
[SA] active: 470. $less(i(l11(p(p(nl8)))),alength) | zero = p(nl8) | zero = nl8 (3:14:1) [resolution 197,93]
[SA] passive: 1518. zero = p(nl8) | zero = nl8 | i(l11(sK2(X0))) = X0 | alength = X0 | $less(X0,0) | $less(i(l11(p(p(nl8)))),X0) (4:26) [resolution 470,261]
[SA] passive: 1519. zero = p(nl8) | zero = nl8 | $less(i(l11(p(p(nl8)))),0) | i(l11(p(p(nl8)))) = i(l11(sK2(i(l11(p(p(nl8))))))) (4:28) [resolution 470,91]
[SA] passive: 1523. zero = p(nl8) | zero = nl8 | alength = X1 | $less(X1,alength) | $less(i(l11(p(p(nl8)))),X1) (4:20) [resolution 470,130]
[SA] passive: 1524. zero = p(nl8) | zero = nl8 | ~$less(X2,i(l11(p(p(nl8))))) | $less(X2,alength) (4:17) [resolution 470,18]
[SA] active: 448. 0 != sK0 | sK2(0) = sK2(sK0) | alength = 0 (4:11:1) [superposition 340,171]
[SA] active: 286. $less(sK0,X1) | $less(X1,X0) | $less(X0,alength) | X0 = X1 | sK0 = X0 (3:15:3) [resolution 133,130]
[SA] passive: 1529. $less(X7,X8) | $less(X8,alength) | X7 = X8 | sK0 = X8 | X7 = X9 | $less(X9,X7) | $less(sK0,X9) (4:21) [resolution 286,130]
[SA] passive: 1530. $less(X10,X11) | $less(X11,alength) | X10 = X11 | sK0 = X11 | ~$less(X12,sK0) | $less(X12,X10) (4:18) [resolution 286,18]
[SA] passive: 1546. $less(sK0,X7) | $less(X8,alength) | X7 = X8 | sK0 = X8 | X8 = X9 | $less(X9,X8) | $less(X7,X9) (4:21) [resolution 286,130]
[SA] passive: 1547. $less(sK0,X10) | $less(X11,alength) | X10 = X11 | sK0 = X11 | ~$less(X12,X10) | $less(X12,X11) (4:18) [resolution 286,18]
[SA] passive: 1560. $less(sK0,X0) | $less(X0,X1) | X0 = X1 | sK0 = X1 | i(l11(sK2(X2))) = X2 | alength = X2 | $less(X2,0) | $less(X1,X2) (4:27) [resolution 286,261]
[SA] passive: 1561. $less(sK0,X3) | $less(X3,X4) | X3 = X4 | sK0 = X4 | $less(X4,0) | i(l11(sK2(X4))) = X4 (4:21) [resolution 286,91]
[SA] passive: 1567. $less(sK0,X15) | $less(X15,X16) | X15 = X16 | sK0 = X16 | ~$less(X17,X16) | $less(X17,alength) (4:18) [resolution 286,18]
[SA] active: 633. v = a(l10(s(s(sK2(sK0)))),$sum(sK0,1)) (4:11:1) [superposition 518,174]
[SA] passive: 1584. $sum(sK0,1) = $sum($sum(sK0,1),1) | v = a(l11(s(s(sK2(sK0)))),$sum(sK0,1)) (5:20) [forward demodulation 1581,333]
[SA] active: 291. $less(sK0,X0) | $less(X0,0) | sK0 = X0 | i(l11(sK2(X0))) = X0 (3:15:2) [resolution 133,91]
[SA] passive: 1591. $less(X5,0) | sK0 = X5 | i(l11(sK2(X5))) = X5 | ~$less(X6,sK0) | $less(X6,X5) (4:18) [resolution 291,18]
[SA] passive: 1612. $less(sK0,X6) | sK0 = X6 | i(l11(sK2(X6))) = X6 | 0 = X7 | $less(X7,0) | $less(X6,X7) (4:21) [resolution 291,130]
[SA] passive: 1613. $less(sK0,X8) | sK0 = X8 | i(l11(sK2(X8))) = X8 | ~$less(X9,X8) | $less(X9,0) (4:18) [resolution 291,18]
[SA] active: 636. v = a(l10(s(sK2(0))),0) | 0 = sK0 (4:11:1) [superposition 518,172]
[SA] passive: 1625. v = a(l11(s(sK2(0))),0) | 0 = sK0 | i(l11(s(sK2(0)))) = 0 (5:18) [superposition 636,88]
[SA] active: 293. $less(sK0,X1) | $less(X2,alength) | $less(X1,X2) | alength = X2 | sK0 = X1 (3:15:3) [resolution 133,130]
[SA] passive: 1634. $less(X10,alength) | $less(X11,X10) | alength = X10 | sK0 = X11 | ~$less(X12,sK0) | $less(X12,X11) (4:18) [resolution 293,18]
[SA] passive: 1646. $less(sK0,X0) | $less(X0,X1) | alength = X1 | sK0 = X0 | i(l11(sK2(X2))) = X2 | alength = X2 | $less(X2,0) | $less(X1,X2) (4:27) [resolution 293,261]
[SA] passive: 1652. $less(sK0,X12) | $less(X12,X13) | alength = X13 | sK0 = X12 | alength = X14 | $less(X14,alength) | $less(X13,X14) (4:21) [resolution 293,130]
[SA] passive: 1653. $less(sK0,X15) | $less(X15,X16) | alength = X16 | sK0 = X15 | ~$less(X17,X16) | $less(X17,alength) (4:18) [resolution 293,18]
[SA] passive: 1658. $less(sK0,X7) | $less(X8,alength) | alength = X8 | sK0 = X7 | X8 = X9 | $less(X9,X8) | $less(X7,X9) (4:21) [resolution 293,130]
[SA] passive: 1659. $less(sK0,X10) | $less(X11,alength) | alength = X11 | sK0 = X10 | ~$less(X12,X10) | $less(X12,X11) (4:18) [resolution 293,18]
[SA] active: 637. v = a(l10(s(sK2(0))),0) | alength = 0 (4:11:1) [superposition 518,171]
[SA] passive: 1685. v = a(l11(s(sK2(0))),0) | alength = 0 | i(l11(s(sK2(0)))) = 0 (5:18) [superposition 637,88]
[SA] active: 344. $less(sK0,X1) | $less(X1,X0) | $less(X0,alength) | X0 = X1 | alength = X0 (3:15:3) [resolution 251,130]
[SA] passive: 1693. $less(X7,X8) | $less(X8,alength) | X7 = X8 | alength = X8 | X7 = X9 | $less(X9,X7) | $less(sK0,X9) (4:21) [resolution 344,130]
[SA] passive: 1694. $less(X10,X11) | $less(X11,alength) | X10 = X11 | alength = X11 | ~$less(X12,sK0) | $less(X12,X10) (4:18) [resolution 344,18]
[SA] passive: 1710. $less(sK0,X7) | $less(X8,alength) | X7 = X8 | alength = X8 | X8 = X9 | $less(X9,X8) | $less(X7,X9) (4:21) [resolution 344,130]
[SA] passive: 1711. $less(sK0,X10) | $less(X11,alength) | X10 = X11 | alength = X11 | ~$less(X12,X10) | $less(X12,X11) (4:18) [resolution 344,18]
[SA] passive: 1724. $less(sK0,X0) | $less(X0,X1) | X0 = X1 | alength = X1 | i(l11(sK2(X2))) = X2 | alength = X2 | $less(X2,0) | $less(X1,X2) (4:27) [resolution 344,261]
[SA] passive: 1725. $less(sK0,X3) | $less(X3,X4) | X3 = X4 | alength = X4 | $less(X4,0) | i(l11(sK2(X4))) = X4 (4:21) [resolution 344,91]
[SA] passive: 1730. $less(sK0,X12) | $less(X12,X13) | X12 = X13 | alength = X13 | alength = X14 | $less(X14,alength) | $less(X13,X14) (4:21) [resolution 344,130]
[SA] passive: 1731. $less(sK0,X15) | $less(X15,X16) | X15 = X16 | alength = X16 | ~$less(X17,X16) | $less(X17,alength) (4:18) [resolution 344,18]
[SA] active: 901. v = a(l11(s(sK2(0))),1) | alength = 0 (4:11:1) [superposition 89,319]
[SA] active: 351. $less(sK0,X0) | $less(X0,0) | alength = X0 | i(l11(sK2(X0))) = X0 (3:15:2) [resolution 251,91]
[SA] passive: 1753. $less(X5,0) | alength = X5 | i(l11(sK2(X5))) = X5 | ~$less(X6,sK0) | $less(X6,X5) (4:18) [resolution 351,18]
[SA] passive: 1768. $less(sK0,X1) | alength = X1 | i(l11(sK2(X1))) = X1 | 0 = sK0 | $less(X1,sK0) (4:18) [resolution 351,149]
[SA] passive: 1775. $less(sK0,X8) | alength = X8 | i(l11(sK2(X8))) = X8 | ~$less(X9,X8) | $less(X9,0) (4:18) [resolution 351,18]
[SA] active: 904. alength != 1 | s(sK2(0)) = nl8 | alength = 0 (4:11:1) [superposition 191,319]
[SA] active: 353. $less(sK0,X1) | $less(X2,alength) | $less(X1,X2) | alength = X2 | alength = X1 (3:15:3) [resolution 251,130]
[SA] passive: 1788. $less(X10,alength) | $less(X11,X10) | alength = X10 | alength = X11 | ~$less(X12,sK0) | $less(X12,X11) (4:18) [resolution 353,18]
[SA] passive: 1800. $less(sK0,X0) | $less(X0,X1) | alength = X1 | alength = X0 | i(l11(sK2(X2))) = X2 | alength = X2 | $less(X2,0) | $less(X1,X2) (4:27) [resolution 353,261]
[SA] passive: 1806. $less(sK0,X12) | $less(X12,X13) | alength = X13 | alength = X12 | alength = X14 | $less(X14,alength) | $less(X13,X14) (4:21) [resolution 353,130]
[SA] passive: 1807. $less(sK0,X15) | $less(X15,X16) | alength = X16 | alength = X15 | ~$less(X17,X16) | $less(X17,alength) (4:18) [resolution 353,18]
[SA] passive: 1812. $less(sK0,X7) | $less(X8,alength) | alength = X8 | alength = X7 | X8 = X9 | $less(X9,X8) | $less(X7,X9) (4:21) [resolution 353,130]
[SA] passive: 1813. $less(sK0,X10) | $less(X11,alength) | alength = X11 | alength = X10 | ~$less(X12,X10) | $less(X12,X11) (4:18) [resolution 353,18]
[SA] active: 913. i(l11(s(s(sK2(0))))) = 2 | alength = 0 (4:11:1) [evaluation 902]
[SA] passive: 1840. v = a(l11(s(s(sK2(0)))),2) | alength = 0 (5:12) [superposition 89,913]
[SA] passive: 1864. i(l11(s(s(s(sK2(0)))))) = 3 | alength = 0 (5:12) [evaluation 1841]
[SA] passive: 1842. i(l8(X0)) != 2 | s(s(sK2(0))) = X0 | alength = 0 (5:14) [superposition 143,913]
[SA] passive: 1843. alength != 2 | s(s(sK2(0))) = nl8 | alength = 0 (5:12) [superposition 191,913]
[SA] passive: 1845. i(l11(X1)) != 2 | s(s(sK2(0))) = X1 | alength = 0 (5:14) [superposition 210,913]
[SA] passive: 1847. sK0 != 2 | s(s(sK2(0))) = sK2(sK0) | alength = 0 (5:13) [superposition 340,913]
[SA] passive: 1853. v = a(l10(s(s(s(sK2(0))))),2) | alength = 0 (5:13) [superposition 518,913]
[SA] active: 425. ~$less(X3,0) | alength = X2 | alength = 0 | $less(X2,alength) | $less(X3,X2) (3:15:1) [resolution 249,18]
[SA] passive: 1868. alength = X4 | alength = 0 | $less(X4,alength) | $less(X5,X4) | $less(alength,X6) | $less(X6,X5) | 0 = X5 | alength = X6 | i(l11(sK2(X6))) = X6 (4:30) [resolution 425,268]
[SA] passive: 1873. alength = X16 | alength = 0 | $less(X16,alength) | $less(X17,X16) | $less(X18,X17) | $less(0,X18) | 0 = X17 | 0 = X18 (4:24) [resolution 425,245]
[SA] passive: 1874. alength = X19 | alength = 0 | $less(X19,alength) | $less(X20,X19) | $less(0,X21) | $less(X21,X20) | X20 = X21 | 0 = X20 (4:24) [resolution 425,245]
[SA] passive: 1875. alength = X22 | alength = 0 | $less(X22,alength) | $less(X23,X22) | $less(X24,X23) | $less(0,X24) | 0 = X24 | X23 = X24 (4:24) [resolution 425,245]
[SA] active: 976. $sum(alength,1) != $sum(sK0,1) | nl8 = sK2(sK0) (4:11:1) [term algebras injectivity 964]
[SA] active: 434. ~$less(X4,X3) | alength = X3 | alength = 0 | $less(0,X3) | $less(X4,alength) (3:15:1) [resolution 249,18]
[SA] passive: 1923. alength = X0 | alength = 0 | $less(0,X0) | $less(X1,alength) | $less(X0,X1) | X0 = X1 (4:18) [resolution 434,19]
[SA] passive: 1925. alength = X4 | alength = 0 | $less(0,X4) | $less(X5,alength) | $less(X6,X5) | $less(X4,X6) | X4 = X5 | X4 = X6 (4:24) [resolution 434,245]
[SA] passive: 1926. alength = X7 | alength = 0 | $less(0,X7) | $less(X8,alength) | $less(X7,X9) | $less(X9,X8) | X8 = X9 | X7 = X8 (4:24) [resolution 434,245]
[SA] passive: 1927. alength = X10 | alength = 0 | $less(0,X10) | $less(X11,alength) | $less(X12,X11) | $less(X10,X12) | X10 = X12 | X11 = X12 (4:24) [resolution 434,245]
[SA] passive: 2011. alength = X50 | alength = 0 | $less(0,X50) | $less(X51,0) | $less(X50,X51) | 0 = X51 | i(l11(sK2(X50))) = X50 (4:24) [subsumption resolution 1994,17]
[SA] active: 1045. v = a(l11(s(sK2(0))),1) | 0 = sK0 (5:11:1) [superposition 89,402]
[SA] active: 451. ~$less(X3,0) | sK0 = X2 | 0 = sK0 | $less(X2,sK0) | $less(X3,X2) (3:15:1) [resolution 250,18]
[SA] passive: 2017. sK0 = X0 | 0 = sK0 | $less(X0,sK0) | $less(X1,X0) | $less(sK0,X1) | alength = X1 | i(l11(sK2(X1))) = X1 (4:24) [resolution 451,351]
[SA] passive: 2019. sK0 = X4 | 0 = sK0 | $less(X4,sK0) | $less(X5,X4) | $less(alength,X6) | $less(X6,X5) | 0 = X5 | alength = X6 | i(l11(sK2(X6))) = X6 (4:30) [resolution 451,268]
[SA] passive: 2020. sK0 = X7 | 0 = sK0 | $less(X7,sK0) | $less(X8,X7) | $less(alength,X9) | $less(X9,X8) | X8 = X9 | alength = X8 | i(l11(sK2(X8))) = X8 (4:30) [resolution 451,260]
[SA] passive: 2021. sK0 = X10 | 0 = sK0 | $less(X10,sK0) | $less(X11,X10) | $less(alength,X11) | i(l11(sK2(X11))) = X11 | alength = X11 (4:24) [resolution 451,167]
[SA] passive: 2024. sK0 = X16 | 0 = sK0 | $less(X16,sK0) | $less(X17,X16) | $less(X18,X17) | $less(0,X18) | 0 = X17 | 0 = X18 (4:24) [resolution 451,245]
[SA] passive: 2025. sK0 = X19 | 0 = sK0 | $less(X19,sK0) | $less(X20,X19) | $less(0,X21) | $less(X21,X20) | X20 = X21 | 0 = X20 (4:24) [resolution 451,245]
[SA] passive: 2026. sK0 = X22 | 0 = sK0 | $less(X22,sK0) | $less(X23,X22) | $less(X24,X23) | $less(0,X24) | 0 = X24 | X23 = X24 (4:24) [resolution 451,245]
[SA] passive: 2036. sK0 = X42 | 0 = sK0 | $less(X42,sK0) | $less(i(l11(sK2(X43))),X42) | $less(X43,0) | ~$less(X43,alength) | i(l11(sK2(X43))) = i(l11(sK2(i(l11(sK2(X43)))))) (4:33) [resolution 451,242]
[SA] active: 1048. alength != 1 | s(sK2(0)) = nl8 | 0 = sK0 (5:11:1) [superposition 191,402]
[SA] active: 460. ~$less(X4,X3) | sK0 = X3 | 0 = sK0 | $less(0,X3) | $less(X4,sK0) (3:15:1) [resolution 250,18]
[SA] passive: 2068. sK0 = X0 | 0 = sK0 | $less(0,X0) | $less(X1,sK0) | $less(X0,X1) | X0 = X1 (4:18) [resolution 460,19]
[SA] passive: 2070. sK0 = X4 | 0 = sK0 | $less(0,X4) | $less(X5,sK0) | $less(X6,X5) | $less(X4,X6) | X4 = X5 | X4 = X6 (4:24) [resolution 460,245]
[SA] passive: 2071. sK0 = X7 | 0 = sK0 | $less(0,X7) | $less(X8,sK0) | $less(X7,X9) | $less(X9,X8) | X8 = X9 | X7 = X8 (4:24) [resolution 460,245]
[SA] passive: 2072. sK0 = X10 | 0 = sK0 | $less(0,X10) | $less(X11,sK0) | $less(X12,X11) | $less(X10,X12) | X10 = X12 | X11 = X12 (4:24) [resolution 460,245]
[SA] passive: 2098. sK0 = X50 | 0 = sK0 | $less(0,X50) | $less(alength,sK0) | $less(X51,0) | $less(X50,X51) | 0 = X51 | alength = X50 | i(l11(sK2(X50))) = X50 (4:30) [resolution 460,268]
[SA] active: 1061. i(l11(s(s(sK2(0))))) = 2 | 0 = sK0 (5:11:1) [evaluation 1046]
[SA] passive: 2155. v = a(l11(s(s(sK2(0)))),2) | 0 = sK0 (6:12) [superposition 89,1061]
[SA] passive: 2178. i(l11(s(s(s(sK2(0)))))) = 3 | 0 = sK0 (6:12) [evaluation 2156]
[SA] passive: 2157. i(l8(X0)) != 2 | s(s(sK2(0))) = X0 | 0 = sK0 (6:14) [superposition 143,1061]
[SA] passive: 2158. alength != 2 | s(s(sK2(0))) = nl8 | 0 = sK0 (6:12) [superposition 191,1061]
[SA] passive: 2160. i(l11(X1)) != 2 | s(s(sK2(0))) = X1 | 0 = sK0 (6:14) [superposition 210,1061]
[SA] passive: 2162. sK0 != 2 | s(s(sK2(0))) = sK2(sK0) | 0 = sK0 (6:13) [superposition 340,1061]
[SA] passive: 2168. v = a(l10(s(s(s(sK2(0))))),2) | 0 = sK0 (6:13) [superposition 518,1061]
[SA] active: 572. $less(sK0,X23) | $less(X24,alength) | $less(X23,X24) | X23 = X24 | sK0 = X23 (3:15:3) [resolution 245,128]
[SA] passive: 2181. $less(X1,alength) | $less(X2,X1) | X1 = X2 | sK0 = X2 | ~$less(X3,sK0) | $less(X3,X2) (4:18) [resolution 572,18]
[SA] passive: 2182. $less(X4,alength) | $less(X5,X4) | X4 = X5 | sK0 = X5 | X5 = X6 | $less(X6,X5) | $less(sK0,X6) (4:21) [resolution 572,130]
[SA] passive: 2246. $less(X15,alength) | $less(X16,X15) | X15 = X16 | sK0 = X16 | 0 = sK0 | $less(0,X16) (4:18) [subsumption resolution 2244,17]
[SA] passive: 2201. $less(sK0,X0) | $less(X0,X1) | X0 = X1 | sK0 = X0 | i(l11(sK2(X2))) = X2 | alength = X2 | $less(X2,0) | $less(X1,X2) (4:27) [resolution 572,261]
[SA] passive: 2202. $less(sK0,X3) | $less(X3,X4) | X3 = X4 | sK0 = X3 | $less(X4,0) | i(l11(sK2(X4))) = X4 (4:21) [resolution 572,91]
[SA] passive: 2204. $less(sK0,X6) | $less(X6,X7) | X6 = X7 | sK0 = X6 | ~$less(X8,X7) | $less(X8,alength) (4:18) [resolution 572,18]
[SA] passive: 2212. $less(sK0,X1) | $less(X2,alength) | X1 = X2 | sK0 = X1 | ~$less(X3,X1) | $less(X3,X2) (4:18) [resolution 572,18]
[SA] active: 273. i(l8(X0)) != 3 | s(s(s(zero))) = X0 (5:11:1) [superposition 143,180]
[SA] active: 594. $less(sK0,X23) | $less(X24,alength) | $less(X23,X24) | sK0 = X23 | sK0 = X24 (3:15:3) [resolution 245,128]
[SA] passive: 2251. $less(X1,alength) | $less(X2,X1) | sK0 = X2 | sK0 = X1 | ~$less(X3,sK0) | $less(X3,X2) (4:18) [resolution 594,18]
[SA] passive: 2319. $less(X7,alength) | $less(X8,X7) | sK0 = X8 | sK0 = X7 | i(l11(sK2(X8))) = X8 | alength = X8 | $less(alength,X8) (4:24) [subsumption resolution 2253,55]
[SA] passive: 2271. $less(sK0,X0) | $less(X0,X1) | sK0 = X0 | sK0 = X1 | i(l11(sK2(X2))) = X2 | alength = X2 | $less(X2,0) | $less(X1,X2) (4:27) [resolution 594,261]
[SA] passive: 2272. $less(sK0,X3) | $less(X3,X4) | sK0 = X3 | sK0 = X4 | $less(X4,0) | i(l11(sK2(X4))) = X4 (4:21) [resolution 594,91]
[SA] passive: 2274. $less(sK0,X6) | $less(X6,X7) | sK0 = X6 | sK0 = X7 | ~$less(X8,X7) | $less(X8,alength) (4:18) [resolution 594,18]
[SA] passive: 2282. $less(sK0,X1) | $less(X2,alength) | sK0 = X1 | sK0 = X2 | ~$less(X3,X1) | $less(X3,X2) (4:18) [resolution 594,18]
[SA] passive: 2283. $less(sK0,X4) | $less(X5,alength) | sK0 = X4 | sK0 = X5 | X5 = X6 | $less(X6,X5) | $less(X4,X6) (4:21) [resolution 594,130]
[SA] active: 376. i(l11(X3)) != 3 | s(s(s(zero))) = X3 (5:11:1) [superposition 210,180]
[SA] active: 358. $less(i(l11(p(nl8))),X0) | alength = X0 | $less(X0,alength) | zero = nl8 (3:15:1) [resolution 121,130]
[SA] passive: 2355. alength = X0 | $less(X0,alength) | zero = nl8 | ~$less(X1,i(l11(p(nl8)))) | $less(X1,X0) (4:18) [resolution 358,18]
[SA] passive: 2356. alength = X2 | $less(X2,alength) | zero = nl8 | X2 = X3 | $less(X3,X2) | $less(i(l11(p(nl8))),X3) (4:21) [resolution 358,130]
[SA] active: 926. sK0 != 4 | s(s(s(s(zero)))) = sK2(sK0) (6:11:1) [superposition 340,277]
[SA] active: 360. 'Sub'(X0,s(s(s(p(p(X0)))))) | zero = p(X0) | zero = X0 (3:15:1) [resolution 215,114]
[SA] passive: 2380. zero = p(p(X0)) | zero = p(X0) | 'Sub'(X0,s(s(s(s(p(p(p(X0)))))))) | zero = X0 (4:22) [resolution 360,114]
[SA] passive: 2385. zero = X0 | 'Sub'(s(X0),s(s(s(p(X0))))) (4:11) [forward demodulation 2384,231]
[SA] active: 2385. 'Sub'(s(X0),s(s(s(p(X0))))) | zero = X0 (4:11:1) [forward demodulation 2384,231]
[SA] active: 492. ~'Sub'(p(X1),p(X0)) | zero = X1 | 'Sub'(X1,s(X0)) | zero = X0 (3:15:1) [superposition 216,52]
[SA] passive: 2397. zero = p(p(X1)) | 'Sub'(p(p(X1)),s(X1)) | zero = X1 | zero = p(X1) (4:18) [duplicate literal removal 2392]
[SA] passive: 2396. ~'Sub'(X0,p(X1)) | 'Sub'(s(X0),s(X1)) | zero = X1 (4:12) [term algebras distinctness 2393]
[SA] active: 298. ~'Sub'(s(s(s(s(X0)))),p(X0)) | zero = X0 (6:11:1) [superposition 158,52]
[SA] passive: 2399. ~'Sub'(s(s(s(X0))),p(p(X0))) | zero = p(X0) | zero = X0 (7:15) [superposition 298,52]
[SA] active: 526. $less(i(l11(s(p(p(nl8))))),alength) | zero = p(nl8) | zero = nl8 (3:15:1) [resolution 228,93]
[SA] passive: 2402. zero = p(nl8) | zero = nl8 | i(l11(sK2(X0))) = X0 | alength = X0 | $less(X0,0) | $less(i(l11(s(p(p(nl8))))),X0) (4:27) [resolution 526,261]
[SA] passive: 2403. zero = p(nl8) | zero = nl8 | $less(i(l11(s(p(p(nl8))))),0) | i(l11(s(p(p(nl8))))) = i(l11(sK2(i(l11(s(p(p(nl8)))))))) (4:31) [resolution 526,91]
[SA] passive: 2404. zero = p(nl8) | zero = nl8 | ~$less(X1,i(l11(s(p(p(nl8)))))) | $less(X1,alength) (4:18) [resolution 526,18]
[SA] passive: 2405. zero = p(nl8) | zero = nl8 | alength = X2 | $less(X2,alength) | $less(i(l11(s(p(p(nl8))))),X2) (4:21) [resolution 526,130]
[SA] active: 928. v = a(l10(s(s(s(s(s(zero)))))),4) (6:11:1) [superposition 518,277]
[SA] passive: 2418. v = a(l11(s(s(s(s(s(zero)))))),4) (7:11) [evaluation 2417]
[SA] active: 306. $less(alength,X0) | $less(X0,sK0) | 0 = sK0 | i(l11(sK2(X0))) = X0 | alength = X0 (3:18:2) [resolution 149,167]
[SA] passive: 2420. $less(X0,sK0) | 0 = sK0 | i(l11(sK2(X0))) = X0 | alength = X0 | ~$less(X1,alength) | $less(X1,X0) (4:21) [resolution 306,18]
[SA] passive: 2462. $less(X7,sK0) | 0 = sK0 | i(l11(sK2(X7))) = X7 | alength = X7 | alength = 0 | $less(0,X7) (4:21) [subsumption resolution 2459,17]
[SA] passive: 2463. $less(i(l11(p(nl8))),sK0) | 0 = sK0 | i(l11(p(nl8))) = i(l11(sK2(i(l11(p(nl8)))))) | alength = i(l11(p(nl8))) | zero = nl8 (4:30) [subsumption resolution 2430,17]
[SA] passive: 2444. $less(alength,X5) | 0 = sK0 | i(l11(sK2(X5))) = X5 | alength = X5 | ~$less(X6,X5) | $less(X6,sK0) (4:21) [resolution 306,18]
[SA] active: 1432. v = a(l11(s(s(s(s(s(zero)))))),5) (7:11:1) [superposition 89,929]
[SA] active: 424. $less(0,X1) | $less(X1,X0) | $less(X0,alength) | X0 = X1 | alength = 0 | alength = X0 (3:18:3) [resolution 249,130]
[SA] passive: 2467. $less(X1,X2) | $less(X2,alength) | X1 = X2 | alength = 0 | alength = X2 | ~$less(X3,0) | $less(X3,X1) (4:21) [resolution 424,18]
[SA] passive: 2468. $less(X4,X5) | $less(X5,alength) | X4 = X5 | alength = 0 | alength = X5 | X4 = X6 | $less(X6,X4) | $less(0,X6) (4:24) [resolution 424,130]
[SA] passive: 2487. $less(0,X1) | $less(X2,alength) | X1 = X2 | alength = 0 | alength = X2 | ~$less(X3,X1) | $less(X3,X2) (4:21) [resolution 424,18]
[SA] passive: 2488. $less(0,X4) | $less(X5,alength) | X4 = X5 | alength = 0 | alength = X5 | X5 = X6 | $less(X6,X5) | $less(X4,X6) (4:24) [resolution 424,130]
[SA] passive: 2508. $less(0,X0) | $less(X0,X1) | X0 = X1 | alength = 0 | alength = X1 | i(l11(sK2(X2))) = X2 | alength = X2 | $less(X2,0) | $less(X1,X2) (4:30) [resolution 424,261]
[SA] passive: 2509. $less(0,X3) | $less(X3,X4) | X3 = X4 | alength = 0 | alength = X4 | $less(X4,0) | i(l11(sK2(X4))) = X4 (4:24) [resolution 424,91]
[SA] passive: 2511. $less(0,X6) | $less(X6,X7) | X6 = X7 | alength = 0 | alength = X7 | ~$less(X8,X7) | $less(X8,alength) (4:21) [resolution 424,18]
[SA] passive: 2512. $less(0,X9) | $less(X9,X10) | X9 = X10 | alength = 0 | alength = X10 | alength = X11 | $less(X11,alength) | $less(X10,X11) (4:24) [resolution 424,130]
[SA] active: 1435. alength != 5 | s(s(s(s(s(zero))))) = nl8 (7:11:1) [superposition 191,929]
[SA] active: 431. $less(0,X0) | $less(X0,0) | alength = 0 | alength = X0 | i(l11(sK2(X0))) = X0 (3:18:2) [resolution 249,91]
[SA] passive: 2543. $less(X0,0) | alength = 0 | alength = X0 | i(l11(sK2(X0))) = X0 | ~$less(X1,0) | $less(X1,X0) (4:21) [resolution 431,18]
[SA] passive: 2564. $less(0,X0) | alength = 0 | alength = X0 | i(l11(sK2(X0))) = X0 | sK0 = X1 | 0 = sK0 | $less(X1,sK0) | $less(X0,X1) (4:27) [resolution 431,451]
[SA] passive: 2570. $less(0,X7) | alength = 0 | alength = X7 | i(l11(sK2(X7))) = X7 | ~$less(X8,X7) | $less(X8,0) (4:21) [resolution 431,18]
[SA] active: 1452. i(l11(s(s(s(s(s(s(zero)))))))) = 6 (7:11:1) [evaluation 1433]
[SA] passive: 2606. v = a(l11(s(s(s(s(s(s(zero))))))),6) (8:12) [superposition 89,1452]
[SA] passive: 2627. i(l11(s(s(s(s(s(s(s(zero))))))))) = 7 (8:12) [evaluation 2607]
[SA] passive: 2608. i(l8(X0)) != 6 | s(s(s(s(s(s(zero)))))) = X0 (8:14) [superposition 143,1452]
[SA] passive: 2609. alength != 6 | s(s(s(s(s(s(zero)))))) = nl8 (8:12) [superposition 191,1452]
[SA] passive: 2611. i(l11(X1)) != 6 | s(s(s(s(s(s(zero)))))) = X1 (8:14) [superposition 210,1452]
[SA] passive: 2613. sK0 != 6 | s(s(s(s(s(s(zero)))))) = sK2(sK0) (8:13) [superposition 340,1452]
[SA] passive: 2620. v = a(l10(s(s(s(s(s(s(s(zero)))))))),6) (8:13) [superposition 518,1452]
[SA] active: 433. $less(0,X1) | $less(X2,alength) | $less(X1,X2) | alength = X2 | alength = 0 | alength = X1 (3:18:3) [resolution 249,130]
[SA] passive: 2629. $less(X1,alength) | $less(X2,X1) | alength = X1 | alength = 0 | alength = X2 | ~$less(X3,0) | $less(X3,X2) (4:21) [resolution 433,18]
[SA] passive: 2648. $less(0,X0) | $less(X0,X1) | alength = X1 | alength = 0 | alength = X0 | i(l11(sK2(X2))) = X2 | alength = X2 | $less(X2,0) | $less(X1,X2) (4:30) [resolution 433,261]
[SA] passive: 2651. $less(0,X6) | $less(X6,X7) | alength = X7 | alength = 0 | alength = X6 | ~$less(X8,X7) | $less(X8,alength) (4:21) [resolution 433,18]
[SA] passive: 2652. $less(0,X9) | $less(X9,X10) | alength = X10 | alength = 0 | alength = X9 | alength = X11 | $less(X11,alength) | $less(X10,X11) (4:24) [resolution 433,130]
[SA] passive: 2659. $less(0,X1) | $less(X2,alength) | alength = X2 | alength = 0 | alength = X1 | ~$less(X3,X1) | $less(X3,X2) (4:21) [resolution 433,18]
[SA] passive: 2660. $less(0,X4) | $less(X5,alength) | alength = X5 | alength = 0 | alength = X4 | X5 = X6 | $less(X6,X5) | $less(X4,X6) (4:24) [resolution 433,130]
[SA] active: 2418. v = a(l11(s(s(s(s(s(zero)))))),4) (7:11:1) [evaluation 2417]
[SA] active: 450. $less(0,X1) | $less(X1,X0) | $less(X0,sK0) | X0 = X1 | 0 = sK0 | sK0 = X0 (3:18:3) [resolution 250,130]
[SA] passive: 2712. $less(X1,X2) | $less(X2,sK0) | X1 = X2 | 0 = sK0 | sK0 = X2 | ~$less(X3,0) | $less(X3,X1) (4:21) [resolution 450,18]
[SA] passive: 2713. $less(X4,X5) | $less(X5,sK0) | X4 = X5 | 0 = sK0 | sK0 = X5 | X4 = X6 | $less(X6,X4) | $less(0,X6) (4:24) [resolution 450,130]
[SA] passive: 2732. $less(0,X1) | $less(X2,sK0) | X1 = X2 | 0 = sK0 | sK0 = X2 | ~$less(X3,X1) | $less(X3,X2) (4:21) [resolution 450,18]
[SA] passive: 2733. $less(0,X4) | $less(X5,sK0) | X4 = X5 | 0 = sK0 | sK0 = X5 | X5 = X6 | $less(X6,X5) | $less(X4,X6) (4:24) [resolution 450,130]
[SA] passive: 2741. $less(0,X20) | $less(i(l11(sK2(X21))),sK0) | i(l11(sK2(X21))) = X20 | 0 = sK0 | i(l11(sK2(X21))) = sK0 | $less(X21,0) | ~$less(X21,alength) | $less(X20,alength) (4:33) [resolution 450,243]
[SA] passive: 2742. $less(0,X22) | $less(i(l11(p(nl8))),sK0) | i(l11(p(nl8))) = X22 | 0 = sK0 | i(l11(p(nl8))) = sK0 | zero = nl8 | $less(X22,alength) (4:30) [resolution 450,359]
[SA] passive: 2753. $less(0,X0) | $less(X0,X1) | X0 = X1 | 0 = sK0 | sK0 = X1 | alength = X2 | $less(X2,alength) | $less(X1,X2) (4:24) [resolution 450,345]
[SA] passive: 2754. $less(0,X3) | $less(X3,X4) | X3 = X4 | 0 = sK0 | sK0 = X4 | sK0 = X5 | $less(X5,alength) | $less(X4,X5) (4:24) [resolution 450,287]
[SA] passive: 2757. $less(0,X9) | $less(X9,X10) | X9 = X10 | 0 = sK0 | sK0 = X10 | ~$less(X11,X10) | $less(X11,sK0) (4:21) [resolution 450,18]
[SA] passive: 2758. $less(0,X12) | $less(X12,X13) | X12 = X13 | 0 = sK0 | sK0 = X13 | sK0 = X14 | $less(X14,sK0) | $less(X13,X14) (4:24) [resolution 450,130]
[SA] active: 1507. ~'Sub'(s(s(s(s(s(s(s(s(X1)))))))),X1) (8:11:1) [resolution 994,78]
[SA] passive: 2790. ~'Sub'(s(s(s(s(s(s(s(s(s(X1))))))))),X1) (9:12) [resolution 1507,78]
[SA] passive: 2792. ~'Sub'(s(s(s(s(s(s(s(X0))))))),p(X0)) | zero = X0 (9:14) [superposition 1507,52]
[SA] active: 459. $less(0,X1) | $less(X2,sK0) | $less(X1,X2) | sK0 = X2 | 0 = sK0 | sK0 = X1 (3:18:3) [resolution 250,130]
[SA] passive: 2796. $less(X1,sK0) | $less(X2,X1) | sK0 = X1 | 0 = sK0 | sK0 = X2 | ~$less(X3,0) | $less(X3,X2) (4:21) [resolution 459,18]
[SA] passive: 2815. $less(0,X0) | $less(X0,X1) | sK0 = X1 | 0 = sK0 | sK0 = X0 | alength = X2 | $less(X2,alength) | $less(X1,X2) (4:24) [resolution 459,345]
[SA] passive: 2816. $less(0,X3) | $less(X3,X4) | sK0 = X4 | 0 = sK0 | sK0 = X3 | sK0 = X5 | $less(X5,alength) | $less(X4,X5) (4:24) [resolution 459,287]
[SA] passive: 2819. $less(0,X9) | $less(X9,X10) | sK0 = X10 | 0 = sK0 | sK0 = X9 | ~$less(X11,X10) | $less(X11,sK0) (4:21) [resolution 459,18]
[SA] passive: 2820. $less(0,X12) | $less(X12,X13) | sK0 = X13 | 0 = sK0 | sK0 = X12 | sK0 = X14 | $less(X14,sK0) | $less(X13,X14) (4:24) [resolution 459,130]
[SA] passive: 2827. $less(0,X1) | $less(X2,sK0) | sK0 = X2 | 0 = sK0 | sK0 = X1 | ~$less(X3,X1) | $less(X3,X2) (4:21) [resolution 459,18]
[SA] passive: 2828. $less(0,X4) | $less(X5,sK0) | sK0 = X5 | 0 = sK0 | sK0 = X4 | X5 = X6 | $less(X6,X5) | $less(X4,X6) (4:24) [resolution 459,130]
[SA] active: 366. $less(0,X1) | $less(X1,alength) | 0 = sK0 | 0 = X1 (4:12:2) [resolution 151,19]
[SA] passive: 2877. $less(X0,alength) | 0 = sK0 | 0 = X0 | ~$less(X1,0) | $less(X1,X0) (5:15) [resolution 366,18]
[SA] passive: 2896. $less(0,X0) | 0 = sK0 | 0 = X0 | i(l11(sK2(X1))) = X1 | alength = X1 | $less(X1,0) | $less(X0,X1) (5:24) [resolution 366,261]
[SA] passive: 2899. $less(0,X3) | 0 = sK0 | 0 = X3 | ~$less(X4,X3) | $less(X4,alength) (5:15) [resolution 366,18]
[SA] active: 570. $less(0,X19) | $less(X20,sK0) | $less(X19,X20) | X19 = X20 | 0 = sK0 | 0 = X19 (3:18:3) [resolution 245,149]
[SA] passive: 2918. $less(X1,sK0) | $less(X2,X1) | X1 = X2 | 0 = sK0 | 0 = X2 | ~$less(X3,0) | $less(X3,X2) (4:21) [resolution 570,18]
[SA] passive: 2919. $less(X4,sK0) | $less(X5,X4) | X4 = X5 | 0 = sK0 | 0 = X5 | X5 = X6 | $less(X6,X5) | $less(0,X6) (4:24) [resolution 570,130]
[SA] passive: 2993. $less(X7,sK0) | $less(X8,X7) | X7 = X8 | 0 = sK0 | 0 = X8 | i(l11(sK2(X8))) = X8 | alength = X8 | $less(alength,X8) (4:27) [evaluation 2920]
[SA] passive: 2937. $less(0,X0) | $less(X0,X1) | X0 = X1 | 0 = sK0 | 0 = X0 | alength = X2 | $less(X2,alength) | $less(X1,X2) (4:24) [resolution 570,345]
[SA] passive: 2938. $less(0,X3) | $less(X3,X4) | X3 = X4 | 0 = sK0 | 0 = X3 | sK0 = X5 | $less(X5,alength) | $less(X4,X5) (4:24) [resolution 570,287]
[SA] passive: 2941. $less(0,X9) | $less(X9,X10) | X9 = X10 | 0 = sK0 | 0 = X9 | ~$less(X11,X10) | $less(X11,sK0) (4:21) [resolution 570,18]
[SA] passive: 2949. $less(0,X1) | $less(X2,sK0) | X1 = X2 | 0 = sK0 | 0 = X1 | ~$less(X3,X1) | $less(X3,X2) (4:21) [resolution 570,18]
[SA] passive: 2950. $less(0,X4) | $less(X5,sK0) | X4 = X5 | 0 = sK0 | 0 = X4 | X5 = X6 | $less(X6,X5) | $less(X4,X6) (4:24) [resolution 570,130]
[SA] active: 396. i(l8(X0)) != 0 | sK2(0) = X0 | 0 = sK0 (4:12:1) [superposition 143,172]
[SA] active: 571. $less(0,X21) | $less(X22,alength) | $less(X21,X22) | X21 = X22 | alength = 0 | 0 = X21 (3:18:3) [resolution 245,140]
[SA] passive: 2999. $less(X1,alength) | $less(X2,X1) | X1 = X2 | alength = 0 | 0 = X2 | ~$less(X3,0) | $less(X3,X2) (4:21) [resolution 571,18]
[SA] passive: 3000. $less(X4,alength) | $less(X5,X4) | X4 = X5 | alength = 0 | 0 = X5 | X5 = X6 | $less(X6,X5) | $less(0,X6) (4:24) [resolution 571,130]
[SA] passive: 3018. $less(0,X0) | $less(X0,X1) | X0 = X1 | alength = 0 | 0 = X0 | i(l11(sK2(X2))) = X2 | alength = X2 | $less(X2,0) | $less(X1,X2) (4:30) [resolution 571,261]
[SA] passive: 3021. $less(0,X6) | $less(X6,X7) | X6 = X7 | alength = 0 | 0 = X6 | ~$less(X8,X7) | $less(X8,alength) (4:21) [resolution 571,18]
[SA] passive: 3029. $less(0,X1) | $less(X2,alength) | X1 = X2 | alength = 0 | 0 = X1 | ~$less(X3,X1) | $less(X3,X2) (4:21) [resolution 571,18]
[SA] passive: 3030. $less(0,X4) | $less(X5,alength) | X4 = X5 | alength = 0 | 0 = X4 | X5 = X6 | $less(X6,X5) | $less(X4,X6) (4:24) [resolution 571,130]
[SA] active: 399. i(l11(X1)) != 0 | sK2(0) = X1 | 0 = sK0 (4:12:1) [superposition 210,172]
[SA] active: 581. $less(0,X19) | $less(X20,sK0) | $less(X19,X20) | 0 = X20 | 0 = sK0 | X19 = X20 (3:18:3) [resolution 245,149]
[SA] passive: 3105. $less(X1,sK0) | $less(X2,X1) | 0 = X1 | 0 = sK0 | X1 = X2 | ~$less(X3,0) | $less(X3,X2) (4:21) [resolution 581,18]
[SA] passive: 3106. $less(X4,sK0) | $less(X5,X4) | 0 = X4 | 0 = sK0 | X4 = X5 | X5 = X6 | $less(X6,X5) | $less(0,X6) (4:24) [resolution 581,130]
[SA] passive: 3124. $less(0,X0) | $less(X0,X1) | 0 = X1 | 0 = sK0 | X0 = X1 | alength = X2 | $less(X2,alength) | $less(X1,X2) (4:24) [resolution 581,345]
[SA] passive: 3125. $less(0,X3) | $less(X3,X4) | 0 = X4 | 0 = sK0 | X3 = X4 | sK0 = X5 | $less(X5,alength) | $less(X4,X5) (4:24) [resolution 581,287]
[SA] passive: 3128. $less(0,X9) | $less(X9,X10) | 0 = X10 | 0 = sK0 | X9 = X10 | ~$less(X11,X10) | $less(X11,sK0) (4:21) [resolution 581,18]
[SA] passive: 3136. $less(0,X1) | $less(X2,sK0) | 0 = X2 | 0 = sK0 | X1 = X2 | ~$less(X3,X1) | $less(X3,X2) (4:21) [resolution 581,18]
[SA] passive: 3137. $less(0,X4) | $less(X5,sK0) | 0 = X5 | 0 = sK0 | X4 = X5 | X5 = X6 | $less(X6,X5) | $less(X4,X6) (4:24) [resolution 581,130]
[SA] passive: 3140. $less(0,X11) | $less(X12,sK0) | 0 = X12 | 0 = sK0 | X11 = X12 | alength = X12 | $less(sK0,X12) | $less(X11,alength) (4:24) [resolution 581,354]
[SA] passive: 3145. $less(0,X20) | $less(i(l11(sK2(X21))),sK0) | i(l11(sK2(X21))) = 0 | 0 = sK0 | i(l11(sK2(X21))) = X20 | $less(X21,0) | ~$less(X21,alength) | $less(X20,alength) (4:33) [resolution 581,243]
[SA] passive: 3146. $less(0,X22) | $less(i(l11(p(nl8))),sK0) | i(l11(p(nl8))) = 0 | 0 = sK0 | i(l11(p(nl8))) = X22 | zero = nl8 | $less(X22,alength) (4:30) [resolution 581,359]
[SA] active: 908. 1 != sK0 | s(sK2(0)) = sK2(sK0) | alength = 0 (4:12:1) [superposition 340,319]
[SA] active: 582. $less(0,X21) | $less(X22,alength) | $less(X21,X22) | 0 = X22 | alength = 0 | X21 = X22 (3:18:3) [resolution 245,140]
[SA] passive: 3183. $less(X1,alength) | $less(X2,X1) | 0 = X1 | alength = 0 | X1 = X2 | ~$less(X3,0) | $less(X3,X2) (4:21) [resolution 582,18]
[SA] passive: 3184. $less(X4,alength) | $less(X5,X4) | 0 = X4 | alength = 0 | X4 = X5 | X5 = X6 | $less(X6,X5) | $less(0,X6) (4:24) [resolution 582,130]
[SA] passive: 3202. $less(0,X0) | $less(X0,X1) | 0 = X1 | alength = 0 | X0 = X1 | i(l11(sK2(X2))) = X2 | alength = X2 | $less(X2,0) | $less(X1,X2) (4:30) [resolution 582,261]
[SA] passive: 3205. $less(0,X6) | $less(X6,X7) | 0 = X7 | alength = 0 | X6 = X7 | ~$less(X8,X7) | $less(X8,alength) (4:21) [resolution 582,18]
[SA] passive: 3213. $less(0,X1) | $less(X2,alength) | 0 = X2 | alength = 0 | X1 = X2 | ~$less(X3,X1) | $less(X3,X2) (4:21) [resolution 582,18]
[SA] passive: 3214. $less(0,X4) | $less(X5,alength) | 0 = X5 | alength = 0 | X4 = X5 | X5 = X6 | $less(X6,X5) | $less(X4,X6) (4:24) [resolution 582,130]
[SA] active: 910. v = a(l10(s(s(sK2(0)))),1) | alength = 0 (4:12:1) [superposition 518,319]
[SA] passive: 3260. v = a(l11(s(s(sK2(0)))),1) | alength = 0 | i(l11(s(s(sK2(0))))) = 1 (5:20) [superposition 910,88]
[SA] active: 592. $less(0,X19) | $less(X20,sK0) | $less(X19,X20) | 0 = X19 | 0 = sK0 | 0 = X20 (3:18:3) [resolution 245,149]
[SA] passive: 3265. $less(X1,sK0) | $less(X2,X1) | 0 = X2 | 0 = sK0 | 0 = X1 | ~$less(X3,0) | $less(X3,X2) (4:21) [resolution 592,18]
[SA] passive: 3266. $less(X4,sK0) | $less(X5,X4) | 0 = X5 | 0 = sK0 | 0 = X4 | X5 = X6 | $less(X6,X5) | $less(0,X6) (4:24) [resolution 592,130]
[SA] passive: 3284. $less(0,X0) | $less(X0,X1) | 0 = X0 | 0 = sK0 | 0 = X1 | alength = X2 | $less(X2,alength) | $less(X1,X2) (4:24) [resolution 592,345]
[SA] passive: 3285. $less(0,X3) | $less(X3,X4) | 0 = X3 | 0 = sK0 | 0 = X4 | sK0 = X5 | $less(X5,alength) | $less(X4,X5) (4:24) [resolution 592,287]
[SA] passive: 3288. $less(0,X9) | $less(X9,X10) | 0 = X9 | 0 = sK0 | 0 = X10 | ~$less(X11,X10) | $less(X11,sK0) (4:21) [resolution 592,18]
[SA] passive: 3296. $less(0,X1) | $less(X2,sK0) | 0 = X1 | 0 = sK0 | 0 = X2 | ~$less(X3,X1) | $less(X3,X2) (4:21) [resolution 592,18]
[SA] passive: 3297. $less(0,X4) | $less(X5,sK0) | 0 = X4 | 0 = sK0 | 0 = X5 | X5 = X6 | $less(X6,X5) | $less(X4,X6) (4:24) [resolution 592,130]
[SA] active: 1470. $less(0,X4) | $less(X4,alength) | sK0 = X4 | 0 = X4 (4:12:2) [resolution 568,128]
[SA] passive: 3344. $less(X0,alength) | sK0 = X0 | 0 = X0 | ~$less(X1,0) | $less(X1,X0) (5:15) [resolution 1470,18]
[SA] passive: 3345. $less(X2,alength) | sK0 = X2 | 0 = X2 | X2 = X3 | $less(X3,X2) | $less(0,X3) (5:18) [resolution 1470,130]
[SA] passive: 3363. $less(0,X0) | sK0 = X0 | 0 = X0 | i(l11(sK2(X1))) = X1 | alength = X1 | $less(X1,0) | $less(X0,X1) (5:24) [resolution 1470,261]
[SA] passive: 3366. $less(0,X3) | sK0 = X3 | 0 = X3 | ~$less(X4,X3) | $less(X4,alength) (5:15) [resolution 1470,18]
[SA] active: 593. $less(0,X21) | $less(X22,alength) | $less(X21,X22) | 0 = X21 | alength = 0 | 0 = X22 (3:18:3) [resolution 245,140]
[SA] passive: 3382. $less(X1,alength) | $less(X2,X1) | 0 = X2 | alength = 0 | 0 = X1 | ~$less(X3,0) | $less(X3,X2) (4:21) [resolution 593,18]
[SA] passive: 3383. $less(X4,alength) | $less(X5,X4) | 0 = X5 | alength = 0 | 0 = X4 | X5 = X6 | $less(X6,X5) | $less(0,X6) (4:24) [resolution 593,130]
[SA] passive: 3401. $less(0,X0) | $less(X0,X1) | 0 = X0 | alength = 0 | 0 = X1 | i(l11(sK2(X2))) = X2 | alength = X2 | $less(X2,0) | $less(X1,X2) (4:30) [resolution 593,261]
[SA] passive: 3404. $less(0,X6) | $less(X6,X7) | 0 = X6 | alength = 0 | 0 = X7 | ~$less(X8,X7) | $less(X8,alength) (4:21) [resolution 593,18]
[SA] passive: 3412. $less(0,X1) | $less(X2,alength) | 0 = X1 | alength = 0 | 0 = X2 | ~$less(X3,X1) | $less(X3,X2) (4:21) [resolution 593,18]
[SA] passive: 3413. $less(0,X4) | $less(X5,alength) | 0 = X4 | alength = 0 | 0 = X5 | X5 = X6 | $less(X6,X5) | $less(X4,X6) (4:24) [resolution 593,130]
[SA] active: 252. ~'Sub'(X0,p(p(X0))) | zero = p(X0) | zero = X0 (4:12:1) [superposition 119,52]
[SA] active: 683. $less(0,X25) | $less(X25,0) | alength = X25 | i(l11(sK2(X25))) = X25 | 0 = sK0 (3:18:2) [resolution 261,148]
[SA] passive: 3459. $less(X0,0) | alength = X0 | i(l11(sK2(X0))) = X0 | 0 = sK0 | ~$less(X1,0) | $less(X1,X0) (4:21) [resolution 683,18]
[SA] passive: 3460. $less(X2,0) | alength = X2 | i(l11(sK2(X2))) = X2 | 0 = sK0 | X2 = X3 | $less(X3,X2) | $less(0,X3) (4:24) [resolution 683,130]
[SA] passive: 3503. $less(0,X0) | alength = X0 | i(l11(sK2(X0))) = X0 | 0 = sK0 | sK0 = X1 | $less(X1,sK0) | $less(X0,X1) (4:24) [duplicate literal removal 3480]
[SA] passive: 3501. $less(0,X5) | alength = X5 | i(l11(sK2(X5))) = X5 | 0 = sK0 | $less(X5,sK0) (4:18) [duplicate literal removal 3483]
[SA] passive: 3486. $less(0,X7) | alength = X7 | i(l11(sK2(X7))) = X7 | 0 = sK0 | ~$less(X8,X7) | $less(X8,0) (4:21) [resolution 683,18]
[SA] passive: 3487. $less(0,X9) | alength = X9 | i(l11(sK2(X9))) = X9 | 0 = sK0 | 0 = X10 | $less(X10,0) | $less(X9,X10) (4:24) [resolution 683,130]
[SA] active: 930. v = a(l11(s(s(nl8))),$sum($sum(alength,1),1)) (4:12:1) [superposition 89,195]
[SA] active: 469. 'Sub'(p(p(p(X3))),X3) | zero = p(p(X3)) | zero = p(X3) | zero = X3 (3:18:1) [resolution 197,116]
[SA] passive: 3532. zero = p(p(p(X2))) | zero = p(p(X2)) | zero = p(X2) | 'Sub'(p(p(p(X2))),s(X2)) | zero = X2 (4:25) [duplicate literal removal 3524]
[SA] passive: 3525. zero = p(p(p(X3))) | zero = p(p(X3)) | zero = p(X3) | 'Sub'(s(p(p(p(p(X3))))),X3) | zero = X3 (4:26) [resolution 469,115]
[SA] passive: 3526. zero = p(p(p(X4))) | zero = p(p(X4)) | zero = p(X4) | 'Sub'(p(p(p(p(X4)))),X4) | zero = X4 (4:25) [resolution 469,116]
[SA] passive: 3527. zero = p(p(nl8)) | zero = p(nl8) | zero = nl8 | i(l11(sK1(s(p(p(p(nl8)))),X5))) = X5 | a(main_end,X5) = a(l11(p(p(p(nl8)))),X5) (4:34) [resolution 469,283]
[SA] passive: 3528. zero = p(p(nl8)) | zero = p(nl8) | zero = nl8 | i(l11(sK1(p(p(p(nl8))),X6))) = X6 | a(main_end,X6) = a(l10(p(p(p(nl8)))),X6) (4:33) [resolution 469,279]
[SA] passive: 3529. zero = p(p(nl8)) | zero = p(nl8) | zero = nl8 | $less(i(l11(p(p(p(nl8))))),alength) (4:20) [resolution 469,93]
[SA] passive: 3540. zero = X0 | zero = p(X0) | 'Sub'(p(p(X0)),s(X0)) (4:13) [forward demodulation 3539,231]
[SA] active: 933. alength != $sum($sum(alength,1),1) | s(s(nl8)) = nl8 (4:12:1) [superposition 191,195]
[SA] active: 491. 'Sub'(X6,s(s(s(X7)))) | ~'Sub'(p(p(X6)),X7) | zero = p(X6) | zero = X6 (3:18:1) [resolution 216,114]
[SA] passive: 3550. ~'Sub'(p(p(p(X9))),X10) | zero = p(p(X9)) | zero = p(X9) | 'Sub'(X9,s(s(s(s(X10))))) | zero = X9 (4:25) [resolution 491,114]
[SA] passive: 3551. 'Sub'(X1,s(s(X0))) | ~'Sub'(p(p(X1)),p(X0)) | zero = p(X1) | zero = X1 | zero = X0 (4:21) [superposition 491,52]
[SA] active: 2396. 'Sub'(s(X0),s(X1)) | ~'Sub'(X0,p(X1)) | zero = X1 (4:12:1) [term algebras distinctness 2393]
[SA] passive: 3591. a(main_end,X1) = a(l11(X0),X1) | ~'Sub'(X0,p(nl8)) | zero = nl8 | i(l11(sK1(s(X0),X1))) = X1 (5:23) [forward demodulation 3579,153]
[SA] passive: 3587. ~'Sub'(s(s(s(s(s(s(s(s(X9)))))))),p(X9)) | zero = X9 (9:15) [resolution 2396,1507]
[SA] passive: 3590. 'Sub'(s(X1),X0) | ~'Sub'(X1,p(p(X0))) | zero = p(X0) | zero = X0 (5:16) [superposition 2396,52]
[SA] active: 564. ~$less(X9,X7) | $less(X8,X6) | X6 = X8 | X6 = X7 | $less(X6,X7) | $less(X9,X8) (3:18:1) T [resolution 245,18]
[SA] passive: 3594. $less(X8,X9) | X8 = X9 | X9 = X10 | $less(X9,X10) | $less(X11,X8) | $less(X12,X11) | $less(X10,X12) | X10 = X11 | X10 = X12 (4:27) T [resolution 564,245]
[SA] passive: 3595. $less(X13,X14) | X13 = X14 | X14 = X15 | $less(X14,X15) | $less(X16,X13) | $less(X15,X17) | $less(X17,X16) | X16 = X17 | X15 = X16 (4:27) T [resolution 564,245]
[SA] passive: 3596. $less(X18,X19) | X18 = X19 | X19 = X20 | $less(X19,X20) | $less(X21,X18) | $less(X22,X21) | $less(X20,X22) | X20 = X22 | X21 = X22 (4:27) T [resolution 564,245]
[SA] passive: 3597. $less(X23,X24) | X23 = X24 | X24 = X25 | $less(X24,X25) | $less(X26,X23) | $less(alength,X26) | $less(X25,0) | X25 = X26 | alength = X25 | i(l11(sK2(X25))) = X25 (4:33) [resolution 564,260]
[SA] passive: 3598. $less(X27,X28) | X27 = X28 | X28 = X29 | $less(X28,X29) | $less(X30,X27) | $less(alength,X30) | $less(X29,0) | 0 = X29 | alength = X30 | i(l11(sK2(X30))) = X30 (4:33) [resolution 564,268]
[SA] passive: 3599. $less(X31,X32) | X31 = X32 | X32 = X33 | $less(X32,X33) | $less(X34,X31) | $less(sK0,X34) | $less(X33,alength) | X33 = X34 | sK0 = X33 (4:27) [resolution 564,286]
[SA] passive: 3600. $less(X35,X36) | X35 = X36 | X36 = X37 | $less(X36,X37) | $less(X38,X35) | $less(sK0,X38) | $less(X37,alength) | alength = X37 | sK0 = X38 (4:27) [resolution 564,293]
[SA] passive: 3601. $less(X39,X40) | X39 = X40 | X40 = X41 | $less(X40,X41) | $less(X42,X39) | $less(sK0,X42) | $less(X41,alength) | X41 = X42 | alength = X41 (4:27) [resolution 564,344]
[SA] passive: 3602. $less(X43,X44) | X43 = X44 | X44 = X45 | $less(X44,X45) | $less(X46,X43) | $less(sK0,X46) | $less(X45,alength) | alength = X45 | alength = X46 (4:27) [resolution 564,353]
[SA] passive: 3603. $less(X47,X48) | X47 = X48 | X48 = X49 | $less(X48,X49) | $less(X50,X47) | $less(0,X50) | $less(X49,alength) | X49 = X50 | alength = 0 | alength = X49 (4:30) [resolution 564,424]
[SA] passive: 3604. $less(X51,X52) | X51 = X52 | X52 = X53 | $less(X52,X53) | $less(X54,X51) | $less(0,X54) | $less(X53,alength) | alength = X53 | alength = 0 | alength = X54 (4:30) [resolution 564,433]
[SA] passive: 3605. $less(X55,X56) | X55 = X56 | X56 = X57 | $less(X56,X57) | $less(X58,X55) | $less(0,X58) | $less(X57,sK0) | X57 = X58 | 0 = sK0 | sK0 = X57 (4:30) [resolution 564,450]
[SA] passive: 3606. $less(X59,X60) | X59 = X60 | X60 = X61 | $less(X60,X61) | $less(X62,X59) | $less(0,X62) | $less(X61,sK0) | sK0 = X61 | 0 = sK0 | sK0 = X62 (4:30) [resolution 564,459]
[SA] passive: 3607. $less(X63,X64) | X63 = X64 | X64 = X65 | $less(X64,X65) | $less(X66,X63) | $less(0,X66) | $less(X65,sK0) | X65 = X66 | 0 = sK0 | 0 = X66 (4:30) [resolution 564,570]
[SA] passive: 3608. $less(X67,X68) | X67 = X68 | X68 = X69 | $less(X68,X69) | $less(X70,X67) | $less(0,X70) | $less(X69,alength) | X69 = X70 | alength = 0 | 0 = X70 (4:30) [resolution 564,571]
[SA] passive: 3609. $less(X71,X72) | X71 = X72 | X72 = X73 | $less(X72,X73) | $less(X74,X71) | $less(sK0,X74) | $less(X73,alength) | X73 = X74 | sK0 = X74 (4:27) [resolution 564,572]
[SA] passive: 3610. $less(X75,X76) | X75 = X76 | X76 = X77 | $less(X76,X77) | $less(X78,X75) | $less(0,X78) | $less(X77,sK0) | 0 = X77 | 0 = sK0 | X77 = X78 (4:30) [resolution 564,581]
[SA] passive: 3611. $less(X79,X80) | X79 = X80 | X80 = X81 | $less(X80,X81) | $less(X82,X79) | $less(0,X82) | $less(X81,alength) | 0 = X81 | alength = 0 | X81 = X82 (4:30) [resolution 564,582]
[SA] passive: 3612. $less(X83,X84) | X83 = X84 | X84 = X85 | $less(X84,X85) | $less(X86,X83) | $less(0,X86) | $less(X85,sK0) | 0 = X86 | 0 = sK0 | 0 = X85 (4:30) [resolution 564,592]
[SA] passive: 3613. $less(X87,X88) | X87 = X88 | X88 = X89 | $less(X88,X89) | $less(X90,X87) | $less(0,X90) | $less(X89,alength) | 0 = X90 | alength = 0 | 0 = X89 (4:30) [resolution 564,593]
[SA] passive: 3614. $less(X91,X92) | X91 = X92 | X92 = X93 | $less(X92,X93) | $less(X94,X91) | $less(sK0,X94) | $less(X93,alength) | sK0 = X94 | sK0 = X93 (4:27) [resolution 564,594]
[SA] passive: 3616. $less(X98,X99) | X98 = X99 | alength = X99 | $less(X99,alength) | $less(X100,X98) | $less(0,X100) | alength = X100 | 0 = sK0 (4:24) [resolution 564,248]
[SA] passive: 3619. $less(X107,X108) | X107 = X108 | alength = X108 | $less(X108,alength) | $less(X109,X107) | $less(sK0,X110) | $less(X110,X109) | X109 = X110 | sK0 = X109 (4:27) [resolution 564,286]
[SA] passive: 3620. $less(X111,X112) | X111 = X112 | alength = X112 | $less(X112,alength) | $less(X113,X111) | $less(sK0,X114) | $less(X114,X113) | alength = X113 | sK0 = X114 (4:27) [resolution 564,293]
[SA] passive: 3621. $less(X115,X116) | X115 = X116 | alength = X116 | $less(X116,alength) | $less(X117,X115) | $less(0,X117) | alength = 0 | 0 = X117 (4:24) [resolution 564,300]
[SA] passive: 3622. $less(X118,X119) | X118 = X119 | alength = X119 | $less(X119,alength) | $less(X120,X118) | $less(sK0,X121) | $less(X121,X120) | X120 = X121 | alength = X120 (4:27) [resolution 564,344]
[SA] passive: 3623. $less(X122,X123) | X122 = X123 | alength = X123 | $less(X123,alength) | $less(X124,X122) | $less(sK0,X125) | $less(X125,X124) | alength = X124 | alength = X125 (4:27) [resolution 564,353]
[SA] passive: 3624. $less(X126,X127) | X126 = X127 | alength = X127 | $less(X127,alength) | $less(X128,X126) | $less(0,X128) | 0 = sK0 | 0 = X128 (5:24) [resolution 564,366]
[SA] passive: 3625. $less(X129,X130) | X129 = X130 | alength = X130 | $less(X130,alength) | $less(X131,X129) | $less(0,X132) | $less(X132,X131) | X131 = X132 | alength = 0 | alength = X131 (4:30) [resolution 564,424]
[SA] passive: 3626. $less(X133,X134) | X133 = X134 | alength = X134 | $less(X134,alength) | $less(X135,X133) | $less(0,X136) | $less(X136,X135) | alength = X135 | alength = 0 | alength = X136 (4:30) [resolution 564,433]
[SA] passive: 3627. $less(X137,X138) | X137 = X138 | alength = X138 | $less(X138,alength) | $less(X139,X137) | $less(0,X139) | 0 = sK0 | sK0 = X139 (4:24) [resolution 564,457]
[SA] passive: 3628. $less(X140,X141) | X140 = X141 | alength = X141 | $less(X141,alength) | $less(X142,X140) | $less(0,X142) | 0 = X142 | alength = X142 (4:24) [resolution 564,567]
[SA] passive: 3629. $less(X143,X144) | X143 = X144 | alength = X144 | $less(X144,alength) | $less(X145,X143) | $less(0,X146) | $less(X146,X145) | X145 = X146 | alength = 0 | 0 = X146 (4:30) [resolution 564,571]
[SA] passive: 3630. $less(X147,X148) | X147 = X148 | alength = X148 | $less(X148,alength) | $less(X149,X147) | $less(sK0,X150) | $less(X150,X149) | X149 = X150 | sK0 = X150 (4:27) [resolution 564,572]
[SA] passive: 3631. $less(X151,X152) | X151 = X152 | alength = X152 | $less(X152,alength) | $less(X153,X151) | $less(0,X154) | $less(X154,X153) | 0 = X153 | alength = 0 | X153 = X154 (4:30) [resolution 564,582]
[SA] passive: 3632. $less(X155,X156) | X155 = X156 | alength = X156 | $less(X156,alength) | $less(X157,X155) | $less(0,X158) | $less(X158,X157) | 0 = X158 | alength = 0 | 0 = X157 (4:30) [resolution 564,593]
[SA] passive: 3633. $less(X159,X160) | X159 = X160 | alength = X160 | $less(X160,alength) | $less(X161,X159) | $less(sK0,X162) | $less(X162,X161) | sK0 = X162 | sK0 = X161 (4:27) [resolution 564,594]
[SA] passive: 3634. $less(X163,X164) | X163 = X164 | alength = X164 | $less(X164,alength) | $less(X165,X163) | $less(0,X165) | sK0 = X165 | 0 = X165 (5:24) [resolution 564,1470]
[SA] passive: 3636. $less(X169,X170) | X169 = X170 | 0 = X170 | $less(X170,0) | $less(X171,X169) | $less(alength,X172) | $less(X172,X171) | X171 = X172 | alength = X171 | i(l11(sK2(X171))) = X171 (4:33) [resolution 564,260]
[SA] passive: 3637. $less(X173,X174) | X173 = X174 | 0 = X174 | $less(X174,0) | $less(X175,X173) | $less(alength,X176) | $less(X176,X175) | 0 = X175 | alength = X176 | i(l11(sK2(X176))) = X176 (4:33) [resolution 564,268]
[SA] passive: 3638. $less(X177,X178) | X177 = X178 | 0 = X178 | $less(X178,0) | $less(X179,X177) | $less(sK0,X179) | sK0 = X179 | i(l11(sK2(X179))) = X179 (4:27) [resolution 564,291]
[SA] passive: 3639. $less(X180,X181) | X180 = X181 | 0 = X181 | $less(X181,0) | $less(X182,X180) | $less(sK0,X182) | alength = X182 | i(l11(sK2(X182))) = X182 (4:27) [resolution 564,351]
[SA] passive: 3640. $less(X183,X184) | X183 = X184 | 0 = X184 | $less(X184,0) | $less(X185,X183) | $less(0,X185) | alength = 0 | alength = X185 | i(l11(sK2(X185))) = X185 (4:30) [resolution 564,431]
[SA] passive: 3641. $less(X186,X187) | X186 = X187 | 0 = X187 | $less(X187,0) | $less(X188,X186) | $less(0,X188) | alength = X188 | i(l11(sK2(X188))) = X188 | 0 = sK0 (4:30) [resolution 564,683]
[SA] passive: 3643. $less(X192,X193) | X192 = X193 | sK0 = X193 | $less(X193,sK0) | $less(X194,X192) | $less(alength,X194) | 0 = sK0 | i(l11(sK2(X194))) = X194 | alength = X194 (4:30) [resolution 564,306]
[SA] passive: 3644. $less(X195,X196) | X195 = X196 | sK0 = X196 | $less(X196,sK0) | $less(X197,X195) | $less(0,X197) | 0 = sK0 | 0 = X197 (4:24) [resolution 564,307]
[SA] passive: 3645. $less(X198,X199) | X198 = X199 | sK0 = X199 | $less(X199,sK0) | $less(X200,X198) | $less(0,X201) | $less(X201,X200) | X200 = X201 | 0 = sK0 | sK0 = X200 (4:30) [resolution 564,450]
[SA] passive: 3646. $less(X202,X203) | X202 = X203 | sK0 = X203 | $less(X203,sK0) | $less(X204,X202) | $less(0,X205) | $less(X205,X204) | sK0 = X204 | 0 = sK0 | sK0 = X205 (4:30) [resolution 564,459]
[SA] passive: 3647. $less(X206,X207) | X206 = X207 | sK0 = X207 | $less(X207,sK0) | $less(X208,X206) | $less(0,X208) | 0 = X208 | sK0 = X208 (4:24) [resolution 564,568]
[SA] passive: 3648. $less(X209,X210) | X209 = X210 | sK0 = X210 | $less(X210,sK0) | $less(X211,X209) | $less(0,X212) | $less(X212,X211) | X211 = X212 | 0 = sK0 | 0 = X212 (4:30) [resolution 564,570]
[SA] passive: 3649. $less(X213,X214) | X213 = X214 | sK0 = X214 | $less(X214,sK0) | $less(X215,X213) | $less(0,X216) | $less(X216,X215) | 0 = X215 | 0 = sK0 | X215 = X216 (4:30) [resolution 564,581]
[SA] passive: 3650. $less(X217,X218) | X217 = X218 | sK0 = X218 | $less(X218,sK0) | $less(X219,X217) | $less(0,X220) | $less(X220,X219) | 0 = X220 | 0 = sK0 | 0 = X219 (4:30) [resolution 564,592]
[SA] passive: 3651. $less(X221,X222) | X221 = X222 | X222 = X223 | $less(X222,X223) | $less(alength,X221) | $less(X223,sK0) | 0 = sK0 | i(l11(sK2(X223))) = X223 | alength = X223 (4:30) [resolution 564,306]
[SA] passive: 3652. $less(X224,X225) | X224 = X225 | X225 = X226 | $less(X225,X226) | $less(alength,X224) | $less(X227,0) | $less(X226,X227) | 0 = X227 | alength = X226 | i(l11(sK2(X226))) = X226 (4:33) [resolution 564,268]
[SA] passive: 3653. $less(X228,X229) | X228 = X229 | X229 = X230 | $less(X229,X230) | $less(alength,X228) | $less(X230,X231) | $less(X231,0) | X230 = X231 | alength = X231 | i(l11(sK2(X231))) = X231 (4:33) [resolution 564,260]
[SA] passive: 3655. $less(X235,X236) | X235 = X236 | X236 = X237 | $less(X236,X237) | $less(i(l11(p(nl8))),X235) | alength = X237 | $less(X237,alength) | zero = nl8 (4:27) [resolution 564,358]
[SA] passive: 3656. $less(X238,X239) | X238 = X239 | X239 = X240 | $less(X239,X240) | $less(i(l11(sK2(X241))),X238) | $less(X240,alength) | alength = X240 | ~$less(X241,alength) | $less(X241,0) (4:30) [resolution 564,247]
[SA] passive: 3659. $less(X247,X248) | X247 = X248 | alength = X248 | $less(X248,alength) | $less(i(l11(p(p(nl8)))),X247) | zero = p(nl8) | zero = nl8 (4:26) [resolution 564,470]
[SA] passive: 3660. $less(X249,X250) | X249 = X250 | alength = X250 | $less(X250,alength) | $less(i(l11(s(p(p(nl8))))),X249) | zero = p(nl8) | zero = nl8 (4:27) [resolution 564,526]
[SA] passive: 3661. $less(X251,X252) | X251 = X252 | 0 = X252 | $less(X252,0) | $less(i(l11(sK2(X253))),X251) | $less(X253,0) | ~$less(X253,alength) | i(l11(sK2(X253))) = i(l11(sK2(i(l11(sK2(X253)))))) (4:36) [resolution 564,242]
[SA] passive: 3665. $less(X260,X261) | X260 = X261 | X261 = X262 | $less(X261,X262) | $less(0,X260) | $less(X262,alength) | alength = X262 | 0 = sK0 (4:24) [resolution 564,248]
[SA] passive: 3670. $less(X275,X276) | X275 = X276 | X276 = X277 | $less(X276,X277) | $less(0,X275) | $less(X277,alength) | 0 = sK0 | 0 = X277 (5:24) [resolution 564,366]
[SA] passive: 3671. $less(X278,X279) | X278 = X279 | X279 = X280 | $less(X279,X280) | $less(0,X278) | $less(X280,X281) | $less(X281,alength) | X280 = X281 | alength = 0 | alength = X281 (4:30) [resolution 564,424]
[SA] passive: 3672. $less(X282,X283) | X282 = X283 | X283 = X284 | $less(X283,X284) | $less(0,X282) | $less(X284,0) | alength = 0 | alength = X284 | i(l11(sK2(X284))) = X284 (4:30) [resolution 564,431]
[SA] passive: 3673. $less(X285,X286) | X285 = X286 | X286 = X287 | $less(X286,X287) | $less(0,X285) | $less(X288,alength) | $less(X287,X288) | alength = X288 | alength = 0 | alength = X287 (4:30) [resolution 564,433]
[SA] passive: 3674. $less(X289,X290) | X289 = X290 | X290 = X291 | $less(X290,X291) | $less(0,X289) | $less(X291,X292) | $less(X292,sK0) | X291 = X292 | 0 = sK0 | sK0 = X292 (4:30) [resolution 564,450]
[SA] passive: 3675. $less(X293,X294) | X293 = X294 | X294 = X295 | $less(X294,X295) | $less(0,X293) | $less(X295,alength) | 0 = sK0 | sK0 = X295 (4:24) [resolution 564,457]
[SA] passive: 3676. $less(X296,X297) | X296 = X297 | X297 = X298 | $less(X297,X298) | $less(0,X296) | $less(X299,sK0) | $less(X298,X299) | sK0 = X299 | 0 = sK0 | sK0 = X298 (4:30) [resolution 564,459]
[SA] passive: 3677. $less(X300,X301) | X300 = X301 | X301 = X302 | $less(X301,X302) | $less(0,X300) | $less(X302,alength) | 0 = X302 | alength = X302 (4:24) [resolution 564,567]
[SA] passive: 3678. $less(X303,X304) | X303 = X304 | X304 = X305 | $less(X304,X305) | $less(0,X303) | $less(X305,sK0) | 0 = X305 | sK0 = X305 (4:24) [resolution 564,568]
[SA] passive: 3679. $less(X306,X307) | X306 = X307 | X307 = X308 | $less(X307,X308) | $less(0,X306) | $less(X309,sK0) | $less(X308,X309) | X308 = X309 | 0 = sK0 | 0 = X308 (4:30) [resolution 564,570]
[SA] passive: 3680. $less(X310,X311) | X310 = X311 | X311 = X312 | $less(X311,X312) | $less(0,X310) | $less(X313,alength) | $less(X312,X313) | X312 = X313 | alength = 0 | 0 = X312 (4:30) [resolution 564,571]
[SA] passive: 3681. $less(X314,X315) | X314 = X315 | X315 = X316 | $less(X315,X316) | $less(0,X314) | $less(X317,sK0) | $less(X316,X317) | 0 = X317 | 0 = sK0 | X316 = X317 (4:30) [resolution 564,581]
[SA] passive: 3682. $less(X318,X319) | X318 = X319 | X319 = X320 | $less(X319,X320) | $less(0,X318) | $less(X321,alength) | $less(X320,X321) | 0 = X321 | alength = 0 | X320 = X321 (4:30) [resolution 564,582]
[SA] passive: 3683. $less(X322,X323) | X322 = X323 | X323 = X324 | $less(X323,X324) | $less(0,X322) | $less(X325,sK0) | $less(X324,X325) | 0 = X324 | 0 = sK0 | 0 = X325 (4:30) [resolution 564,592]
[SA] passive: 3684. $less(X326,X327) | X326 = X327 | X327 = X328 | $less(X327,X328) | $less(0,X326) | $less(X329,alength) | $less(X328,X329) | 0 = X328 | alength = 0 | 0 = X329 (4:30) [resolution 564,593]
[SA] passive: 3685. $less(X330,X331) | X330 = X331 | X331 = X332 | $less(X331,X332) | $less(0,X330) | $less(X332,0) | alength = X332 | i(l11(sK2(X332))) = X332 | 0 = sK0 (4:30) [resolution 564,683]
[SA] passive: 3686. $less(X333,X334) | X333 = X334 | X334 = X335 | $less(X334,X335) | $less(0,X333) | $less(X335,alength) | sK0 = X335 | 0 = X335 (5:24) [resolution 564,1470]
[SA] passive: 3690. $less(X344,X345) | X344 = X345 | X345 = X346 | $less(X345,X346) | $less(sK0,X344) | $less(X346,X347) | $less(X347,alength) | X346 = X347 | sK0 = X347 (4:27) [resolution 564,286]
[SA] passive: 3691. $less(X348,X349) | X348 = X349 | X349 = X350 | $less(X349,X350) | $less(sK0,X348) | $less(X350,0) | sK0 = X350 | i(l11(sK2(X350))) = X350 (4:27) [resolution 564,291]
[SA] passive: 3692. $less(X351,X352) | X351 = X352 | X352 = X353 | $less(X352,X353) | $less(sK0,X351) | $less(X354,alength) | $less(X353,X354) | alength = X354 | sK0 = X353 (4:27) [resolution 564,293]
[SA] passive: 3693. $less(X355,X356) | X355 = X356 | X356 = X357 | $less(X356,X357) | $less(sK0,X355) | $less(X357,X358) | $less(X358,alength) | X357 = X358 | alength = X358 (4:27) [resolution 564,344]
[SA] passive: 3694. $less(X359,X360) | X359 = X360 | X360 = X361 | $less(X360,X361) | $less(sK0,X359) | $less(X361,0) | alength = X361 | i(l11(sK2(X361))) = X361 (4:27) [resolution 564,351]
[SA] passive: 3695. $less(X362,X363) | X362 = X363 | X363 = X364 | $less(X363,X364) | $less(sK0,X362) | $less(X365,alength) | $less(X364,X365) | alength = X365 | alength = X364 (4:27) [resolution 564,353]
[SA] passive: 3696. $less(X366,X367) | X366 = X367 | X367 = X368 | $less(X367,X368) | $less(sK0,X366) | $less(X369,alength) | $less(X368,X369) | X368 = X369 | sK0 = X368 (4:27) [resolution 564,572]
[SA] passive: 3697. $less(X370,X371) | X370 = X371 | X371 = X372 | $less(X371,X372) | $less(sK0,X370) | $less(X373,alength) | $less(X372,X373) | sK0 = X372 | sK0 = X373 (4:27) [resolution 564,594]
[SA] active: 1052. 1 != sK0 | s(sK2(0)) = sK2(sK0) | 0 = sK0 (5:12:1) [superposition 340,402]
[SA] active: 575. ~$less(X9,X8) | $less(X7,X8) | X7 = X8 | X6 = X8 | $less(X6,X7) | $less(X9,X6) (3:18:1) T [resolution 245,18]
[SA] passive: 3700. $less(X8,X9) | X8 = X9 | X9 = X10 | $less(X10,X8) | $less(X11,X10) | $less(X12,X11) | $less(X9,X12) | X9 = X11 | X9 = X12 (4:27) T [resolution 575,245]
[SA] passive: 3701. $less(X13,X14) | X13 = X14 | X14 = X15 | $less(X15,X13) | $less(X16,X15) | $less(X14,X17) | $less(X17,X16) | X16 = X17 | X14 = X16 (4:27) T [resolution 575,245]
[SA] passive: 3702. $less(X18,X19) | X18 = X19 | X19 = X20 | $less(X20,X18) | $less(X21,X20) | $less(X22,X21) | $less(X19,X22) | X19 = X22 | X21 = X22 (4:27) T [resolution 575,245]
[SA] passive: 3703. $less(X23,X24) | X23 = X24 | X24 = X25 | $less(X25,X23) | $less(X26,X25) | $less(alength,X26) | $less(X24,0) | X24 = X26 | alength = X24 | i(l11(sK2(X24))) = X24 (4:33) [resolution 575,260]
[SA] passive: 3704. $less(X27,X28) | X27 = X28 | X28 = X29 | $less(X29,X27) | $less(X30,X29) | $less(alength,X30) | $less(X28,0) | 0 = X28 | alength = X30 | i(l11(sK2(X30))) = X30 (4:33) [resolution 575,268]
[SA] passive: 3705. $less(X31,X32) | X31 = X32 | X32 = X33 | $less(X33,X31) | $less(X34,X33) | $less(sK0,X34) | $less(X32,alength) | X32 = X34 | sK0 = X32 (4:27) [resolution 575,286]
[SA] passive: 3706. $less(X35,X36) | X35 = X36 | X36 = X37 | $less(X37,X35) | $less(X38,X37) | $less(sK0,X38) | $less(X36,alength) | alength = X36 | sK0 = X38 (4:27) [resolution 575,293]
[SA] passive: 3707. $less(X39,X40) | X39 = X40 | X40 = X41 | $less(X41,X39) | $less(X42,X41) | $less(sK0,X42) | $less(X40,alength) | X40 = X42 | alength = X40 (4:27) [resolution 575,344]
[SA] passive: 3708. $less(X43,X44) | X43 = X44 | X44 = X45 | $less(X45,X43) | $less(X46,X45) | $less(sK0,X46) | $less(X44,alength) | alength = X44 | alength = X46 (4:27) [resolution 575,353]
[SA] passive: 3709. $less(X47,X48) | X47 = X48 | X48 = X49 | $less(X49,X47) | $less(X50,X49) | $less(0,X50) | $less(X48,alength) | X48 = X50 | alength = 0 | alength = X48 (4:30) [resolution 575,424]
[SA] passive: 3710. $less(X51,X52) | X51 = X52 | X52 = X53 | $less(X53,X51) | $less(X54,X53) | $less(0,X54) | $less(X52,alength) | alength = X52 | alength = 0 | alength = X54 (4:30) [resolution 575,433]
[SA] passive: 3711. $less(X55,X56) | X55 = X56 | X56 = X57 | $less(X57,X55) | $less(X58,X57) | $less(0,X58) | $less(X56,sK0) | X56 = X58 | 0 = sK0 | sK0 = X56 (4:30) [resolution 575,450]
[SA] passive: 3712. $less(X59,X60) | X59 = X60 | X60 = X61 | $less(X61,X59) | $less(X62,X61) | $less(0,X62) | $less(X60,sK0) | sK0 = X60 | 0 = sK0 | sK0 = X62 (4:30) [resolution 575,459]
[SA] passive: 3713. $less(X63,X64) | X63 = X64 | X64 = X65 | $less(X65,X63) | $less(X66,X65) | $less(0,X66) | $less(X64,sK0) | X64 = X66 | 0 = sK0 | 0 = X66 (4:30) [resolution 575,570]
[SA] passive: 3714. $less(X67,X68) | X67 = X68 | X68 = X69 | $less(X69,X67) | $less(X70,X69) | $less(0,X70) | $less(X68,alength) | X68 = X70 | alength = 0 | 0 = X70 (4:30) [resolution 575,571]
[SA] passive: 3715. $less(X71,X72) | X71 = X72 | X72 = X73 | $less(X73,X71) | $less(X74,X73) | $less(sK0,X74) | $less(X72,alength) | X72 = X74 | sK0 = X74 (4:27) [resolution 575,572]
[SA] passive: 3716. $less(X75,X76) | X75 = X76 | X76 = X77 | $less(X77,X75) | $less(X78,X77) | $less(0,X78) | $less(X76,sK0) | 0 = X76 | 0 = sK0 | X76 = X78 (4:30) [resolution 575,581]
[SA] passive: 3717. $less(X79,X80) | X79 = X80 | X80 = X81 | $less(X81,X79) | $less(X82,X81) | $less(0,X82) | $less(X80,alength) | 0 = X80 | alength = 0 | X80 = X82 (4:30) [resolution 575,582]
[SA] passive: 3718. $less(X83,X84) | X83 = X84 | X84 = X85 | $less(X85,X83) | $less(X86,X85) | $less(0,X86) | $less(X84,sK0) | 0 = X86 | 0 = sK0 | 0 = X84 (4:30) [resolution 575,592]
[SA] passive: 3719. $less(X87,X88) | X87 = X88 | X88 = X89 | $less(X89,X87) | $less(X90,X89) | $less(0,X90) | $less(X88,alength) | 0 = X90 | alength = 0 | 0 = X88 (4:30) [resolution 575,593]
[SA] passive: 3720. $less(X91,X92) | X91 = X92 | X92 = X93 | $less(X93,X91) | $less(X94,X93) | $less(sK0,X94) | $less(X92,alength) | sK0 = X94 | sK0 = X92 (4:27) [resolution 575,594]
[SA] passive: 3722. $less(X98,alength) | alength = X98 | alength = X99 | $less(X99,X98) | $less(X100,X99) | $less(0,X100) | alength = X100 | 0 = sK0 (4:24) [resolution 575,248]
[SA] passive: 3725. $less(X107,alength) | alength = X107 | alength = X108 | $less(X108,X107) | $less(X109,X108) | $less(sK0,X110) | $less(X110,X109) | X109 = X110 | sK0 = X109 (4:27) [resolution 575,286]
[SA] passive: 3726. $less(X111,alength) | alength = X111 | alength = X112 | $less(X112,X111) | $less(X113,X112) | $less(sK0,X114) | $less(X114,X113) | alength = X113 | sK0 = X114 (4:27) [resolution 575,293]
[SA] passive: 3727. $less(X115,alength) | alength = X115 | alength = X116 | $less(X116,X115) | $less(X117,X116) | $less(0,X117) | alength = 0 | 0 = X117 (4:24) [resolution 575,300]
[SA] passive: 3728. $less(X118,alength) | alength = X118 | alength = X119 | $less(X119,X118) | $less(X120,X119) | $less(sK0,X121) | $less(X121,X120) | X120 = X121 | alength = X120 (4:27) [resolution 575,344]
[SA] passive: 3729. $less(X122,alength) | alength = X122 | alength = X123 | $less(X123,X122) | $less(X124,X123) | $less(sK0,X125) | $less(X125,X124) | alength = X124 | alength = X125 (4:27) [resolution 575,353]
[SA] passive: 3730. $less(X126,alength) | alength = X126 | alength = X127 | $less(X127,X126) | $less(X128,X127) | $less(0,X128) | 0 = sK0 | 0 = X128 (5:24) [resolution 575,366]
[SA] passive: 3731. $less(X129,alength) | alength = X129 | alength = X130 | $less(X130,X129) | $less(X131,X130) | $less(0,X132) | $less(X132,X131) | X131 = X132 | alength = 0 | alength = X131 (4:30) [resolution 575,424]
[SA] passive: 3732. $less(X133,alength) | alength = X133 | alength = X134 | $less(X134,X133) | $less(X135,X134) | $less(0,X136) | $less(X136,X135) | alength = X135 | alength = 0 | alength = X136 (4:30) [resolution 575,433]
[SA] passive: 3733. $less(X137,alength) | alength = X137 | alength = X138 | $less(X138,X137) | $less(X139,X138) | $less(0,X139) | 0 = sK0 | sK0 = X139 (4:24) [resolution 575,457]
[SA] passive: 3734. $less(X140,alength) | alength = X140 | alength = X141 | $less(X141,X140) | $less(X142,X141) | $less(0,X142) | 0 = X142 | alength = X142 (4:24) [resolution 575,567]
[SA] passive: 3735. $less(X143,alength) | alength = X143 | alength = X144 | $less(X144,X143) | $less(X145,X144) | $less(0,X146) | $less(X146,X145) | X145 = X146 | alength = 0 | 0 = X146 (4:30) [resolution 575,571]
[SA] passive: 3736. $less(X147,alength) | alength = X147 | alength = X148 | $less(X148,X147) | $less(X149,X148) | $less(sK0,X150) | $less(X150,X149) | X149 = X150 | sK0 = X150 (4:27) [resolution 575,572]
[SA] passive: 3737. $less(X151,alength) | alength = X151 | alength = X152 | $less(X152,X151) | $less(X153,X152) | $less(0,X154) | $less(X154,X153) | 0 = X153 | alength = 0 | X153 = X154 (4:30) [resolution 575,582]
[SA] passive: 3738. $less(X155,alength) | alength = X155 | alength = X156 | $less(X156,X155) | $less(X157,X156) | $less(0,X158) | $less(X158,X157) | 0 = X158 | alength = 0 | 0 = X157 (4:30) [resolution 575,593]
[SA] passive: 3739. $less(X159,alength) | alength = X159 | alength = X160 | $less(X160,X159) | $less(X161,X160) | $less(sK0,X162) | $less(X162,X161) | sK0 = X162 | sK0 = X161 (4:27) [resolution 575,594]
[SA] passive: 3740. $less(X163,alength) | alength = X163 | alength = X164 | $less(X164,X163) | $less(X165,X164) | $less(0,X165) | sK0 = X165 | 0 = X165 (5:24) [resolution 575,1470]
[SA] passive: 3742. $less(X169,0) | 0 = X169 | 0 = X170 | $less(X170,X169) | $less(X171,X170) | $less(alength,X172) | $less(X172,X171) | X171 = X172 | alength = X171 | i(l11(sK2(X171))) = X171 (4:33) [resolution 575,260]
[SA] passive: 3743. $less(X173,0) | 0 = X173 | 0 = X174 | $less(X174,X173) | $less(X175,X174) | $less(alength,X176) | $less(X176,X175) | 0 = X175 | alength = X176 | i(l11(sK2(X176))) = X176 (4:33) [resolution 575,268]
[SA] passive: 3744. $less(X177,0) | 0 = X177 | 0 = X178 | $less(X178,X177) | $less(X179,X178) | $less(sK0,X179) | sK0 = X179 | i(l11(sK2(X179))) = X179 (4:27) [resolution 575,291]
[SA] passive: 3745. $less(X180,0) | 0 = X180 | 0 = X181 | $less(X181,X180) | $less(X182,X181) | $less(sK0,X182) | alength = X182 | i(l11(sK2(X182))) = X182 (4:27) [resolution 575,351]
[SA] passive: 3746. $less(X183,0) | 0 = X183 | 0 = X184 | $less(X184,X183) | $less(X185,X184) | $less(0,X185) | alength = 0 | alength = X185 | i(l11(sK2(X185))) = X185 (4:30) [resolution 575,431]
[SA] passive: 3747. $less(X186,0) | 0 = X186 | 0 = X187 | $less(X187,X186) | $less(X188,X187) | $less(0,X188) | alength = X188 | i(l11(sK2(X188))) = X188 | 0 = sK0 (4:30) [resolution 575,683]
[SA] passive: 3749. $less(X192,sK0) | sK0 = X192 | sK0 = X193 | $less(X193,X192) | $less(X194,X193) | $less(alength,X194) | 0 = sK0 | i(l11(sK2(X194))) = X194 | alength = X194 (4:30) [resolution 575,306]
[SA] passive: 3750. $less(X195,sK0) | sK0 = X195 | sK0 = X196 | $less(X196,X195) | $less(X197,X196) | $less(0,X197) | 0 = sK0 | 0 = X197 (4:24) [resolution 575,307]
[SA] passive: 3751. $less(X198,sK0) | sK0 = X198 | sK0 = X199 | $less(X199,X198) | $less(X200,X199) | $less(0,X201) | $less(X201,X200) | X200 = X201 | 0 = sK0 | sK0 = X200 (4:30) [resolution 575,450]
[SA] passive: 3752. $less(X202,sK0) | sK0 = X202 | sK0 = X203 | $less(X203,X202) | $less(X204,X203) | $less(0,X205) | $less(X205,X204) | sK0 = X204 | 0 = sK0 | sK0 = X205 (4:30) [resolution 575,459]
[SA] passive: 3753. $less(X206,sK0) | sK0 = X206 | sK0 = X207 | $less(X207,X206) | $less(X208,X207) | $less(0,X208) | 0 = X208 | sK0 = X208 (4:24) [resolution 575,568]
[SA] passive: 3754. $less(X209,sK0) | sK0 = X209 | sK0 = X210 | $less(X210,X209) | $less(X211,X210) | $less(0,X212) | $less(X212,X211) | X211 = X212 | 0 = sK0 | 0 = X212 (4:30) [resolution 575,570]
[SA] passive: 3755. $less(X213,sK0) | sK0 = X213 | sK0 = X214 | $less(X214,X213) | $less(X215,X214) | $less(0,X216) | $less(X216,X215) | 0 = X215 | 0 = sK0 | X215 = X216 (4:30) [resolution 575,581]
[SA] passive: 3756. $less(X217,sK0) | sK0 = X217 | sK0 = X218 | $less(X218,X217) | $less(X219,X218) | $less(0,X220) | $less(X220,X219) | 0 = X220 | 0 = sK0 | 0 = X219 (4:30) [resolution 575,592]
[SA] passive: 3757. $less(X221,X222) | X221 = X222 | X222 = X223 | $less(X223,X221) | $less(alength,X223) | $less(X222,sK0) | 0 = sK0 | i(l11(sK2(X222))) = X222 | alength = X222 (4:30) [resolution 575,306]
[SA] passive: 3758. $less(X224,X225) | X224 = X225 | X225 = X226 | $less(X226,X224) | $less(alength,X226) | $less(X227,0) | $less(X225,X227) | 0 = X227 | alength = X225 | i(l11(sK2(X225))) = X225 (4:33) [resolution 575,268]
[SA] passive: 3759. $less(X228,X229) | X228 = X229 | X229 = X230 | $less(X230,X228) | $less(alength,X230) | $less(X229,X231) | $less(X231,0) | X229 = X231 | alength = X231 | i(l11(sK2(X231))) = X231 (4:33) [resolution 575,260]
[SA] passive: 3761. $less(X235,X236) | X235 = X236 | X236 = X237 | $less(X237,X235) | $less(i(l11(p(nl8))),X237) | alength = X236 | $less(X236,alength) | zero = nl8 (4:27) [resolution 575,358]
[SA] passive: 3762. $less(X238,X239) | X238 = X239 | X239 = X240 | $less(X240,X238) | $less(i(l11(sK2(X241))),X240) | $less(X239,alength) | alength = X239 | ~$less(X241,alength) | $less(X241,0) (4:30) [resolution 575,247]
[SA] passive: 3763. $less(X242,alength) | alength = X242 | alength = X243 | $less(X243,X242) | $less(i(l11(sK2(X244))),X243) | ~$less(X244,alength) | $less(X244,0) (4:24) [resolution 575,150]
[SA] passive: 3764. $less(X245,alength) | alength = X245 | alength = X246 | $less(X246,X245) | $less(i(l11(p(nl8))),X246) | zero = nl8 (4:21) [resolution 575,121]
[SA] passive: 3765. $less(X247,alength) | alength = X247 | alength = X248 | $less(X248,X247) | $less(i(l11(p(p(nl8)))),X248) | zero = p(nl8) | zero = nl8 (4:26) [resolution 575,470]
[SA] passive: 3766. $less(X249,alength) | alength = X249 | alength = X250 | $less(X250,X249) | $less(i(l11(s(p(p(nl8))))),X250) | zero = p(nl8) | zero = nl8 (4:27) [resolution 575,526]
[SA] passive: 3767. $less(X251,0) | 0 = X251 | 0 = X252 | $less(X252,X251) | $less(i(l11(sK2(X253))),X252) | $less(X253,0) | ~$less(X253,alength) | i(l11(sK2(X253))) = i(l11(sK2(i(l11(sK2(X253)))))) (4:36) [resolution 575,242]
[SA] passive: 3771. $less(X260,X261) | X260 = X261 | X261 = X262 | $less(X262,X260) | $less(0,X262) | $less(X261,alength) | alength = X261 | 0 = sK0 (4:24) [resolution 575,248]
[SA] passive: 3776. $less(X275,X276) | X275 = X276 | X276 = X277 | $less(X277,X275) | $less(0,X277) | $less(X276,alength) | 0 = sK0 | 0 = X276 (5:24) [resolution 575,366]
[SA] passive: 3777. $less(X278,X279) | X278 = X279 | X279 = X280 | $less(X280,X278) | $less(0,X280) | $less(X279,X281) | $less(X281,alength) | X279 = X281 | alength = 0 | alength = X281 (4:30) [resolution 575,424]
[SA] passive: 3778. $less(X282,X283) | X282 = X283 | X283 = X284 | $less(X284,X282) | $less(0,X284) | $less(X283,0) | alength = 0 | alength = X283 | i(l11(sK2(X283))) = X283 (4:30) [resolution 575,431]
[SA] passive: 3779. $less(X285,X286) | X285 = X286 | X286 = X287 | $less(X287,X285) | $less(0,X287) | $less(X288,alength) | $less(X286,X288) | alength = X288 | alength = 0 | alength = X286 (4:30) [resolution 575,433]
[SA] passive: 3780. $less(X289,X290) | X289 = X290 | X290 = X291 | $less(X291,X289) | $less(0,X291) | $less(X290,X292) | $less(X292,sK0) | X290 = X292 | 0 = sK0 | sK0 = X292 (4:30) [resolution 575,450]
[SA] passive: 3781. $less(X293,X294) | X293 = X294 | X294 = X295 | $less(X295,X293) | $less(0,X295) | $less(X294,alength) | 0 = sK0 | sK0 = X294 (4:24) [resolution 575,457]
[SA] passive: 3782. $less(X296,X297) | X296 = X297 | X297 = X298 | $less(X298,X296) | $less(0,X298) | $less(X299,sK0) | $less(X297,X299) | sK0 = X299 | 0 = sK0 | sK0 = X297 (4:30) [resolution 575,459]
[SA] passive: 3783. $less(X300,X301) | X300 = X301 | X301 = X302 | $less(X302,X300) | $less(0,X302) | $less(X301,alength) | 0 = X301 | alength = X301 (4:24) [resolution 575,567]
[SA] passive: 3784. $less(X303,X304) | X303 = X304 | X304 = X305 | $less(X305,X303) | $less(0,X305) | $less(X304,sK0) | 0 = X304 | sK0 = X304 (4:24) [resolution 575,568]
[SA] passive: 3785. $less(X306,X307) | X306 = X307 | X307 = X308 | $less(X308,X306) | $less(0,X308) | $less(X309,sK0) | $less(X307,X309) | X307 = X309 | 0 = sK0 | 0 = X307 (4:30) [resolution 575,570]
[SA] passive: 3786. $less(X310,X311) | X310 = X311 | X311 = X312 | $less(X312,X310) | $less(0,X312) | $less(X313,alength) | $less(X311,X313) | X311 = X313 | alength = 0 | 0 = X311 (4:30) [resolution 575,571]
[SA] passive: 3787. $less(X314,X315) | X314 = X315 | X315 = X316 | $less(X316,X314) | $less(0,X316) | $less(X317,sK0) | $less(X315,X317) | 0 = X317 | 0 = sK0 | X315 = X317 (4:30) [resolution 575,581]
[SA] passive: 3788. $less(X318,X319) | X318 = X319 | X319 = X320 | $less(X320,X318) | $less(0,X320) | $less(X321,alength) | $less(X319,X321) | 0 = X321 | alength = 0 | X319 = X321 (4:30) [resolution 575,582]