[SA] active: 67. ~$less(sK0,0) (0:3:1) [cnf transformation 57]
[SA] active: 88. $less(sK0,alength0) (0:3:1) [definition unfolding 66,83]
[SA] active: 96. ~$less(alength0,0) (0:3:1) [definition unfolding 84,83]
[SA] active: 21. ~$less(X0,X0) (0:3:1) T [theory axiom]
[SA] active: 85. i(zero) = 0 (0:4:1) [cnf transformation 3]
[SA] active: 87. 'Sub'(X0,s(X0)) (0:4:1) [cnf transformation 1]
[SA] active: 90. ~$less(i('$n'),alength0) (0:4:1) [definition unfolding 68,83]
[SA] active: 29. zero != s(X0) (0:4:1) T [term algebras distinctness]
[SA] active: 103. v0 != a('$n',sK0) (0:5:1) [backward demodulation 95,89]
[SA] active: 18. $sum(X0,0) = X0 (0:5:1) T [theory axiom]
[SA] active: 26. $uminus($uminus(X0)) = X0 (0:5:1) T [theory axiom]
[SA] active: 20. 0 = $sum(X0,$uminus(X0)) (0:6:1) T [theory axiom]
[SA] active: 109. 0 = $sum($uminus(X0),X0) (1:6:1) T [superposition 20,26]
[SA] active: 86. 'Sub'(X0,s(X1)) | ~'Sub'(X0,X1) (0:7:1) [cnf transformation 55]
[SA] active: 91. ~'Sub'(X0,'$n') | $less(i(X0),alength0) (0:7:1) [definition unfolding 69,83]
[SA] active: 95. v0 = a(s(X0),i(X0)) (0:7:1) [definition unfolding 82,77]
[SA] active: 113. v0 = a(s(zero),0) (1:6:1) [superposition 95,85]
[SA] active: 16. $sum(X0,X1) = $sum(X1,X0) (0:7:1) T [theory axiom]
[SA] active: 75. i(s(X0)) = $sum(i(X0),1) (0:8:1) [cnf transformation 50]
[SA] active: 78. ~'Sub'(X0,X1) | $less(i(X0),i(X1)) (0:8:1) [cnf transformation 51]
[SA] active: 131. i(s(zero)) = 1 (1:5:1) [evaluation 126]
[SA] active: 79. i(X0) != i(X1) | X0 = X1 (0:8:1) [cnf transformation 52]
[SA] active: 132. $less(i(X0),i(s(X0))) (1:6:1) [resolution 78,87]
[SA] active: 25. $less(X1,$sum(X0,1)) | $less(X0,X1) (0:8:1) T [theory axiom]
[SA] active: 148. $less(X0,$sum(X0,1)) (1:5:1) T [resolution 25,21]
[SA] active: 27. ~$less(X1,$sum(X0,1)) | ~$less(X0,X1) (0:8:1) T [theory axiom]
[SA] active: 154. $less(X0,$sum(1,X0)) (2:5:1) T [superposition 148,16]
[SA] active: 30. s(X0) != s(X1) | X0 = X1 (0:8:1) T [term algebras injectivity]
[SA] active: 136. i(s(s(zero))) = 2 (2:6:1) [evaluation 134]
[SA] active: 64. s(p(X0)) = X0 | zero = X0 (0:8:1) [cnf transformation 28]
[SA] active: 153. $less(X4,0) | $less(-1,X4) (2:6:2) T [evaluation 151]
[SA] active: 22. ~$less(X1,X2) | ~$less(X0,X1) | $less(X0,X2) (0:9:1) T [theory axiom]
[SA] active: 187. $less(-1,sK0) (3:3:1) [resolution 153,67]
[SA] active: 23. $less(X1,X0) | $less(X0,X1) | X0 = X1 (0:9:2) T [theory axiom]
[SA] active: 188. $less(-1,alength0) (3:3:1) [resolution 153,96]
[SA] active: 19. $sum($uminus(X1),$uminus(X0)) = $uminus($sum(X0,X1)) (0:10:1) T [theory axiom]
[SA] active: 199. ~$less(X11,sK0) | $less(X11,alength0) (1:6:1) [resolution 22,88]
[SA] active: 24. $less($sum(X0,X2),$sum(X1,X2)) | ~$less(X0,X1) (0:10:1) T [theory axiom]
[SA] active: 205. $less(0,sK0) | 0 = sK0 (1:6:1) [resolution 23,67]
[SA] active: 73. 'Sub'(sK2(X0),'$n') | a(zero,X0) = a('$n',X0) (0:11:1) [cnf transformation 61]
[SA] active: 206. $less(0,alength0) | alength0 = 0 (1:6:1) [resolution 23,96]
[SA] active: 108. ~$less(X0,i('$n')) | 'Sub'(sK3(X0),'$n') | $less(X0,0) (0:11:1) [backward demodulation 85,81]
[SA] active: 255. $less(0,alength0) | 0 = sK0 (2:6:1) [resolution 205,199]
[SA] active: 17. $sum(X0,$sum(X1,X2)) = $sum($sum(X0,X1),X2) (0:11:1) T [theory axiom]
[SA] active: 164. ~$less(-1,X4) | ~$less(X4,0) (2:6:1) T [evaluation 162]
[SA] active: 74. a(zero,X0) = a('$n',X0) | i(sK2(X0)) = X0 (0:12:1) [cnf transformation 61]
[SA] active: 201. ~$less(X0,-1) | $less(X0,sK0) (4:6:1) [resolution 187,22]
[SA] active: 76. a(X0,X1) = a(s(X0),X1) | i(X0) = X1 (0:12:1) [cnf transformation 50]
[SA] active: 215. ~$less(X0,-1) | $less(X0,alength0) (4:6:1) [resolution 188,22]
[SA] active: 107. ~$less(X0,i('$n')) | i(sK3(X0)) = X0 | $less(X0,0) (0:12:1) [backward demodulation 85,80]
[SA] active: 137. i(X0) != 0 | zero = X0 (1:7:1) [superposition 79,85]
[SA] active: 104. 'Sub'(sK1(X1,i(X1)),'$n') | v0 = a('$n',i(X1)) | ~'Sub'(X1,'$n') (0:15:1) [backward demodulation 95,102]
[SA] active: 184. 'Sub'(p(X7),X7) | zero = X7 (1:7:1) [superposition 87,64]
[SA] active: 105. sK1(X1,i(X1)) != X1 | v0 = a('$n',i(X1)) | ~'Sub'(X1,'$n') (0:15:1) [backward demodulation 95,100]
[SA] active: 135. v0 = a(s(s(zero)),1) (2:7:1) [superposition 95,131]
[SA] active: 106. ~'Sub'(X1,'$n') | i(X1) = i(sK1(X1,i(X1))) | v0 = a('$n',i(X1)) (0:17:1) [backward demodulation 95,98]
[SA] active: 177. i(s(s(s(zero)))) = 3 (3:7:1) [evaluation 175]
[SA] active: 127. i(s(X1)) = $sum(1,i(X1)) (1:8:1) [superposition 75,16]
[SA] active: 204. $less(alength0,i('$n')) | alength0 = i('$n') (1:8:1) [resolution 23,90]
[SA] active: 149. $less(X1,$sum(1,X0)) | $less(X0,X1) (1:8:1) T [superposition 25,16]
[SA] active: 378. $less(X2,1) | $less(0,X2) (2:6:2) T [superposition 149,18]
[SA] active: 160. ~$less(X1,$sum(1,X0)) | ~$less(X0,X1) (1:8:1) T [superposition 27,16]
[SA] active: 410. ~$less(0,X2) | ~$less(X2,1) (2:6:1) T [superposition 160,18]
[SA] active: 288. $sum(X2,$sum($uminus(X2),X3)) = X3 (1:8:1) T [evaluation 265]
[SA] active: 451. $less($sum(-1,X17),X17) (3:5:1) T [evaluation 449]
[SA] active: 133. ~'Sub'(X1,X2) | $less(i(X1),i(s(X2))) (1:9:1) [resolution 78,86]
[SA] active: 472. $less(-2,sK0) (5:3:1) [evaluation 465]
[SA] active: 152. $less(X6,i(s(X5))) | $less(i(X5),X6) (1:9:1) [superposition 25,75]
[SA] active: 473. $less(-2,alength0) (5:3:1) [evaluation 464]
[SA] active: 163. ~$less(X6,i(s(X5))) | ~$less(i(X5),X6) (1:9:1) [superposition 27,75]
[SA] active: 463. $less($sum(-1,sK0),alength0) (4:5:1) [resolution 451,199]
[SA] active: 240. $less(0,$sum(X3,$uminus(X2))) | ~$less(X2,X3) (1:9:1) T [superposition 24,20]
[SA] active: 466. $less($sum(X0,-1),X0) (4:5:1) T [superposition 451,16]
[SA] active: 247. $less($sum(X3,$uminus(X2)),0) | ~$less(X3,X2) (1:9:1) T [superposition 24,20]
[SA] active: 393. $less(sK0,1) | $less(0,alength0) (3:6:1) [resolution 378,199]
[SA] active: 183. ~'Sub'(X6,p(X5)) | 'Sub'(X6,X5) | zero = X5 (1:10:1) [superposition 86,64]
[SA] active: 419. ~$less(sK0,1) | 0 = sK0 (3:6:1) [resolution 410,205]
[SA] active: 185. v0 = a(X8,i(p(X8))) | zero = X8 (1:10:1) [superposition 95,64]
[SA] active: 420. ~$less(alength0,1) | 0 = sK0 (3:6:1) [resolution 410,255]
[SA] active: 303. v0 != a(zero,sK0) | i(sK2(sK0)) = sK0 (1:10:1) [superposition 103,74]
[SA] active: 421. ~$less(alength0,1) | alength0 = 0 (3:6:1) [resolution 410,206]
[SA] active: 216. $sum(X0,$uminus(X1)) = $uminus($sum(X1,$uminus(X0))) (1:10:1) T [superposition 19,26]
[SA] active: 496. $less(1,X1) | $less(X1,2) (3:6:2) [forward demodulation 493,131]
[SA] active: 217. $sum($uminus(X1),X0) = $uminus($sum($uminus(X0),X1)) (1:10:1) T [superposition 19,26]
[SA] active: 511. ~$less(1,X1) | ~$less(X1,2) (3:6:1) [forward demodulation 508,131]
[SA] active: 220. $sum($uminus(X5),$uminus(X4)) = $uminus($sum(X5,X4)) (1:10:1) T [superposition 19,16]
[SA] active: 497. $less(2,X2) | $less(X2,3) (4:6:2) [forward demodulation 494,136]
[SA] active: 241. $less($sum(X5,X4),$sum(X6,X5)) | ~$less(X4,X6) (1:10:1) T [superposition 24,16]
[SA] active: 512. ~$less(2,X2) | ~$less(X2,3) (4:6:1) [forward demodulation 509,136]
[SA] active: 248. $less($sum(X6,X5),$sum(X5,X4)) | ~$less(X6,X4) (1:10:1) T [superposition 24,16]
[SA] active: 614. $less(sK0,2) | $less(1,alength0) (4:6:1) [resolution 496,199]
[SA] active: 245. $less(i(s(X15)),$sum(X16,1)) | ~$less(i(X15),X16) (1:11:1) [superposition 24,75]
[SA] active: 729. $less(sK0,3) | $less(2,alength0) (5:6:1) [resolution 497,199]
[SA] active: 252. $less($sum(X16,1),i(s(X15))) | ~$less(X16,i(X15)) (1:11:1) [superposition 24,75]
[SA] active: 486. ~$less(X0,-2) | $less(X0,sK0) (6:6:1) [resolution 472,22]
[SA] active: 181. s(X2) != X1 | p(X1) = X2 | zero = X1 (1:11:1) [superposition 30,64]
[SA] active: 869. $less(-3,sK0) (7:3:1) [evaluation 864]
[SA] active: 194. ~$less(X0,X1) | $less(X0,$sum(X2,1)) | $less(X2,X1) (1:11:1) T [resolution 22,25]
[SA] active: 873. $less(-3,alength0) (8:3:1) [resolution 869,199]
[SA] active: 266. $sum(X4,$sum(X5,X6)) = $sum($sum(X5,X4),X6) (1:11:1) T [superposition 17,16]
[SA] active: 872. p(s(X0)) = X0 (2:5:1) [term algebras distinctness 871]
[SA] active: 274. $sum(X7,$sum(X8,X9)) = $sum(X9,$sum(X7,X8)) (1:11:1) T [superposition 17,16]
[SA] active: 921. $less(-2,X75) | $less(X75,alength0) (4:6:2) [evaluation 920]
[SA] active: 257. $less(i(sK2(X0)),alength0) | a(zero,X0) = a('$n',X0) (1:12:1) [resolution 73,91]
[SA] active: 1129. $less(-2,i('$n')) (5:4:1) [resolution 921,90]
[SA] active: 268. $sum(i(s(X10)),X11) = $sum(i(X10),$sum(1,X11)) (1:12:1) [superposition 17,75]
[SA] active: 923. $less(-2,X74) | $less(X74,sK0) (4:6:2) [evaluation 922]
[SA] active: 203. ~$less(X3,X2) | X1 = X2 | $less(X1,X2) | $less(X3,X1) (1:12:1) T [resolution 23,22]
[SA] active: 498. ~$less(X0,-2) | $less(X0,alength0) (6:6:1) [resolution 473,22]
[SA] active: 281. ~$less(X18,$sum(X16,$sum(X17,1))) | ~$less($sum(X16,X17),X18) (1:12:1) T [superposition 27,17]
[SA] active: 917. $less(-3,X77) | $less(X77,alength0) (6:6:2) [evaluation 916]
[SA] active: 283. $less(X23,$sum(X21,$sum(X22,1))) | $less($sum(X21,X22),X23) (1:12:1) T [superposition 25,17]
[SA] active: 1338. $less(-3,i('$n')) (7:4:1) [resolution 917,90]
[SA] active: 258. $less(i(sK2(X1)),i('$n')) | a(zero,X1) = a('$n',X1) (1:13:1) [resolution 73,78]
[SA] active: 919. $less(-3,X76) | $less(X76,sK0) (6:6:2) [evaluation 918]
[SA] active: 326. 'Sub'(sK1(zero,0),'$n') | v0 = a('$n',0) | ~'Sub'(zero,'$n') (1:13:1) [forward demodulation 323,85]
[SA] active: 874. ~$less(X0,-3) | $less(X0,sK0) (8:6:1) [resolution 869,22]
[SA] active: 334. zero != sK1(zero,0) | v0 = a('$n',0) | ~'Sub'(zero,'$n') (1:13:1) [forward demodulation 331,85]
[SA] active: 1443. $less(-4,sK0) (9:3:1) [evaluation 1433]
[SA] active: 207. $less($sum(X4,1),X5) | $sum(X4,1) = X5 | ~$less(X4,X5) (1:13:1) T [resolution 23,27]
[SA] active: 1444. $less(-4,alength0) (10:3:1) [resolution 1443,199]
[SA] active: 238. ~$less(X5,$sum(X3,X6)) | ~$less(X3,X4) | $less(X5,$sum(X4,X6)) (1:13:1) T [resolution 24,22]
[SA] active: 915. $less(-4,X78) | $less(X78,sK0) (8:6:2) [evaluation 914]
[SA] active: 244. $less($uminus($sum(X13,X12)),$sum(X14,$uminus(X13))) | ~$less($uminus(X12),X14) (1:13:1) T [superposition 24,19]
[SA] active: 941. ~$less(X1,-3) | $less(X1,alength0) (9:6:1) [resolution 873,22]
[SA] active: 251. $less($sum(X14,$uminus(X13)),$uminus($sum(X13,X12))) | ~$less(X14,$uminus(X12)) (1:13:1) T [superposition 24,19]
[SA] active: 943. $less(-4,X0) | $less(X0,alength0) (9:6:2) [evaluation 942]
[SA] active: 271. $sum($uminus($sum(X19,X18)),X20) = $sum($uminus(X18),$sum($uminus(X19),X20)) (1:14:1) T [superposition 17,19]
[SA] active: 1706. $less(-4,i('$n')) (10:4:1) [resolution 943,90]
[SA] active: 276. $less($sum(X3,X2),$sum(X0,$sum(X1,X2))) | ~$less(X3,$sum(X0,X1)) (1:14:1) T [superposition 24,17]
[SA] active: 1447. ~$less(X2,-4) | $less(X2,sK0) (10:6:1) [resolution 1443,22]
[SA] active: 277. $less($sum(X4,$sum(X5,X6)),$sum(X7,X6)) | ~$less($sum(X4,X5),X7) (1:14:1) T [superposition 24,17]
[SA] active: 1906. $less(-5,sK0) (11:3:1) [evaluation 1891]
[SA] active: 260. 'Sub'(sK3(X0),'$n') | $less(X0,0) | $less(i('$n'),X0) | i('$n') = X0 (1:15:1) [resolution 108,23]
[SA] active: 1970. $less(-5,alength0) (12:3:1) [resolution 1906,199]
[SA] active: 307. a(X0,X1) = a(p(X0),X1) | i(p(X0)) = X1 | zero = X0 (1:16:1) [superposition 76,64]
[SA] active: 1449. $less(-5,X1) | $less(X1,sK0) (10:6:2) [evaluation 1448]
[SA] active: 315. $less(i('$n'),X0) | $less(X0,0) | i(sK3(X0)) = X0 | i('$n') = X0 (1:16:1) [resolution 107,23]
[SA] active: 1494. ~$less(X2,-4) | $less(X2,alength0) (11:6:1) [resolution 1444,22]
[SA] active: 321. ~'Sub'(X0,'$n') | v0 = a('$n',i(X0)) | $less(i(sK1(X0,i(X0))),alength0) (1:16:1) [resolution 104,91]
[SA] active: 1496. $less(-5,X1) | $less(X1,alength0) (11:6:2) [evaluation 1495]
[SA] active: 322. ~'Sub'(X1,'$n') | v0 = a('$n',i(X1)) | $less(i(sK1(X1,i(X1))),i('$n')) (1:17:1) [resolution 104,78]
[SA] active: 2095. $less(-5,i('$n')) (12:4:1) [resolution 1496,90]
[SA] active: 339. i(sK2(X0)) = i(sK1(sK2(X0),i(sK2(X0)))) | v0 = a('$n',i(sK2(X0))) | a(zero,X0) = a('$n',X0) (1:25:1) [resolution 106,73]
[SA] active: 1973. ~$less(X2,-5) | $less(X2,sK0) (12:6:1) [resolution 1906,22]
[SA] active: 340. ~'Sub'(X1,'$n') | v0 = a('$n',i(sK1(X1,i(X1)))) | v0 = a('$n',i(X1)) | i(sK1(X1,i(X1))) = i(sK1(sK1(X1,i(X1)),i(sK1(X1,i(X1))))) (1:35:1) [resolution 106,104]
[SA] active: 2155. $less(-6,sK0) (13:3:1) [evaluation 2138]
[SA] active: 480. $less(i(X0),i(s(s(X0)))) (2:7:1) [resolution 133,87]
[SA] active: 2161. $less(-6,alength0) (14:3:1) [resolution 2155,199]
[SA] active: 514. ~$less($uminus(X0),0) | ~$less(X0,1) (2:7:1) T [resolution 240,160]
[SA] active: 2167. ~$less(i(s(X0)),i(X0)) (3:6:1) [resolution 480,163]
[SA] active: 138. i(X1) != 1 | s(zero) = X1 (2:8:1) [superposition 79,131]
[SA] active: 1975. $less(-6,X1) | $less(X1,sK0) (12:6:2) [evaluation 1974]
[SA] active: 329. $less(i(p('$n')),alength0) | zero = '$n' (2:8:1) [resolution 184,91]
[SA] active: 1982. ~$less(X2,-5) | $less(X2,alength0) (13:6:1) [resolution 1970,22]
[SA] active: 372. 'Sub'(sK3(alength0),'$n') | alength0 = i('$n') (2:8:1) [subsumption resolution 369,96]
[SA] active: 1984. $less(-6,X1) | $less(X1,alength0) (13:6:2) [evaluation 1983]
[SA] active: 903. $less(sK0,$sum(X66,1)) | $less(X66,alength0) (2:8:1) [resolution 194,88]
[SA] active: 2307. $less(-6,i('$n')) (14:4:1) [resolution 1984,90]
[SA] active: 1177. i(X15) = $sum(i(s(X15)),-1) (2:8:1) [evaluation 1158]
[SA] active: 2164. ~$less(X2,-6) | $less(X2,sK0) (14:6:1) [resolution 2155,22]
[SA] active: 195. $less(X3,$sum(X4,1)) | ~$less(X3,X4) (2:8:1) T [resolution 22,148]
[SA] active: 2394. $less(-7,sK0) (15:3:1) [evaluation 2375]
[SA] active: 287. $sum($uminus(X16),$sum(X16,X17)) = X17 (2:8:1) T [evaluation 270]
[SA] active: 2421. $less(-7,alength0) (16:3:1) [resolution 2394,199]
[SA] active: 407. ~$less(X9,$sum(X10,X9)) | ~$less(X10,1) (2:8:1) T [resolution 160,24]
[SA] active: 2398. ~$less($sum(X6,1),X6) (3:5:1) T [resolution 195,21]
[SA] active: 438. $sum(X3,$sum(X4,$uminus(X3))) = X4 (2:8:1) T [superposition 288,16]
[SA] active: 2621. $less(X9,$sum(2,X9)) (4:5:1) [forward demodulation 2620,1181]
[SA] active: 452. $less($sum(-1,X15),X16) | $less(X16,X15) (2:8:1) T [evaluation 448]
[SA] active: 2570. ~$less($sum(1,X5),X5) (4:5:1) T [resolution 2398,248]
[SA] active: 453. ~$less($sum(-1,X13),X14) | ~$less(X14,X13) (2:8:1) T [evaluation 447]
[SA] active: 2824. ~$less(alength0,sK0) (5:3:1) [resolution 453,463]
[SA] active: 744. ~$less(X3,$sum(X3,X2)) | ~$less(X2,1) (2:8:1) T [resolution 241,160]
[SA] active: 2713. $less(X0,$sum(X0,2)) (5:5:1) [superposition 2621,16]
[SA] active: 749. $less(X6,$sum(X7,X6)) | ~$less(0,X7) (2:8:1) T [superposition 241,18]
[SA] active: 2721. $less($sum(X7,-2),X7) (5:5:1) [evaluation 2719]
[SA] active: 804. $less($sum(X7,X6),X6) | ~$less(X7,0) (2:8:1) T [superposition 248,18]
[SA] active: 3063. $less(-8,sK0) (15:3:1) [evaluation 3057]
[SA] active: 1086. $sum($uminus(X19),$sum(X20,X19)) = X20 (2:8:1) T [evaluation 1008]
[SA] active: 3134. $less(-8,alength0) (16:3:1) [resolution 3063,199]
[SA] active: 1306. ~$less($sum(X18,-1),X17) | ~$less(X17,X18) (2:8:1) T [evaluation 1296]
[SA] active: 2722. $less($sum(-2,X2),X2) (5:5:1) [evaluation 2716]
[SA] active: 1374. $less($sum(X18,-1),X17) | $less(X17,X18) (2:8:1) T [evaluation 1364]
[SA] active: 2796. ~$less(X8,$sum(X8,-1)) (5:5:1) T [evaluation 2794]
[SA] active: 186. $less(i(p(X9)),i(X9)) | zero = X9 (2:9:1) [superposition 132,64]
[SA] active: 2797. ~$less(X3,$sum(-1,X3)) (5:5:1) T [evaluation 2791]
[SA] active: 198. $less(X9,i(s(X10))) | ~$less(X9,i(X10)) (2:9:1) [resolution 22,132]
[SA] active: 3088. $less($sum(-2,sK0),alength0) (6:5:1) [forward demodulation 3046,16]
[SA] active: 232. $less(sK0,X0) | $less(X0,alength0) | sK0 = X0 (2:9:2) [resolution 199,23]
[SA] active: 3536. ~$less(X0,0) | $less(X0,1) (3:6:1) [forward demodulation 3530,85]
[SA] active: 256. ~$less(X0,0) | 0 = sK0 | $less(X0,sK0) (2:9:1) [resolution 205,22]
[SA] active: 3537. ~$less(X1,1) | $less(X1,2) (3:6:1) [forward demodulation 3531,131]
[SA] active: 259. ~$less(X0,0) | alength0 = 0 | $less(X0,alength0) (2:9:1) [resolution 206,22]
[SA] active: 2395. ~$less(X1,X0) | ~$less(X0,X1) (3:6:1) T [resolution 195,27]
[SA] active: 371. alength0 = i(sK3(alength0)) | alength0 = i('$n') (2:9:1) [subsumption resolution 368,96]
[SA] active: 3811. ~$less(sK0,-1) (4:3:1) [resolution 2395,187]
[SA] active: 856. $less(0,i(s(X9))) | ~$less(-1,i(X9)) (2:9:1) [evaluation 845]
[SA] active: 3812. ~$less(alength0,-1) (4:3:1) [resolution 2395,188]
[SA] active: 931. $less(-1,X40) | $less(X40,sK0) | 0 = sK0 (2:9:2) [evaluation 930]
[SA] active: 3815. ~$less(alength0,-2) (6:3:1) [resolution 2395,473]
[SA] active: 935. $less(-1,X36) | $less(X36,alength0) | alength0 = 0 (2:9:2) [evaluation 934]
[SA] active: 3817. ~$less(sK0,-2) (6:3:1) [resolution 2395,472]
[SA] active: 1181. $sum(2,X1) = $sum(1,$sum(1,X1)) (2:9:1) [forward demodulation 1148,136]
[SA] active: 3822. ~$less(sK0,-3) (8:3:1) [resolution 2395,869]
[SA] active: 1248. $less(sK0,X72) | $less(X72,alength0) | alength0 = X72 (2:9:2) [resolution 203,88]
[SA] active: 3820. ~$less(alength0,-3) (9:3:1) [resolution 2395,873]
[SA] active: 2047. i(sK3(sK0)) = sK0 | i('$n') = sK0 (2:9:1) [subsumption resolution 2046,90]
[SA] active: 3827. ~$less(sK0,-4) (10:3:1) [resolution 2395,1443]
[SA] active: 243. $less(0,$sum(X11,X10)) | ~$less($uminus(X10),X11) (2:9:1) T [superposition 24,109]
[SA] active: 3825. ~$less(alength0,-4) (11:3:1) [resolution 2395,1444]
[SA] active: 250. $less($sum(X11,X10),0) | ~$less(X11,$uminus(X10)) (2:9:1) T [superposition 24,109]
[SA] active: 3832. ~$less(sK0,-5) (12:3:1) [resolution 2395,1906]
[SA] active: 282. $less($sum(X19,X20),$sum(X19,$sum(X20,1))) (2:9:1) T [superposition 148,17]
[SA] active: 3830. ~$less(alength0,-5) (13:3:1) [resolution 2395,1970]
[SA] active: 437. $sum(X1,$uminus($sum(X2,X1))) = $uminus(X2) (2:9:1) T [superposition 288,19]
[SA] active: 3837. ~$less(sK0,-6) (14:3:1) [resolution 2395,2155]
[SA] active: 519. $less(0,$sum($uminus(X2),X1)) | ~$less(X2,X1) (2:9:1) T [superposition 240,16]
[SA] active: 3835. ~$less(alength0,-6) (15:3:1) [resolution 2395,2161]
[SA] active: 557. $less($sum($uminus(X2),X1),0) | ~$less(X1,X2) (2:9:1) T [superposition 247,16]
[SA] active: 3838. ~$less(sK0,-7) (16:3:1) [resolution 2395,2394]
[SA] active: 656. $sum($uminus($sum(X8,X9)),X8) = $uminus(X9) (2:9:1) T [forward demodulation 655,19]
[SA] active: 3840. ~$less(sK0,-8) (16:3:1) [resolution 2395,3063]
[SA] active: 686. $uminus($sum(X2,X3)) = $uminus($sum(X3,X2)) (2:9:1) T [superposition 220,19]
[SA] active: 3839. ~$less(alength0,-7) (17:3:1) [resolution 2395,2421]
[SA] active: 694. $sum(X5,$uminus($sum(X5,X6))) = $uminus(X6) (2:9:1) T [superposition 288,220]
[SA] active: 3841. ~$less(alength0,-8) (17:3:1) [resolution 2395,3134]
[SA] active: 766. $less($sum(X21,X22),0) | ~$less(X22,$uminus(X21)) (2:9:1) T [superposition 241,109]
[SA] active: 3816. ~$less(i('$n'),-2) (6:4:1) [resolution 2395,1129]
[SA] active: 799. $less(0,$sum(X21,X22)) | ~$less($uminus(X21),X22) (2:9:1) T [superposition 248,109]
[SA] active: 3821. ~$less(i('$n'),-3) (8:4:1) [resolution 2395,1338]
[SA] active: 976. $less($sum(X48,X49),$sum(X49,$sum(X48,1))) (2:9:1) T [superposition 148,266]
[SA] active: 3826. ~$less(i('$n'),-4) (11:4:1) [resolution 2395,1706]
[SA] active: 1473. ~$less(-1,X12) | 0 = X12 | $less(0,X12) (2:9:1) T [evaluation 1472]
[SA] active: 3831. ~$less(i('$n'),-5) (13:4:1) [resolution 2395,2095]
[SA] active: 481. ~'Sub'(X1,X2) | $less(i(X1),i(s(s(X2)))) (2:10:1) [resolution 133,86]
[SA] active: 3836. ~$less(i('$n'),-6) (15:4:1) [resolution 2395,2307]
[SA] active: 484. $less(i(p(X5)),i(s(X5))) | zero = X5 (2:10:1) [resolution 133,184]
[SA] active: 3754. ~$less($sum(2,X17),X17) (5:5:1) [resolution 2395,2621]
[SA] active: 522. $less(0,$uminus($sum(X9,X8))) | ~$less(X9,$uminus(X8)) (2:10:1) T [superposition 240,19]
[SA] active: 3792. ~$less(alength0,$sum(-1,sK0)) (5:5:1) [resolution 2395,463]
[SA] active: 560. $less($uminus($sum(X9,X8)),0) | ~$less($uminus(X8),X9) (2:10:1) T [superposition 247,19]
[SA] active: 3761. ~$less($sum(X30,2),X30) (6:5:1) [resolution 2395,2713]
[SA] active: 586. $sum(X2,$uminus(X1)) = $uminus($sum($uminus(X2),X1)) (2:10:1) T [superposition 216,16]
[SA] active: 3786. ~$less(X52,$sum(X52,-2)) (6:5:1) [resolution 2395,2721]
[SA] active: 632. $sum($uminus(X4),X3) = $uminus($sum(X4,$uminus(X3))) (2:10:1) T [superposition 217,16]
[SA] active: 3787. ~$less(X53,$sum(-2,X53)) (6:5:1) [resolution 2395,2722]
[SA] active: 697. $less($uminus($sum(X10,X11)),0) | ~$less($uminus(X10),X11) (2:10:1) T [superposition 247,220]
[SA] active: 3793. ~$less(alength0,$sum(-2,sK0)) (7:5:1) [resolution 2395,3088]
[SA] active: 698. $less(0,$uminus($sum(X12,X13))) | ~$less(X13,$uminus(X12)) (2:10:1) T [superposition 240,220]
[SA] active: 2412. ~$less(X12,-1) | $less(X12,0) (3:6:1) T [evaluation 2411]
[SA] active: 758. $less($sum(X1,X2),$sum(X1,X0)) | ~$less(X2,X0) (2:10:1) T [superposition 241,16]
[SA] active: 3538. ~$less(X2,2) | $less(X2,3) (4:6:1) [forward demodulation 3532,136]
[SA] active: 881. $less(X16,$sum(X17,1)) | $less(X17,$sum(X16,1)) (2:10:2) T [resolution 194,148]
[SA] active: 6280. $less(X12,0) | $less(-2,X12) (3:6:2) T [evaluation 6279]
[SA] active: 1277. ~$less($sum(1,X1),X2) | $less($sum(X1,1),X2) (2:10:1) T [resolution 281,149]
[SA] active: 6505. $less(X42,-1) | $less(-2,X42) (3:6:2) T [evaluation 6504]
[SA] active: 1350. ~$less($sum(X3,1),X4) | $less($sum(1,X3),X4) (2:10:1) T [resolution 283,160]
[SA] active: 3590. $less(X0,1) | $less(-1,X0) (4:6:2) [resolution 3536,153]
[SA] active: 362. $less($sum(X4,i(X3)),i(s(X3))) | ~$less(X4,1) (2:11:1) [superposition 24,127]
[SA] active: 3621. $less(X7,2) | $less(0,X7) (4:6:2) [evaluation 3620]
[SA] active: 363. $less(i(s(X5)),$sum(X6,i(X5))) | ~$less(1,X6) (2:11:1) [superposition 24,127]
[SA] active: 3679. $less(sK0,2) | $less(0,alength0) (4:6:1) [resolution 3537,393]
[SA] active: 370. $less(X0,i('$n')) | ~$less(X0,alength0) | alength0 = i('$n') (2:11:1) [resolution 204,22]
[SA] active: 3696. $less(X5,3) | $less(1,X5) (4:6:2) [evaluation 3695]
[SA] active: 763. $less($sum(1,X14),i(s(X13))) | ~$less(X14,i(X13)) (2:11:1) [superposition 241,75]
[SA] active: 3767. ~$less(2,X36) | $less(1,X36) (4:6:1) [resolution 2395,496]
[SA] active: 764. $less($sum(i(X15),X16),i(s(X15))) | ~$less(X16,1) (2:11:1) [superposition 241,127]
[SA] active: 3805. ~$less(1,sK0) | $less(0,alength0) (4:6:1) [resolution 2395,393]
[SA] active: 796. $less(i(s(X13)),$sum(1,X14)) | ~$less(i(X13),X14) (2:11:1) [superposition 248,75]
[SA] active: 6345. $less(-2,X2) | $less(X2,1) (4:6:2) [resolution 6280,3536]
[SA] active: 797. $less(i(s(X15)),$sum(i(X15),X16)) | ~$less(1,X16) (2:11:1) [superposition 248,127]
[SA] active: 3751. ~$less(0,X14) | $less(-1,X14) (4:6:1) T [resolution 2395,153]
[SA] active: 887. $less(i(X28),$sum(X29,1)) | $less(X29,i(s(X28))) (2:11:1) [resolution 194,132]
[SA] active: 3752. ~$less(1,X15) | $less(0,X15) (4:6:1) T [resolution 2395,378]
[SA] active: 376. ~$less(X5,X4) | $less(X3,X4) | $less(X5,$sum(1,X3)) (2:11:1) T [resolution 149,22]
[SA] active: 6359. ~$less(0,X4) | $less(-2,X4) (4:6:1) T [resolution 6280,2395]
[SA] active: 445. $less($sum(X5,$sum($uminus(X3),X4)),X4) | ~$less(X5,X3) (2:11:1) T [superposition 24,288]
[SA] active: 6366. ~$less(X0,-2) | $less(X0,0) (4:6:1) T [resolution 6280,2395]
[SA] active: 446. $less(X7,$sum(X8,$sum($uminus(X6),X7))) | ~$less(X6,X8) (2:11:1) T [superposition 24,288]
[SA] active: 6404. ~$less(X13,-2) | $less(X13,-1) (4:6:1) T [evaluation 6403]
[SA] active: 593. 0 = $sum($sum(X6,$uminus(X7)),$sum(X7,$uminus(X6))) (2:11:1) T [superposition 20,216]
[SA] active: 6507. $less(X41,-1) | $less(-3,X41) (4:6:2) T [evaluation 6506]
[SA] active: 641. 0 = $sum($sum($uminus(X8),X9),$sum($uminus(X9),X8)) (2:11:1) T [superposition 20,217]
[SA] active: 6566. ~$less(-1,X3) | $less(-2,X3) (4:6:1) T [resolution 6505,2395]
[SA] active: 750. ~$less($sum($uminus(X8),X9),X10) | $less(X9,$sum(X10,X8)) (2:11:1) T [superposition 241,288]
[SA] active: 6624. $less(-3,X12) | $less(X12,-2) (4:6:2) T [evaluation 6623]
[SA] active: 770. $less($sum($uminus(X8),$sum(X9,X10)),X9) | ~$less(X10,X8) (2:11:1) T [forward demodulation 761,17]
[SA] active: 2873. ~$less(sK0,X49) | $less(-1,X49) (5:6:1) [evaluation 2872]
[SA] active: 805. ~$less(X10,$sum($uminus(X8),X9)) | $less($sum(X10,X8),X9) (2:11:1) T [superposition 248,288]
[SA] active: 2887. ~$less(alength0,X28) | $less(-1,X28) (5:6:1) [evaluation 2886]
[SA] active: 814. $less(X9,$sum($uminus(X8),$sum(X9,X10))) | ~$less(X8,X10) (2:11:1) T [forward demodulation 794,17]
[SA] active: 3768. ~$less(3,X37) | $less(2,X37) (5:6:1) [resolution 2395,497]
[SA] active: 956. $sum(X3,$sum(X4,X5)) = $sum(X4,$sum(X3,X5)) (2:11:1) T [superposition 266,17]
[SA] active: 3807. ~$less(2,sK0) | $less(1,alength0) (5:6:1) [resolution 2395,614]
[SA] active: 957. $sum(X7,$sum(X6,X8)) = $sum(X8,$sum(X6,X7)) (2:11:1) T [superposition 266,16]
[SA] active: 6189. $less(sK0,3) | $less(1,alength0) (5:6:1) [resolution 3538,614]
[SA] active: 1018. $sum(X10,$sum(X8,X9)) = $sum($sum(X9,X10),X8) (2:11:1) T [superposition 274,16]
[SA] active: 6210. $less(X5,4) | $less(2,X5) (5:6:2) [evaluation 6209]
[SA] active: 1498. $less(X3,$sum(X2,X4)) | ~$less(1,X2) | $less(X4,X3) (2:11:1) T [resolution 238,149]
[SA] active: 6746. $less(-1,X0) | $less(X0,2) (5:6:2) [resolution 3590,3537]
[SA] active: 1499. ~$less(X5,X6) | $less(X7,$sum(X6,1)) | $less(X5,X7) (2:11:1) T [resolution 238,25]
[SA] active: 6757. ~$less(1,X1) | $less(-1,X1) (5:6:1) [resolution 3590,2395]
[SA] active: 1566. ~$less($uminus(X5),$uminus($sum(X5,X4))) | ~$less($uminus(X4),1) (2:11:1) T [resolution 244,160]
[SA] active: 6768. ~$less(X2,-1) | $less(X2,1) (5:6:1) [resolution 3590,2395]
[SA] active: 1873. ~$less(X49,$sum(X50,$uminus(X48))) | $less($sum(X49,X48),X50) (2:11:1) T [evaluation 1859]
[SA] active: 6866. $less(0,X0) | $less(X0,3) (5:6:2) [resolution 3621,3538]
[SA] active: 1874. $less($sum(X26,$uminus(X25)),X27) | ~$less(X26,$sum(X27,X25)) (2:11:1) T [evaluation 1853]
[SA] active: 6877. ~$less(2,X1) | $less(0,X1) (5:6:1) [resolution 3621,2395]
[SA] active: 1957. ~$less($sum(X49,$uminus(X48)),X50) | $less(X49,$sum(X50,X48)) (2:11:1) T [evaluation 1926]
[SA] active: 6887. ~$less(X1,0) | $less(X1,2) (5:6:1) [resolution 3621,2395]
[SA] active: 1958. $less(X26,$sum(X27,$uminus(X25))) | ~$less($sum(X26,X25),X27) (2:11:1) T [evaluation 1920]
[SA] active: 6984. $less(sK0,3) | $less(0,alength0) (5:6:1) [resolution 3679,3538]
[SA] active: 361. $sum(i(s(X1)),X2) = $sum(1,$sum(i(X1),X2)) (2:12:1) [superposition 17,127]
[SA] active: 12326. $less(X8,$sum(3,X8)) (6:5:1) [backward demodulation 12323,5412]
[SA] active: 490. ~$less(X5,X4) | $less(i(X3),X4) | $less(X5,i(s(X3))) (2:12:1) [resolution 152,22]
[SA] active: 12327. ~$less($sum(3,X7),X7) (6:5:1) [backward demodulation 12323,5411]
[SA] active: 491. $less(i(p(X0)),X1) | $less(X1,i(X0)) | zero = X0 (2:12:1) [superposition 152,64]
[SA] active: 12663. $less(-9,sK0) (13:3:1) [evaluation 12662]
[SA] active: 506. ~$less(i(p(X0)),X1) | ~$less(X1,i(X0)) | zero = X0 (2:12:1) [superposition 163,64]
[SA] active: 12667. $less(-9,alength0) (14:3:1) [evaluation 12666]
[SA] active: 566. 'Sub'(p(p(X0)),X0) | zero = X0 | zero = p(X0) (2:12:1) [resolution 183,184]
[SA] active: 12751. ~$less(sK0,-9) (14:3:1) [resolution 12663,2395]
[SA] active: 825. $less(i(s(X5)),i(s(X4))) | ~$less(i(X5),i(X4)) (2:12:1) [superposition 245,75]
[SA] active: 12834. ~$less(alength0,-9) (15:3:1) [resolution 12667,2395]
[SA] active: 1009. $sum(X22,i(s(X21))) = $sum(1,$sum(X22,i(X21))) (2:12:1) [superposition 274,75]
[SA] active: 12406. $less($sum(X15,-3),X15) (7:5:1) [evaluation 12402]
[SA] active: 1010. $sum(X24,i(s(X23))) = $sum(i(X23),$sum(X24,1)) (2:12:1) [superposition 274,127]
[SA] active: 12407. $less($sum(-3,X2),X2) (7:5:1) [evaluation 12395]
[SA] active: 1030. $sum(X22,i(s(X21))) = $sum(i(X21),$sum(1,X22)) (2:12:1) [superposition 274,75]
[SA] active: 12419. $less(X19,$sum(X19,3)) (7:5:1) [forward demodulation 12418,1061]
[SA] active: 1031. $sum(X24,i(s(X23))) = $sum(1,$sum(i(X23),X24)) (2:12:1) [superposition 274,127]
[SA] active: 12594. ~$less($sum(X2,3),X2) (7:5:1) [resolution 12327,241]
[SA] active: 1152. $sum(i(s(X3)),X2) = $sum(i(X3),$sum(X2,1)) (2:12:1) [superposition 268,16]
[SA] active: 12652. ~$less(X15,$sum(X15,-3)) (7:5:1) [evaluation 12648]
[SA] active: 1159. $sum(i(s(X2)),X3) = $sum(1,$sum(X3,i(X2))) (2:12:1) [superposition 268,274]
[SA] active: 12653. ~$less(X2,$sum(-3,X2)) (7:5:1) [evaluation 12641]
[SA] active: 1161. $sum(i(s(X6)),X7) = $sum($sum(1,X7),i(X6)) (2:12:1) [superposition 268,16]
[SA] active: 12688. ~$less($sum(4,X13),X13) (7:5:1) [forward demodulation 12687,12329]
[SA] active: 1178. $sum(i(X8),X7) = $sum(i(s(X8)),$sum(-1,X7)) (2:12:1) [evaluation 1155]
[SA] active: 14438. $less(-10,sK0) (13:3:1) [evaluation 14437]
[SA] active: 1236. $less(0,X42) | $less(X42,alength0) | alength0 = X42 | alength0 = 0 (2:12:2) [resolution 203,206]
[SA] active: 14442. $less(-10,alength0) (14:3:1) [evaluation 14441]
[SA] active: 1238. $less(0,X46) | $less(X46,sK0) | sK0 = X46 | 0 = sK0 (2:12:2) [resolution 203,205]
[SA] active: 14657. ~$less(sK0,-10) (14:3:1) [resolution 14438,2395]
[SA] active: 516. $less(X5,$sum(X4,$uminus(X3))) | ~$less(X5,0) | ~$less(X3,X4) (2:12:1) T [resolution 240,22]
[SA] active: 14816. ~$less(alength0,-10) (15:3:1) [resolution 14442,2395]
[SA] active: 554. ~$less(X2,$sum(X0,$uminus(X1))) | ~$less(X0,X1) | $less(X2,0) (2:12:1) T [resolution 247,22]
[SA] active: 12690. $less(X14,$sum(4,X14)) (7:5:1) [forward demodulation 12689,12329]
[SA] active: 933. $less(X37,$sum(X38,$uminus(X39))) | $less(-1,X37) | ~$less(X39,X38) (2:12:1) T [evaluation 932]
[SA] active: 12991. $less(X8,$sum(X8,4)) (7:5:1) [backward demodulation 12990,12680]
[SA] active: 947. $sum(X9,X10) = $sum($sum($uminus(X8),X9),$sum(X8,X10)) (2:12:1) T [superposition 266,288]
[SA] active: 13093. $less($sum(-3,sK0),alength0) (8:5:1) [forward demodulation 13034,16]
[SA] active: 975. ~$less(X47,$sum(X46,$sum(X45,1))) | ~$less($sum(X45,X46),X47) (2:12:1) T [superposition 27,266]
[SA] active: 14369. ~$less($sum(X2,4),X2) (8:5:1) [resolution 12688,241]
[SA] active: 977. $less(X52,$sum(X51,$sum(X50,1))) | $less($sum(X50,X51),X52) (2:12:1) T [superposition 25,266]
[SA] active: 14427. ~$less(X15,$sum(X15,-4)) (8:5:1) [evaluation 14423]
[SA] active: 1005. $sum(X10,X9) = $sum($sum($uminus(X8),X9),$sum(X10,X8)) (2:12:1) T [superposition 274,288]
[SA] active: 14428. ~$less(X2,$sum(-4,X2)) (8:5:1) [evaluation 14416]
[SA] active: 1054. ~$less(X42,$sum(X41,$sum(1,X40))) | ~$less($sum(X40,X41),X42) (2:12:1) T [superposition 160,274]
[SA] active: 15226. $less($sum(X15,-4),X15) (8:5:1) [evaluation 15222]
[SA] active: 1055. $less(X45,$sum(X44,$sum(1,X43))) | $less($sum(X43,X44),X45) (2:12:1) T [superposition 149,274]
[SA] active: 15227. $less($sum(-4,X2),X2) (8:5:1) [evaluation 15215]
[SA] active: 1061. $sum(X62,X63) = $sum(X61,$sum(X63,$sum($uminus(X61),X62))) (2:12:1) T [superposition 288,274]
[SA] active: 15661. ~$less(alength0,$sum(-3,sK0)) (9:5:1) [resolution 13093,2395]
[SA] active: 1080. $sum(X62,X63) = $sum(X61,$sum(X62,$sum(X63,$uminus(X61)))) (2:12:1) T [superposition 288,274]
[SA] active: 16758. $less($sum(-4,sK0),alength0) (9:5:1) [forward demodulation 16697,16]
[SA] active: 1094. $sum(X10,X9) = $sum(X8,$sum($uminus(X8),$sum(X9,X10))) (2:12:1) T [forward demodulation 1026,17]
[SA] active: 17674. ~$less(alength0,$sum(-4,sK0)) (10:5:1) [resolution 16758,2395]
[SA] active: 1282. ~$less($sum(X13,X14),$sum(X15,$sum(X14,1))) | ~$less(X15,X13) (2:12:1) T [resolution 281,24]
[SA] active: 6985. ~$less(2,sK0) | $less(0,alength0) (5:6:1) [resolution 3679,2395]
[SA] active: 1284. ~$less($sum(X19,X20),$sum(X21,X19)) | ~$less(X21,$sum(X20,1)) (2:12:1) T [resolution 281,248]
[SA] active: 7026. ~$less(3,X0) | $less(1,X0) (5:6:1) [resolution 3696,2395]
[SA] active: 1291. ~$less(X1,$sum(X2,$sum(1,X0))) | ~$less($sum(X2,X0),X1) (2:12:1) T [superposition 281,16]
[SA] active: 7036. ~$less(X1,1) | $less(X1,3) (5:6:1) [resolution 3696,2395]
[SA] active: 1307. ~$less($sum(X16,X17),$sum(X17,$sum(1,X18))) | ~$less(X18,X16) (2:12:1) T [forward demodulation 1283,17]
[SA] active: 7285. ~$less(X0,-2) | $less(X0,1) (5:6:1) [resolution 6345,2395]
[SA] active: 1359. $less(X1,$sum(X2,$sum(1,X0))) | $less($sum(X2,X0),X1) (2:12:1) T [superposition 283,16]
[SA] active: 7323. $less(-2,X0) | $less(X0,2) (5:6:2) [resolution 6345,3537]
[SA] active: 1505. ~$less(X27,X25) | $less(0,$sum(X26,$uminus(X27))) | ~$less(X25,X26) (2:12:1) T [resolution 238,240]
[SA] active: 7335. ~$less(1,X1) | $less(-2,X1) (5:6:1) [resolution 6345,2395]
[SA] active: 1532. $less(X49,$sum(X50,X48)) | ~$less($uminus(X48),X50) | ~$less(X49,0) (2:12:1) T [superposition 238,109]
[SA] active: 7346. $less(-3,X10) | $less(X10,0) (5:6:2) [evaluation 7345]
[SA] active: 1827. ~$less($sum(X9,X10),$sum(X8,X10)) | ~$less(X8,$sum(1,X9)) (2:12:1) T [resolution 276,160]
[SA] active: 7876. ~$less(0,X23) | $less(-3,X23) (5:6:1) T [evaluation 7875]
[SA] active: 1910. ~$less(X14,$sum(X12,$sum(X13,X14))) | ~$less($sum(X12,X13),1) (2:12:1) T [resolution 277,160]
[SA] active: 7880. $less(-3,X21) | $less(X21,1) (5:6:2) T [evaluation 7879]
[SA] active: 886. $less(alength0,$sum(X27,1)) | $less(X27,i('$n')) | alength0 = i('$n') (2:13:1) [resolution 194,204]
[SA] active: 8422. ~$less(-1,X5) | $less(-3,X5) (5:6:1) T [resolution 6507,2395]
[SA] active: 1151. $sum(i(X1),i(s(X0))) = $sum(i(s(X1)),i(X0)) (2:13:1) [superposition 268,127]
[SA] active: 8432. ~$less(X2,-3) | $less(X2,-1) (5:6:1) T [resolution 6507,2395]
[SA] active: 1293. ~$less(X7,$sum(X8,i(s(X6)))) | ~$less($sum(X8,i(X6)),X7) (2:13:1) [superposition 281,75]
[SA] active: 8474. ~$less(X15,-3) | $less(X15,-2) (5:6:1) T [evaluation 8473]
[SA] active: 1361. $less(X7,$sum(X8,i(s(X6)))) | $less($sum(X8,i(X6)),X7) (2:13:1) [superposition 283,75]
[SA] active: 8488. $less(-4,X14) | $less(X14,-2) (5:6:2) T [evaluation 8487]
[SA] active: 1460. $less($sum(X21,1),alength0) | ~$less(X21,sK0) | $sum(X21,1) = sK0 (2:13:1) [resolution 207,199]
[SA] active: 9041. ~$less(-2,X6) | $less(-3,X6) (5:6:1) T [resolution 6624,2395]
[SA] active: 1599. $less($uminus(i(s(X26))),$sum(X27,-1)) | ~$less($uminus(i(X26)),X27) (2:13:1) [evaluation 1579]
[SA] active: 9053. $less(-4,X15) | $less(X15,-3) (5:6:2) T [evaluation 9052]
[SA] active: 1600. $less($uminus(i(s(X21))),$sum(X22,$uminus(i(X21)))) | ~$less(-1,X22) (2:13:1) [evaluation 1577]
[SA] active: 3808. ~$less(3,sK0) | $less(2,alength0) (6:6:1) [resolution 2395,729]
[SA] active: 1669. $less($sum(X27,-1),$uminus(i(s(X26)))) | ~$less(X27,$uminus(i(X26))) (2:13:1) [evaluation 1660]
[SA] active: 9669. $less(1,X6) | $less(X6,4) (6:6:2) [evaluation 9668]
[SA] active: 1670. $less($sum(X22,$uminus(i(X21))),$uminus(i(s(X21)))) | ~$less(X22,-1) (2:13:1) [evaluation 1658]
[SA] active: 10477. ~$less(3,sK0) | $less(1,alength0) (6:6:1) [resolution 6189,2395]
[SA] active: 404. $less($sum(1,X4),X5) | ~$less(X4,X5) | $sum(1,X4) = X5 (2:13:1) T [resolution 160,23]
[SA] active: 10761. ~$less(4,X2) | $less(2,X2) (6:6:1) [resolution 6210,2395]
[SA] active: 521. $less(0,$sum(X5,$sum(X6,$uminus(X7)))) | ~$less(X7,$sum(X5,X6)) (2:13:1) T [superposition 240,17]
[SA] active: 10771. ~$less(X1,3) | $less(X1,4) (6:6:1) [resolution 6210,512]
[SA] active: 559. $less($sum(X5,$sum(X6,$uminus(X7))),0) | ~$less($sum(X5,X6),X7) (2:13:1) T [superposition 247,17]
[SA] active: 10774. ~$less(X4,2) | $less(X4,4) (6:6:1) [resolution 6210,2395]
[SA] active: 606. $sum($sum(X18,$uminus(X19)),$sum(X19,$sum($uminus(X18),X20))) = X20 (2:13:1) T [forward demodulation 598,17]
[SA] active: 10796. $less(sK0,4) | $less(1,alength0) (6:6:1) [resolution 6210,3807]
[SA] active: 664. $sum($sum($uminus(X23),X24),$sum($uminus(X24),$sum(X23,X25))) = X25 (2:13:1) T [forward demodulation 647,17]
[SA] active: 10797. $less(sK0,4) | $less(2,alength0) (6:6:1) [resolution 6210,199]
[SA] active: 700. $less($sum(X19,$uminus(X18)),$uminus($sum(X17,X18))) | ~$less(X19,$uminus(X17)) (2:13:1) T [superposition 24,220]
[SA] active: 10834. ~$less(4,X9) | $less(3,X9) (6:6:1) [evaluation 10833]
[SA] active: 701. $less($uminus($sum(X20,X21)),$sum(X22,$uminus(X21))) | ~$less($uminus(X20),X22) (2:13:1) T [superposition 24,220]
[SA] active: 10915. ~$less(X5,-1) | $less(X5,2) (6:6:1) [resolution 6746,2395]
[SA] active: 746. ~$less(X7,$sum(X8,X5)) | ~$less(X5,X6) | $less(X7,$sum(X6,X8)) (2:13:1) T [resolution 241,22]
[SA] active: 10953. $less(-1,X0) | $less(X0,3) (6:6:2) [resolution 6746,3538]
[SA] active: 767. $less($sum($uminus(X24),X25),$uminus($sum(X23,X24))) | ~$less(X25,$uminus(X23)) (2:13:1) T [superposition 241,220]
[SA] active: 10973. ~$less(2,X3) | $less(-1,X3) (6:6:1) [resolution 6746,2395]
[SA] active: 768. $less($sum($uminus(X27),X28),$uminus($sum(X27,X26))) | ~$less(X28,$uminus(X26)) (2:13:1) T [superposition 241,19]
[SA] active: 11642. $less(X0,3) | $less(-2,X0) (6:6:2) [resolution 6866,6359]
[SA] active: 790. ~$less(X7,$sum(X5,X8)) | ~$less(X5,X6) | $less(X7,$sum(X8,X6)) (2:13:1) T [resolution 248,22]
[SA] active: 11648. ~$less(X7,0) | $less(X7,3) (6:6:1) [resolution 6866,2395]
[SA] active: 800. $less($uminus($sum(X23,X24)),$sum($uminus(X24),X25)) | ~$less($uminus(X23),X25) (2:13:1) T [superposition 248,220]
[SA] active: 11709. ~$less(3,X4) | $less(0,X4) (6:6:1) [resolution 6866,2395]
[SA] active: 801. $less($uminus($sum(X27,X26)),$sum($uminus(X27),X28)) | ~$less($uminus(X26),X28) (2:13:1) T [superposition 248,19]
[SA] active: 11843. $less(0,X0) | $less(X0,4) (6:6:2) [resolution 6877,6210]
[SA] active: 880. $less(X13,$sum(X14,1)) | $less(X14,$sum(X15,1)) | $less(X15,X13) (2:13:3) T [resolution 194,25]
[SA] active: 12207. ~$less(3,sK0) | $less(0,alength0) (6:6:1) [resolution 6984,2395]
[SA] active: 883. $less(X20,$sum(X21,1)) | $less(X21,$sum(1,X22)) | $less(X22,X20) (2:13:3) T [resolution 194,149]
[SA] active: 18240. $less(sK0,4) | $less(0,alength0) (6:6:1) [resolution 6985,6210]
[SA] active: 981. $less($sum(X62,$sum(X61,$uminus(X63))),0) | ~$less($sum(X61,X62),X63) (2:13:1) T [superposition 247,266]
[SA] active: 19238. ~$less(X12,-2) | $less(X12,2) (6:6:1) [resolution 7323,2395]
[SA] active: 982. $less(0,$sum(X65,$sum(X64,$uminus(X66)))) | ~$less(X66,$sum(X64,X65)) (2:13:1) T [superposition 240,266]
[SA] active: 19318. ~$less(2,X13) | $less(-2,X13) (6:6:1) [resolution 7323,2395]
[SA] active: 1466. $less($sum(1,X0),X1) | $sum(X0,1) = X1 | ~$less(X0,X1) (2:13:1) T [superposition 207,16]
[SA] active: 19342. ~$less(X14,1) | $less(X14,4) (6:6:1) [evaluation 19341]
[SA] active: 1508. ~$less(X36,X34) | $less($sum(X36,X37),$sum(X35,X37)) | ~$less(X34,X35) (2:13:1) T [resolution 238,24]
[SA] active: 19750. $less(-3,X10) | $less(X10,2) (6:6:2) [evaluation 19749]
[SA] active: 1509. ~$less(X41,X38) | $less($sum(X40,X41),$sum(X39,X40)) | ~$less(X38,X39) (2:13:1) T [resolution 238,241]
[SA] active: 19917. ~$less(X14,-3) | $less(X14,0) (6:6:1) [resolution 7346,2395]
[SA] active: 1510. $less($sum(X44,X42),$sum(X43,X45)) | ~$less(X42,X43) | ~$less(X44,X45) (2:13:1) T [resolution 238,248]
[SA] active: 20018. $less(-4,X22) | $less(X22,-1) (6:6:2) [evaluation 20017]
[SA] active: 1603. $less($uminus(X9),$sum(X10,$uminus(X8))) | ~$less($sum($uminus(X9),X8),X10) (2:13:1) T [forward demodulation 1573,217]
[SA] active: 20202. $less(-3,X2) | $less(X2,3) (6:6:2) [resolution 7876,6866]
[SA] active: 1679. $less($sum(X10,$uminus(X8)),$uminus(X9)) | ~$less(X10,$sum($uminus(X9),X8)) (2:13:1) T [forward demodulation 1654,217]
[SA] active: 19889. ~$less(1,X2) | ~$less(X2,0) (6:6:1) T [evaluation 19759]
[SA] active: 1742. $sum($uminus(X6),X7) = $sum(X5,$sum($uminus($sum(X6,X5)),X7)) (2:13:1) T [superposition 288,271]
[SA] active: 20286. ~$less(0,X30) | $less(-4,X30) (6:6:1) T [evaluation 20285]
[SA] active: 1783. $sum($uminus(X22),X21) = $sum($uminus($sum(X20,X22)),$sum(X20,X21)) (2:13:1) T [forward demodulation 1730,26]
[SA] active: 20290. $less(-4,X28) | $less(X28,1) (6:6:2) T [evaluation 20289]
[SA] active: 1843. $less(0,$sum(X49,$sum(X50,X48))) | ~$less($uminus(X48),$sum(X49,X50)) (2:13:1) T [superposition 276,109]
[SA] active: 20292. $less(-4,X23) | $less(X23,0) (6:6:2) T [evaluation 20291]
[SA] active: 1952. $less($sum(X49,$sum(X50,X48)),0) | ~$less($sum(X49,X50),$uminus(X48)) (2:13:1) T [superposition 277,109]
[SA] active: 20510. ~$less(X14,-3) | $less(X14,1) (6:6:1) T [resolution 7880,2395]
[SA] active: 482. $less(i(sK2(X3)),i(s('$n'))) | a(zero,X3) = a('$n',X3) (2:14:1) [resolution 133,73]
[SA] active: 20733. ~$less(-1,X17) | $less(-4,X17) (6:6:1) T [evaluation 20732]
[SA] active: 500. $less(i(s(X2)),X3) | ~$less(i(X2),X3) | i(s(X2)) = X3 (2:14:1) [resolution 163,23]
[SA] active: 21207. ~$less(X14,-4) | $less(X14,-2) (6:6:1) T [resolution 8488,2395]
[SA] active: 818. ~$less(X5,i(s(X3))) | ~$less(i(X3),X4) | $less(X5,$sum(X4,1)) (2:14:1) [resolution 245,22]
[SA] active: 21302. ~$less(-2,X19) | $less(-4,X19) (6:6:1) T [resolution 8488,2395]
[SA] active: 819. $less(i(X0),$sum(X1,1)) | ~$less(i(p(X0)),X1) | zero = X0 (2:14:1) [superposition 245,64]
[SA] active: 21313. $less(-5,X22) | $less(X22,-3) (6:6:2) T [evaluation 21312]
[SA] active: 840. ~$less(X4,$sum(X2,1)) | ~$less(X2,i(X3)) | $less(X4,i(s(X3))) (2:14:1) [resolution 252,22]
[SA] active: 21328. ~$less(X22,-4) | $less(X22,-3) (6:6:1) T [evaluation 21327]
[SA] active: 846. $less($sum(X1,1),i(X0)) | ~$less(X1,i(p(X0))) | zero = X0 (2:14:1) [superposition 252,64]
[SA] active: 21615. ~$less(-3,X18) | $less(-4,X18) (6:6:1) T [resolution 9053,2395]
[SA] active: 877. $less(X6,$sum(X7,1)) | $less(X7,i(s(X8))) | $less(i(X8),X6) (2:14:3) [resolution 194,152]
[SA] active: 21626. $less(-5,X21) | $less(X21,-4) (6:6:2) T [evaluation 21625]
[SA] active: 1426. ~'Sub'(zero,'$n') | v0 = a('$n',0) | $less(i(sK1(zero,0)),alength0) (2:14:1) [resolution 326,91]
[SA] active: 2875. ~$less(sK0,X48) | $less(-2,X48) (7:6:1) [evaluation 2874]
[SA] active: 1504. ~$less(i(X24),X22) | $less(i(s(X24)),$sum(X23,1)) | ~$less(X22,X23) (2:14:1) [resolution 238,245]
[SA] active: 2889. ~$less(alength0,X27) | $less(-2,X27) (7:6:1) [evaluation 2888]
[SA] active: 1529. $less(X36,$sum(X37,i(X35))) | ~$less(1,X37) | ~$less(X36,i(s(X35))) (2:14:1) [superposition 238,127]
[SA] active: 19888. ~$less(2,X3) | ~$less(X3,0) (7:6:1) [evaluation 19760]
[SA] active: 596. $less(0,$sum(X14,$sum(X13,$uminus(X12)))) | ~$less($sum(X12,$uminus(X13)),X14) (2:14:1) T [superposition 240,216]
[SA] active: 21762. $less(X0,4) | $less(-2,X0) (7:6:2) [resolution 9669,7335]
[SA] active: 597. $less($sum(X17,$sum(X16,$uminus(X15))),0) | ~$less(X17,$sum(X15,$uminus(X16))) (2:14:1) T [superposition 247,216]
[SA] active: 21763. $less(X1,4) | $less(-1,X1) (7:6:2) [resolution 9669,6757]
[SA] active: 601. $sum(X2,$uminus($sum(X4,X3))) = $uminus($sum(X4,$sum(X3,$uminus(X2)))) (2:14:1) T [forward demodulation 600,19]
[SA] active: 21863. ~$less(4,X14) | $less(1,X14) (7:6:1) [resolution 9669,2395]
[SA] active: 605. $sum(X4,$uminus($sum(X5,X3))) = $sum($sum(X4,$uminus(X3)),$uminus(X5)) (2:14:1) T [forward demodulation 592,601]
[SA] active: 22171. $less(1,X12) | $less(X12,5) (7:6:2) [evaluation 22170]
[SA] active: 645. $less(0,$sum(X19,$sum($uminus(X18),X17))) | ~$less($sum($uminus(X17),X18),X19) (2:14:1) T [superposition 240,217]
[SA] active: 22363. $less(X8,5) | $less(3,X8) (7:6:2) [evaluation 22362]
[SA] active: 646. $less($sum(X22,$sum($uminus(X21),X20)),0) | ~$less(X22,$sum($uminus(X20),X21)) (2:14:1) T [superposition 247,217]
[SA] active: 22546. $less(X9,5) | $less(2,X9) (7:6:2) [evaluation 22545]
[SA] active: 699. $sum($uminus($sum(X14,X15)),X16) = $sum($uminus(X14),$sum($uminus(X15),X16)) (2:14:1) T [superposition 17,220]
[SA] active: 22936. ~$less(4,sK0) | $less(1,alength0) (7:6:1) [resolution 10796,2395]
[SA] active: 706. $sum(X3,$uminus($sum(X4,X2))) = $uminus($sum(X2,$sum($uminus(X3),X4))) (2:14:1) T [forward demodulation 705,19]
[SA] active: 23367. ~$less(4,sK0) | $less(2,alength0) (7:6:1) [resolution 10797,2395]
[SA] active: 707. $sum(X1,$uminus($sum(X2,X0))) = $sum($uminus(X2),$sum(X1,$uminus(X0))) (2:14:1) T [backward demodulation 706,604]
[SA] active: 23968. ~$less(X12,-1) | $less(X12,3) (7:6:1) [resolution 10953,2395]
[SA] active: 709. $sum($uminus($sum(X6,X7)),X5) = $sum($sum($uminus(X6),X5),$uminus(X7)) (2:14:1) T [forward demodulation 708,217]
[SA] active: 24060. ~$less(3,X9) | $less(-1,X9) (7:6:1) [resolution 10953,2395]
[SA] active: 711. $sum(X5,$uminus($sum(X6,X7))) = $sum($uminus(X7),$sum($uminus(X6),X5)) (2:14:1) T [forward demodulation 685,706]
[SA] active: 24560. ~$less(3,X9) | $less(-2,X9) (7:6:1) [resolution 11642,2395]
[SA] active: 754. $less($sum(X17,$sum(X18,X19)),$sum(X20,$sum(X17,X18))) | ~$less(X19,X20) (2:14:1) T [superposition 241,17]
[SA] active: 24574. ~$less(X9,-2) | $less(X9,3) (7:6:1) [resolution 11642,2395]
[SA] active: 765. $less($sum(X19,X20),$sum(X17,$sum(X18,X19))) | ~$less(X20,$sum(X17,X18)) (2:14:1) T [superposition 241,17]
[SA] active: 24637. ~$less(X11,1) | $less(X11,5) (7:6:1) [evaluation 24636]
[SA] active: 798. $less($sum(X17,$sum(X18,X19)),$sum(X19,X20)) | ~$less($sum(X17,X18),X20) (2:14:1) T [superposition 248,17]
[SA] active: 25384. $less(X0,4) | $less(-3,X0) (7:6:2) [resolution 11843,7876]
[SA] active: 809. $less($sum(X20,$sum(X17,X18)),$sum(X17,$sum(X18,X19))) | ~$less(X20,X19) (2:14:1) T [superposition 248,17]
[SA] active: 25393. ~$less(X12,0) | $less(X12,4) (7:6:1) [resolution 11843,2395]
[SA] active: 870. X0 != X1 | p(X0) = p(X1) | zero = X1 | zero = X0 (2:14:1) [superposition 181,64]
[SA] active: 25488. ~$less(4,X8) | $less(0,X8) (7:6:1) [resolution 11843,2395]
[SA] active: 875. $less(X0,$sum(X1,1)) | $less(X2,X0) | $less(X1,X2) | X0 = X2 (2:14:3) T [resolution 194,23]
[SA] active: 26337. ~$less(4,sK0) | $less(0,alength0) (7:6:1) [resolution 18240,2395]
[SA] active: 899. $less($sum(X51,$uminus(X52)),$sum(X53,1)) | $less(X53,0) | ~$less(X51,X52) (2:14:1) T [resolution 194,247]
[SA] active: 26712. ~$less(2,X16) | $less(-3,X16) (7:6:1) [evaluation 26711]
[SA] active: 966. $less($sum(X16,X15),$sum(X14,$sum(X13,X15))) | ~$less(X16,$sum(X13,X14)) (2:14:1) T [superposition 24,266]
[SA] active: 27173. ~$less(X8,-3) | $less(X8,2) (7:6:1) [resolution 19750,2395]
[SA] active: 967. $less($sum(X18,$sum(X17,X19)),$sum(X20,X19)) | ~$less($sum(X17,X18),X20) (2:14:1) T [superposition 24,266]
[SA] active: 27709. ~$less(X8,-4) | $less(X8,-1) (7:6:1) [resolution 20018,2395]
[SA] active: 968. $less($sum(X22,$sum(X21,X23)),$sum(X24,$sum(X21,X22))) | ~$less(X23,X24) (2:14:1) T [superposition 241,266]
[SA] active: 27826. $less(-5,X16) | $less(X16,-2) (7:6:2) [evaluation 27825]
[SA] active: 969. $less($sum(X27,X28),$sum(X26,$sum(X25,X27))) | ~$less(X28,$sum(X25,X26)) (2:14:1) T [superposition 241,266]
[SA] active: 27925. ~$less(X8,-3) | $less(X8,3) (7:6:1) [resolution 20202,2395]
[SA] active: 970. $less($sum(X32,$sum(X29,X30)),$sum(X30,$sum(X29,X31))) | ~$less(X32,X31) (2:14:1) T [superposition 248,266]
[SA] active: 28025. ~$less(3,X9) | $less(-3,X9) (7:6:1) [resolution 20202,2395]
[SA] active: 971. $less($sum(X34,$sum(X33,X35)),$sum(X35,X36)) | ~$less($sum(X33,X34),X36) (2:14:1) T [superposition 248,266]
[SA] active: 28049. ~$less(X10,1) | $less(X10,6) (7:6:1) [evaluation 28048]
[SA] active: 980. $sum(X60,$uminus($sum(X58,X59))) = $uminus($sum(X59,$sum(X58,$uminus(X60)))) (2:14:1) T [superposition 216,266]
[SA] active: 28733. $less(-4,X0) | $less(X0,4) (7:6:2) [resolution 20286,11843]
[SA] active: 1014. $sum(X37,$uminus($sum(X35,X36))) = $sum($uminus(X36),$sum(X37,$uminus(X35))) (2:14:1) T [superposition 274,220]
[SA] active: 28734. $less(-4,X1) | $less(X1,3) (7:6:2) [resolution 20286,6866]
[SA] active: 1035. $sum(X37,$uminus($sum(X35,X36))) = $sum($uminus(X35),$sum($uminus(X36),X37)) (2:14:1) T [superposition 274,220]
[SA] active: 28201. ~$less(1,X17) | ~$less(X17,-1) (7:6:1) T [evaluation 28200]
[SA] active: 1047. $less($sum(X15,$sum(X13,X14)),$sum(X14,$sum(X12,X13))) | ~$less(X15,X12) (2:14:1) T [superposition 24,274]
[SA] active: 28207. ~$less(0,X7) | ~$less(X7,-1) (7:6:1) T [evaluation 28206]
[SA] active: 1048. $less($sum(X18,$sum(X16,X17)),$sum(X19,$sum(X17,X18))) | ~$less(X16,X19) (2:14:1) T [superposition 24,274]
[SA] active: 28833. ~$less(0,X37) | $less(-5,X37) (7:6:1) T [evaluation 28832]
[SA] active: 1049. $less($sum(X22,$sum(X20,X21)),$sum(X23,X20)) | ~$less($sum(X21,X22),X23) (2:14:1) T [superposition 241,274]
[SA] active: 28837. $less(-5,X29) | $less(X29,0) (7:6:2) T [evaluation 28836]
[SA] active: 1051. $less($sum(X31,X28),$sum(X30,$sum(X28,X29))) | ~$less(X31,$sum(X29,X30)) (2:14:1) T [superposition 248,274]
[SA] active: 29592. ~$less(X8,-4) | $less(X8,1) (7:6:1) T [resolution 20290,2395]
[SA] active: 1060. $sum($uminus($sum(X59,X60)),X58) = $uminus($sum(X60,$sum($uminus(X58),X59))) (2:14:1) T [superposition 217,274]
[SA] active: 29719. ~$less(X20,-4) | $less(X20,0) (7:6:1) T [evaluation 29718]
[SA] active: 1079. $sum($uminus($sum(X59,X60)),X58) = $uminus($sum(X59,$sum(X60,$uminus(X58)))) (2:14:1) T [superposition 217,274]
[SA] active: 30054. $less(-5,X11) | $less(X11,-1) (7:6:2) T [evaluation 30053]
[SA] active: 1101. $less($sum(X25,$sum(X26,X27)),$sum(X26,$sum(X24,X25))) | ~$less(X27,X24) (2:14:1) T [forward demodulation 1050,17]
[SA] active: 30433. ~$less(-1,X19) | $less(-5,X19) (7:6:1) T [evaluation 30432]
[SA] active: 1102. $less($sum(X34,$sum(X32,X33)),$sum(X33,$sum(X34,X35))) | ~$less(X32,X35) (2:14:1) T [forward demodulation 1052,17]
[SA] active: 30718. ~$less(-2,X20) | $less(-5,X20) (7:6:1) T [evaluation 30717]
[SA] active: 1112. $less($sum(X25,$sum(X26,X27)),$sum(X25,$sum(X26,X24))) | ~$less(X27,X24) (2:14:1) T [forward demodulation 1069,17]
[SA] active: 30774. ~$less(X2,-5) | $less(X2,-3) (7:6:1) T [resolution 21313,2395]
[SA] active: 1523. $less(X13,$sum(X14,$sum($uminus(X11),X12))) | ~$less(X11,X14) | ~$less(X13,X12) (2:14:1) T [superposition 238,288]
[SA] active: 30881. ~$less(-3,X8) | $less(-5,X8) (7:6:1) T [resolution 21313,2395]
[SA] active: 1595. $less($uminus($sum(X17,X18)),$uminus($sum(X16,X17))) | ~$less($uminus(X18),$uminus(X16)) (2:14:1) T [superposition 244,220]
[SA] active: 30892. $less(-6,X11) | $less(X11,-4) (7:6:2) T [evaluation 30891]
[SA] active: 1596. $less($uminus($sum(X20,X21)),$uminus($sum(X20,X19))) | ~$less($uminus(X21),$uminus(X19)) (2:14:1) T [superposition 244,19]
[SA] active: 30908. ~$less(X8,-5) | $less(X8,-4) (7:6:1) T [evaluation 30907]
[SA] active: 1648. $less($uminus($sum(X16,X17)),$uminus($sum(X17,X18))) | ~$less($uminus(X16),$uminus(X18)) (2:14:1) T [superposition 251,220]
[SA] active: 31540. ~$less(-4,X6) | $less(-5,X6) (7:6:1) T [resolution 21626,2395]
[SA] active: 1727. $sum($uminus($sum(X12,X14)),X13) = $sum($uminus(X14),$sum(X13,$uminus(X12))) (2:14:1) T [superposition 271,16]
[SA] active: 31551. $less(-6,X9) | $less(X9,-5) (7:6:2) T [evaluation 31550]
[SA] active: 1736. $sum($uminus($sum(X8,X7)),X9) = $sum($uminus(X8),$sum(X9,$uminus(X7))) (2:14:1) T [superposition 271,274]
[SA] active: 13921. ~$less(1,X0) | $less(-3,X0) (8:6:1) [resolution 12652,1498]
[SA] active: 1780. $sum($uminus(X4),$uminus($sum(X2,X3))) = $uminus($sum(X3,$sum(X2,X4))) (2:14:1) T [forward demodulation 1724,19]
[SA] active: 19887. ~$less(3,X4) | ~$less(X4,0) (8:6:1) [evaluation 19761]
[SA] active: 1781. $sum($uminus(X7),$uminus($sum(X6,X5))) = $uminus($sum(X6,$sum(X5,X7))) (2:14:1) T [forward demodulation 1725,19]
[SA] active: 32061. ~$less(2,X15) | ~$less(X15,-1) (8:6:1) [evaluation 32060]
[SA] active: 1847. $less($sum(X2,X1),$sum(X3,$sum(X1,X0))) | ~$less(X2,$sum(X3,X0)) (2:14:1) T [superposition 276,16]
[SA] active: 32067. ~$less(0,X6) | ~$less(X6,-2) (8:6:1) [evaluation 32066]
[SA] active: 1914. $less($sum(X2,$sum(X1,X0)),$sum(X3,X1)) | ~$less($sum(X2,X0),X3) (2:14:1) T [superposition 277,16]
[SA] active: 32246. ~$less(4,X2) | $less(-2,X2) (8:6:1) [resolution 21762,2395]
[SA] active: 327. 'Sub'(sK1(s(zero),1),'$n') | v0 = a('$n',1) | ~'Sub'(s(zero),'$n') (2:15:1) [forward demodulation 324,131]
[SA] active: 32258. ~$less(X4,-2) | $less(X4,4) (8:6:1) [resolution 21762,2395]
[SA] active: 826. $less(i(s(X8)),$sum(X6,$sum(X7,1))) | ~$less(i(X8),$sum(X6,X7)) (2:15:1) [superposition 245,17]
[SA] active: 32333. ~$less(4,X3) | $less(-1,X3) (8:6:1) [evaluation 32332]
[SA] active: 844. $less($sum(X6,$sum(X7,1)),i(s(X8))) | ~$less($sum(X6,X7),i(X8)) (2:15:1) [superposition 252,17]
[SA] active: 32515. ~$less(X7,-1) | $less(X7,4) (8:6:1) [resolution 21763,2395]
[SA] active: 973. $less($sum(X40,$sum(X39,1)),i(s(X41))) | ~$less($sum(X39,X40),i(X41)) (2:15:1) [superposition 252,266]
[SA] active: 33043. $less(0,X14) | $less(X14,5) (8:6:2) [evaluation 33042]
[SA] active: 974. $less(i(s(X44)),$sum(X43,$sum(X42,1))) | ~$less(i(X44),$sum(X42,X43)) (2:15:1) [superposition 245,266]
[SA] active: 33407. ~$less(X0,0) | $less(X0,5) (8:6:1) [resolution 22171,19889]
[SA] active: 1140. ~$less(X3,i(sK2(X2))) | a(zero,X2) = a('$n',X2) | $less(X3,alength0) (2:15:1) [resolution 257,22]
[SA] active: 33408. $less(X1,5) | $less(-2,X1) (8:6:2) [resolution 22171,7335]
[SA] active: 1168. $less($sum(X13,$sum(1,X12)),$sum(i(s(X11)),X12)) | ~$less(X13,i(X11)) (2:15:1) [superposition 24,268]
[SA] active: 33409. $less(X2,5) | $less(-1,X2) (8:6:2) [resolution 22171,6757]
[SA] active: 1169. $less($sum(i(s(X14)),X15),$sum(X16,$sum(1,X15))) | ~$less(i(X14),X16) (2:15:1) [superposition 24,268]
[SA] active: 33411. ~$less(X4,2) | $less(X4,5) (8:6:1) [resolution 22171,511]
[SA] active: 1170. $less($sum(i(s(X17)),X18),$sum(X19,i(X17))) | ~$less($sum(1,X18),X19) (2:15:1) [superposition 241,268]
[SA] active: 33446. $less(sK0,5) | $less(0,alength0) (8:6:1) [resolution 22171,3805]
[SA] active: 1172. $less($sum(X25,i(X23)),$sum(i(s(X23)),X24)) | ~$less(X25,$sum(1,X24)) (2:15:1) [superposition 248,268]
[SA] active: 33447. $less(sK0,5) | $less(1,alength0) (8:6:1) [resolution 22171,199]
[SA] active: 1190. $less($sum(1,$sum(X21,X22)),$sum(i(s(X20)),X21)) | ~$less(X22,i(X20)) (2:15:1) [forward demodulation 1171,17]
[SA] active: 33519. ~$less(5,X2) | $less(1,X2) (8:6:1) [resolution 22171,2395]
[SA] active: 1191. $less($sum(i(s(X26)),X27),$sum(1,$sum(X27,X28))) | ~$less(i(X26),X28) (2:15:1) [forward demodulation 1173,17]
[SA] active: 33534. ~$less(5,X3) | $less(2,X3) (8:6:1) [evaluation 33533]
[SA] active: 1230. $less(X31,i('$n')) | i('$n') = X31 | $less(alength0,X31) | alength0 = i('$n') (2:15:1) [resolution 203,204]
[SA] active: 33773. ~$less(5,X2) | $less(3,X2) (8:6:1) [resolution 22363,2395]
[SA] active: 1428. ~'Sub'(zero,'$n') | v0 = a('$n',0) | $less(i(sK1(zero,0)),i('$n')) (2:15:1) [resolution 326,78]
[SA] active: 33786. ~$less(X5,3) | $less(X5,5) (8:6:1) [resolution 22363,2395]
[SA] active: 1840. $less(i(s(X35)),$sum(X36,$sum(X37,i(X35)))) | ~$less(1,$sum(X36,X37)) (2:15:1) [superposition 276,127]
[SA] active: 33820. $less(sK0,5) | $less(2,alength0) (8:6:1) [resolution 22363,3808]
[SA] active: 1854. $less($sum(X29,1),$sum(X30,i(s(X28)))) | ~$less(X29,$sum(X30,i(X28))) (2:15:1) [superposition 276,75]
[SA] active: 33821. $less(sK0,5) | $less(3,alength0) (8:6:1) [resolution 22363,199]
[SA] active: 1856. $less($sum(X36,i(X35)),$sum(X37,i(s(X35)))) | ~$less(X36,$sum(X37,1)) (2:15:1) [superposition 276,127]
[SA] active: 33864. ~$less(5,X3) | $less(4,X3) (8:6:1) [evaluation 33863]
[SA] active: 1921. $less($sum(X29,i(s(X28))),$sum(X30,1)) | ~$less($sum(X29,i(X28)),X30) (2:15:1) [superposition 277,75]
[SA] active: 37564. $less(X0,5) | $less(-4,X0) (8:6:2) [resolution 24637,20290]
[SA] active: 1923. $less($sum(X36,i(s(X35))),$sum(X37,i(X35))) | ~$less($sum(X36,1),X37) (2:15:1) [superposition 277,127]
[SA] active: 37565. $less(X1,5) | $less(-3,X1) (8:6:2) [resolution 24637,7880]
[SA] active: 1949. $less($sum(X36,$sum(X37,i(X35))),i(s(X35))) | ~$less($sum(X36,X37),1) (2:15:1) [superposition 277,127]
[SA] active: 37639. $less(X8,6) | $less(1,X8) (8:6:2) [evaluation 37638]
[SA] active: 900. ~$less(X54,X57) | $less(X56,$sum(X57,X55)) | $less($sum(X54,X55),$sum(X56,1)) (2:15:1) T [resolution 194,24]
[SA] active: 37986. ~$less(4,X2) | $less(-3,X2) (8:6:1) [resolution 25384,2395]
[SA] active: 901. ~$less(X59,X61) | $less(X60,$sum(X61,X58)) | $less($sum(X58,X59),$sum(X60,1)) (2:15:1) T [resolution 194,241]
[SA] active: 37996. ~$less(X3,-3) | $less(X3,4) (8:6:1) [resolution 25384,2395]
[SA] active: 902. ~$less(X62,X65) | $less(X64,$sum(X63,X65)) | $less($sum(X62,X63),$sum(X64,1)) (2:15:1) T [resolution 194,248]
[SA] active: 38065. ~$less(X5,1) | $less(X5,7) (8:6:1) [evaluation 38064]
[SA] active: 989. $sum(X19,$sum(X17,$sum(X18,X20))) = $sum($sum(X18,$sum(X17,X19)),X20) (2:15:1) T [forward demodulation 951,17]
[SA] active: 39045. ~$less(2,X14) | $less(-4,X14) (8:6:1) [evaluation 39044]
[SA] active: 991. $sum(X23,$sum(X21,$sum(X22,X24))) = $sum(X23,$sum(X22,$sum(X21,X24))) (2:15:1) T [forward demodulation 990,17]
[SA] active: 39300. $less(X0,2) | $less(-5,X0) (8:6:2) [resolution 27173,21313]
[SA] active: 1006. $sum(X14,$sum(X12,$sum(X13,X11))) = $sum($sum(X12,X13),$sum(X14,X11)) (2:15:1) T [superposition 274,274]
[SA] active: 40113. ~$less(X2,-5) | $less(X2,-2) (8:6:1) [resolution 27826,2395]
[SA] active: 1007. $sum(X18,$sum(X17,$sum(X15,X16))) = $sum($sum(X16,X17),$sum(X18,X15)) (2:15:1) T [superposition 274,274]
[SA] active: 40184. $less(-5,X6) | $less(X6,3) (8:6:2) [resolution 27826,24574]
[SA] active: 1011. $sum(X27,$sum(X28,$sum(X25,X26))) = $sum(X28,$sum(X26,$sum(X25,X27))) (2:15:1) T [superposition 274,266]
[SA] active: 40634. ~$less(X4,0) | $less(X4,6) (8:6:1) [evaluation 40633]
[SA] active: 1012. $sum(X31,$sum(X32,$sum(X29,X30))) = $sum(X32,$sum(X29,$sum(X30,X31))) (2:15:1) T [superposition 274,17]
[SA] active: 41130. ~$less(3,X12) | $less(-4,X12) (8:6:1) [evaluation 41129]
[SA] active: 1032. $sum(X28,$sum(X26,$sum(X25,X27))) = $sum($sum(X25,X26),$sum(X27,X28)) (2:15:1) T [superposition 274,266]
[SA] active: 41136. ~$less(0,X5) | $less(-6,X5) (8:6:1) [evaluation 41135]
[SA] active: 1033. $sum(X32,$sum(X29,$sum(X30,X31))) = $sum($sum(X29,X30),$sum(X31,X32)) (2:15:1) T [superposition 274,17]
[SA] active: 41435. $less(X0,6) | $less(-4,X0) (8:6:2) [resolution 28049,20290]
[SA] active: 1041. $sum(X18,$sum(X19,$sum(X20,X21))) = $sum(X20,$sum(X21,$sum(X18,X19))) (2:15:1) T [superposition 274,17]
[SA] active: 41436. $less(X1,6) | $less(-3,X1) (8:6:2) [resolution 28049,7880]
[SA] active: 1090. $sum(X15,$sum(X14,$sum(X16,X17))) = $sum(X17,$sum(X14,$sum(X15,X16))) (2:15:1) T [forward demodulation 1020,17]
[SA] active: 41437. $less(X2,6) | $less(-2,X2) (8:6:2) [resolution 28049,6345]
[SA] active: 1095. $sum(X11,$sum(X12,$sum(X13,X14))) = $sum(X14,$sum(X12,$sum(X13,X11))) (2:15:1) T [forward demodulation 1027,17]
[SA] active: 41438. $less(X3,6) | $less(-1,X3) (8:6:2) [resolution 28049,3590]
[SA] active: 1096. $sum(X15,$sum(X16,$sum(X17,X18))) = $sum(X18,$sum(X17,$sum(X15,X16))) (2:15:1) T [forward demodulation 1028,17]
[SA] active: 41439. $less(X4,6) | $less(0,X4) (8:6:2) [resolution 28049,378]
[SA] active: 1100. $sum(X8,$sum(X9,$sum(X10,X11))) = $sum($sum(X10,$sum(X8,X9)),X11) (2:15:1) T [forward demodulation 1046,17]
[SA] active: 41464. $less(sK0,6) | $less(0,alength0) (8:6:1) [resolution 28049,393]
[SA] active: 1834. $less(X12,$sum(X13,$sum(X14,$sum($uminus(X11),X12)))) | ~$less(X11,$sum(X13,X14)) (2:15:1) T [superposition 276,288]
[SA] active: 41512. $less(X8,7) | $less(1,X8) (8:6:2) [evaluation 41511]
[SA] active: 1850. $less($sum(X13,$sum($uminus(X11),X12)),$sum(X14,X12)) | ~$less(X13,$sum(X14,X11)) (2:15:1) T [superposition 276,288]
[SA] active: 41980. ~$less(X2,-4) | $less(X2,4) (8:6:1) [resolution 28733,2395]
[SA] active: 1917. $less($sum(X13,X12),$sum(X14,$sum($uminus(X11),X12))) | ~$less($sum(X13,X11),X14) (2:15:1) T [superposition 277,288]
[SA] active: 42096. ~$less(4,X2) | $less(-4,X2) (8:6:1) [resolution 28733,2395]
[SA] active: 1943. $less($sum(X13,$sum(X14,$sum($uminus(X11),X12))),X12) | ~$less($sum(X13,X14),X11) (2:15:1) T [superposition 277,288]
[SA] active: 42110. ~$less(X8,-4) | $less(X8,3) (8:6:1) [evaluation 42109]
[SA] active: 1156. $sum(i(s(X11)),$sum(X9,X10)) = $sum(i(X11),$sum(X9,$sum(X10,1))) (2:16:1) [superposition 268,274]
[SA] active: 42114. ~$less(X4,1) | $less(X4,8) (8:6:1) [evaluation 42113]
[SA] active: 1157. $sum(i(s(X14)),$sum(X12,X13)) = $sum(i(X14),$sum(X13,$sum(1,X12))) (2:16:1) [superposition 268,274]
[SA] active: 42774. ~$less(X8,-4) | $less(X8,2) (8:6:1) [evaluation 42773]
[SA] active: 1176. $sum(X37,$sum(i(s(X35)),X36)) = $sum($sum(1,X36),$sum(X37,i(X35))) (2:16:1) [superposition 274,268]
[SA] active: 43394. ~$less(X0,-1) | $less(X0,5) (8:6:1) [resolution 28201,22171]
[SA] active: 1192. $sum(i(s(X29)),$sum(X30,X31)) = $sum($sum(1,X30),$sum(i(X29),X31)) (2:16:1) [forward demodulation 1174,17]
[SA] active: 44511. $less(-5,X0) | $less(X0,4) (8:6:2) [resolution 28833,11843]
[SA] active: 1193. $sum(X34,$sum(i(s(X32)),X33)) = $sum(i(X32),$sum(1,$sum(X33,X34))) (2:16:1) [forward demodulation 1175,17]
[SA] active: 43482. ~$less(1,X15) | ~$less(X15,-2) (8:6:1) T [evaluation 43481]
[SA] active: 1908. ~$less($sum(X5,X6),$sum(X3,$sum(X4,$sum(X6,1)))) | ~$less($sum(X3,X4),X5) (2:16:1) T [resolution 277,281]
[SA] active: 44618. $less(-6,X37) | $less(X37,0) (8:6:2) T [evaluation 44617]
[SA] active: 1580. $less($uminus($sum(X29,$sum(X28,X30))),$sum(X31,$uminus($sum(X28,X29)))) | ~$less($uminus(X30),X31) (2:17:1) T [superposition 244,266]
[SA] active: 44954. ~$less(X2,-5) | $less(X2,0) (8:6:1) T [resolution 28837,2395]
[SA] active: 1581. $less($uminus($sum(X32,$sum(X33,X34))),$sum(X35,$uminus($sum(X32,X33)))) | ~$less($uminus(X34),X35) (2:17:1) T [superposition 244,17]
[SA] active: 45097. ~$less(X8,-5) | $less(X8,-1) (8:6:1) T [evaluation 45096]
[SA] active: 1614. $less($sum(X4,$uminus($sum(X5,X3))),$sum(X6,$sum(X4,$uminus(X3)))) | ~$less($uminus(X5),X6) (2:17:1) T [forward demodulation 1613,706]
[SA] active: 47393. ~$less(-1,X18) | $less(-6,X18) (8:6:1) T [evaluation 47392]
[SA] active: 1616. $less($sum($uminus($sum(X8,X9)),X7),$sum(X10,$sum($uminus(X8),X7))) | ~$less($uminus(X9),X10) (2:17:1) T [forward demodulation 1615,217]
[SA] active: 47397. $less(-6,X12) | $less(X12,-1) (8:6:2) T [evaluation 47396]
[SA] active: 1661. $less($sum(X31,$uminus($sum(X28,X29))),$uminus($sum(X29,$sum(X28,X30)))) | ~$less(X31,$uminus(X30)) (2:17:1) T [superposition 251,266]
[SA] active: 47884. ~$less(-2,X19) | $less(-6,X19) (8:6:1) T [evaluation 47883]
[SA] active: 1662. $less($sum(X35,$uminus($sum(X32,X33))),$uminus($sum(X32,$sum(X33,X34)))) | ~$less(X35,$uminus(X34)) (2:17:1) T [superposition 251,17]
[SA] active: 47888. $less(-6,X13) | $less(X13,-2) (8:6:2) T [evaluation 47887]
[SA] active: 1675. $less($sum(X5,$sum(X4,$uminus(X3))),$sum(X4,$uminus($sum(X6,X3)))) | ~$less(X5,$uminus(X6)) (2:17:1) T [forward demodulation 1674,706]
[SA] active: 48660. ~$less(-3,X20) | $less(-6,X20) (8:6:1) T [evaluation 48659]
[SA] active: 1677. $less($sum(X9,$sum($uminus(X8),X7)),$sum($uminus($sum(X8,X10)),X7)) | ~$less(X9,$uminus(X10)) (2:17:1) T [forward demodulation 1676,217]
[SA] active: 48664. $less(-6,X14) | $less(X14,-3) (8:6:2) T [evaluation 48663]
[SA] active: 1808. $less($sum($uminus(X50),$sum(X51,X52)),$sum($uminus($sum(X50,X49)),X51)) | ~$less(X52,$uminus(X49)) (2:17:1) T [forward demodulation 1755,17]
[SA] active: 49052. ~$less(X2,-6) | $less(X2,-4) (8:6:1) T [resolution 30892,2395]
[SA] active: 1812. $less($sum($uminus($sum(X62,X61)),X63),$sum($uminus(X62),$sum(X63,X64))) | ~$less($uminus(X61),X64) (2:17:1) T [forward demodulation 1758,17]
[SA] active: 49127. $less(-6,X5) | $less(X5,1) (8:6:2) T [resolution 30892,29592]
[SA] active: 1731. $sum($uminus($sum(X23,X26)),$sum(X24,X25)) = $sum($uminus(X26),$sum(X24,$sum(X25,$uminus(X23)))) (2:18:1) T [superposition 271,274]
[SA] active: 49187. ~$less(-4,X9) | $less(-6,X9) (8:6:1) T [resolution 30892,2395]
[SA] active: 1732. $sum($uminus($sum(X27,X30)),$sum(X28,X29)) = $sum($uminus(X30),$sum(X29,$sum($uminus(X27),X28))) (2:18:1) T [superposition 271,274]
[SA] active: 49213. ~$less(X8,-6) | $less(X8,-5) (8:6:1) T [evaluation 49212]
[SA] active: 1762. $sum(X80,$sum($uminus($sum(X78,X77)),X79)) = $sum($sum($uminus(X78),X79),$sum(X80,$uminus(X77))) (2:18:1) T [superposition 274,271]
[SA] active: 50924. ~$less(-5,X6) | $less(-6,X6) (8:6:1) T [resolution 31551,2395]
[SA] active: 1768. $sum(X4,$sum($uminus($sum(X5,X3)),X6)) = $sum($sum(X4,$uminus(X3)),$sum($uminus(X5),X6)) (2:18:1) T [forward demodulation 1767,17]
[SA] active: 2877. ~$less(sK0,X47) | $less(-3,X47) (9:6:1) [evaluation 2876]
[SA] active: 1770. $sum(X7,$sum($uminus($sum(X8,X9)),X10)) = $sum($sum($uminus(X8),X7),$sum($uminus(X9),X10)) (2:18:1) T [forward demodulation 1769,17]
[SA] active: 16062. ~$less(1,X0) | $less(-4,X0) (9:6:1) [resolution 14427,1498]
[SA] active: 1775. $sum(X4,$sum($uminus($sum(X5,X3)),X6)) = $sum($uminus(X5),$sum(X4,$sum($uminus(X3),X6))) (2:18:1) T [forward demodulation 1774,17]
[SA] active: 19886. ~$less(4,X5) | ~$less(X5,0) (9:6:1) [evaluation 19762]
[SA] active: 1782. $sum($uminus($sum(X9,$sum(X8,X11))),X10) = $sum($uminus(X11),$sum($uminus($sum(X9,X8)),X10)) (2:18:1) T [forward demodulation 1726,271]
[SA] active: 25829. $less(X27,2) | $less(-4,X27) (9:6:2) [evaluation 25591]
[SA] active: 1817. $sum(X76,$sum($uminus($sum(X74,X73)),X75)) = $sum($uminus(X73),$sum($uminus(X74),$sum(X75,X76))) (2:18:1) T [forward demodulation 1761,17]
[SA] active: 25841. $less(-5,X27) | $less(X27,1) (9:6:2) [evaluation 25840]
[SA] active: 1835. $less($sum(X16,$sum(X17,X15)),$sum(X18,$sum(X19,$sum(X16,X17)))) | ~$less(X15,$sum(X18,X19)) (2:18:1) T [superposition 276,274]
[SA] active: 52458. ~$less(3,X6) | ~$less(X6,-1) (9:6:1) [evaluation 52457]
[SA] active: 1836. $less($sum(X22,$sum(X20,X21)),$sum(X23,$sum(X24,$sum(X21,X22)))) | ~$less(X20,$sum(X23,X24)) (2:18:1) T [superposition 276,274]
[SA] active: 52464. ~$less(0,X3) | ~$less(X3,-3) (9:6:1) [evaluation 52463]
[SA] active: 1851. $less($sum(X18,$sum(X16,X17)),$sum(X19,$sum(X16,$sum(X17,X15)))) | ~$less(X18,$sum(X19,X15)) (2:18:1) T [superposition 276,274]
[SA] active: 53253. ~$less(2,X8) | ~$less(X8,-2) (9:6:1) [evaluation 53252]
[SA] active: 1852. $less($sum(X23,$sum(X21,X22)),$sum(X24,$sum(X22,$sum(X20,X21)))) | ~$less(X23,$sum(X24,X20)) (2:18:1) T [superposition 276,274]
[SA] active: 54636. $less(X6,5) | $less(-5,X6) (9:6:2) [resolution 33043,28833]
[SA] active: 1918. $less($sum(X18,$sum(X16,$sum(X17,X15))),$sum(X19,$sum(X16,X17))) | ~$less($sum(X18,X15),X19) (2:18:1) T [superposition 277,274]
[SA] active: 54637. ~$less(X7,-2) | $less(X7,5) (9:6:1) [resolution 33043,32067]
[SA] active: 1919. $less($sum(X23,$sum(X22,$sum(X20,X21))),$sum(X24,$sum(X21,X22))) | ~$less($sum(X23,X20),X24) (2:18:1) T [superposition 277,274]
[SA] active: 54768. ~$less(5,X0) | $less(0,X0) (9:6:1) [resolution 33043,2395]
[SA] active: 1944. $less($sum(X18,$sum(X19,$sum(X16,X17))),$sum(X16,$sum(X17,X15))) | ~$less($sum(X18,X19),X15) (2:18:1) T [superposition 277,274]
[SA] active: 55070. ~$less(5,X0) | $less(-2,X0) (9:6:1) [resolution 33408,2395]
[SA] active: 1945. $less($sum(X23,$sum(X24,$sum(X21,X22))),$sum(X22,$sum(X20,X21))) | ~$less($sum(X23,X24),X20) (2:18:1) T [superposition 277,274]
[SA] active: 55162. ~$less(5,X1) | $less(-1,X1) (9:6:1) [evaluation 55161]
[SA] active: 90804. $less(X188,$sum(X188,$sum(X189,$uminus(X186)))) | ~$less(X186,X189) (3:11:1) T [forward demodulation 90803,438]
[SA] active: 55801. $less(X6,6) | $less(2,X6) (9:6:2) [evaluation 55800]
[SA] active: 90825. $less(X254,$sum(X254,$sum(X255,X252))) | ~$less($uminus(X252),X255) (3:11:1) T [forward demodulation 90824,288]
[SA] active: 55909. ~$less(5,sK0) | $less(0,alength0) (9:6:1) [resolution 33446,2395]
[SA] active: 90777. $less($sum(X2,X1),$sum(X2,$sum(X3,X1))) | ~$less(0,X3) (3:12:1) T [forward demodulation 90275,20]
[SA] active: 56041. ~$less(5,sK0) | $less(1,alength0) (9:6:1) [resolution 33447,2395]
[SA] active: 90778. $less($sum(X5,X6),$sum(X6,$sum(X7,X5))) | ~$less(0,X7) (3:12:1) T [forward demodulation 90276,20]
[SA] active: 56905. $less(X8,6) | $less(3,X8) (9:6:2) [evaluation 56904]
[SA] active: 90828. $less(0,$sum(X257,$sum(X259,X256))) | ~$less($uminus($sum(X256,X257)),X259) (3:13:1) T [forward demodulation 90827,288]
[SA] active: 57046. ~$less(5,sK0) | $less(2,alength0) (9:6:1) [resolution 33820,2395]
[SA] active: 91012. $less($sum(X303,$sum(X304,X302)),0) | ~$less($sum(X303,X304),$uminus(X302)) (3:13:1) T [forward demodulation 90501,288]
[SA] active: 57124. ~$less(5,sK0) | $less(3,alength0) (9:6:1) [resolution 33821,2395]
[SA] active: 90192. $less($sum(X2,X1),$sum(X1,$sum(X3,$uminus(X0)))) | ~$less($sum(X2,X0),X3) (3:15:1) T [superposition 1945,288]
[SA] active: 57439. ~$less(5,X0) | $less(-4,X0) (9:6:1) [resolution 37564,2395]
[SA] active: 90201. $less($sum(X46,X45),$sum($uminus(X44),$sum(X47,X45))) | ~$less($sum(X46,X44),X47) (3:15:1) T [superposition 1945,438]
[SA] active: 57448. ~$less(X1,-4) | $less(X1,5) (9:6:1) [resolution 37564,2395]
[SA] active: 90263. $less($sum(X376,X375),$sum(X375,$sum(X377,X374))) | ~$less($sum(X376,$uminus(X374)),X377) (3:15:1) T [superposition 1945,287]
[SA] active: 57527. ~$less(X3,1) | $less(X3,9) (9:6:1) [evaluation 57526]
[SA] active: 90264. $less($sum(X380,X379),$sum(X378,$sum(X381,X379))) | ~$less($sum(X380,$uminus(X378)),X381) (3:15:1) T [superposition 1945,1086]
[SA] active: 57531. ~$less(5,X1) | $less(-3,X1) (9:6:1) [evaluation 57530]
[SA] active: 90278. $less($sum(X15,X14),$sum(X15,$sum(X16,$uminus(X13)))) | ~$less($sum(X13,X14),X16) (3:15:1) T [superposition 1945,1061]
[SA] active: 57696. ~$less(X2,-3) | $less(X2,5) (9:6:1) [resolution 37565,2395]
[SA] active: 90279. $less($sum(X18,X19),$sum($uminus(X17),$sum(X20,X19))) | ~$less($sum(X17,X18),X20) (3:15:1) T [superposition 1945,1080]
[SA] active: 57971. ~$less(6,X0) | $less(1,X0) (9:6:1) [resolution 37639,2395]
[SA] active: 90342. ~$less($sum($uminus($sum(X331,X332)),X333),X334) | $less(X333,$sum(X332,$sum(X334,X331))) (3:15:1) T [superposition 1945,1086]
[SA] active: 57978. ~$less(X0,2) | $less(X0,6) (9:6:1) [resolution 37639,511]
[SA] active: 90444. $less($sum(X2,$sum(X3,$sum(X1,X0))),X1) | ~$less($sum(X2,X3),$uminus(X0)) (3:15:1) T [superposition 1945,288]
[SA] active: 57984. ~$less(X6,-1) | $less(X6,6) (9:6:1) [resolution 37639,28201]
[SA] active: 90527. $less($sum(X453,$sum(X454,$sum(X452,X455))),X455) | ~$less($sum(X453,X454),$uminus(X452)) (3:15:1) T [evaluation 90422]
[SA] active: 58018. $less(sK0,6) | $less(1,alength0) (9:6:1) [resolution 37639,199]
[SA] active: 90533. ~$less($sum(X422,$uminus($sum(X420,X421))),X423) | $less(X422,$sum(X421,$sum(X423,X420))) (3:15:1) T [evaluation 90272]
[SA] active: 58071. ~$less(6,X1) | $less(2,X1) (9:6:1) [evaluation 58070]
[SA] active: 91010. $less($sum(X237,$sum(X238,$sum(X236,$uminus(X235)))),X236) | ~$less($sum(X237,X238),X235) (3:15:1) T [forward demodulation 91009,288]
[SA] active: 58602. ~$less(0,X2) | $less(-7,X2) (9:6:1) [evaluation 58601]
[SA] active: 90232. $less($sum(X212,$sum(2,X211)),$sum(X211,$sum(X213,1))) | ~$less($sum(X212,1),X213) (3:16:1) [superposition 1945,1181]
[SA] active: 59024. ~$less(X2,0) | $less(X2,7) (9:6:1) [evaluation 59023]
[SA] active: 90484. $less($sum(X212,$sum(X213,$sum(X211,1))),$sum(2,X211)) | ~$less($sum(X212,X213),1) (3:16:1) [superposition 1945,1181]
[SA] active: 59337. $less(X0,7) | $less(-4,X0) (9:6:2) [resolution 38065,20290]
[SA] active: 90076. ~$less($sum(X9,X10),$sum(X7,$sum(X8,$sum(1,X10)))) | ~$less($sum(X7,X8),X9) (3:16:1) T [resolution 1945,975]
[SA] active: 59338. $less(X1,7) | $less(-3,X1) (9:6:2) [resolution 38065,7880]
[SA] active: 90077. ~$less($sum(X14,X13),$sum(X11,$sum(X12,$sum(1,X14)))) | ~$less($sum(X11,X12),X13) (3:16:1) T [resolution 1945,281]
[SA] active: 59339. $less(X2,7) | $less(-2,X2) (9:6:2) [resolution 38065,6345]
[SA] active: 90078. ~$less($sum(X17,X18),$sum(X15,$sum(X16,$sum(X18,X17)))) | ~$less($sum(X15,X16),1) (3:16:1) T [resolution 1945,1291]
[SA] active: 59340. $less(X3,7) | $less(-1,X3) (9:6:2) [resolution 38065,3590]
[SA] active: 90087. ~$less($sum(X44,$sum(X43,X45)),$sum(X42,$sum(X45,X44))) | ~$less($sum(-1,X42),X43) (3:16:1) T [resolution 1945,453]
[SA] active: 59341. $less(X4,7) | $less(0,X4) (9:6:2) [resolution 38065,378]
[SA] active: 90534. ~$less($sum(X303,$sum($uminus(X301),X302)),X304) | $less(X303,$sum(X301,$sum(X304,$uminus(X302)))) (3:16:1) T [evaluation 90249]
[SA] active: 59366. $less(sK0,7) | $less(0,alength0) (9:6:1) [resolution 38065,393]
[SA] active: 90535. ~$less($sum(X232,$sum(X230,$uminus(X231))),X233) | $less(X232,$sum($uminus(X230),$sum(X233,X231))) (3:16:1) T [evaluation 90237]
[SA] active: 59420. $less(X6,8) | $less(1,X6) (9:6:2) [evaluation 59419]
[SA] active: 90854. ~$less($uminus($sum(X345,$sum(X346,X347))),X348) | $less($uminus(X347),$sum(X346,$sum(X348,X345))) (3:16:1) T [forward demodulation 90853,17]
[SA] active: 60354. ~$less(2,X8) | $less(-5,X8) (9:6:1) [evaluation 60353]
[SA] active: 90793. $less($sum(X153,$sum(i(s(X151)),X152)),$sum(X153,$sum(X154,X152))) | ~$less(i(s(X151)),X154) (3:18:1) [forward demodulation 90792,127]
[SA] active: 61908. ~$less(X0,-5) | $less(X0,2) (9:6:1) [resolution 39300,2395]
[SA] active: 90094. $less($sum(X2,$sum(X3,$sum(X1,X0))),$sum(X1,$sum(X4,X0))) | ~$less($sum(X2,X3),X4) (3:18:1) T [superposition 1945,16]
[SA] active: 61985. ~$less(X4,-5) | $less(X4,1) (9:6:1) [evaluation 61984]
[SA] active: 90195. $less($sum(X17,$sum(X16,$sum(X15,X14))),$sum(X16,$sum(X18,X15))) | ~$less($sum(X17,X14),X18) (3:18:1) T [superposition 1945,957]
[SA] active: 65509. ~$less(X0,-5) | $less(X0,3) (9:6:1) [resolution 40184,2395]
[SA] active: 90197. $less($sum(X27,$sum(X25,$sum(X24,X26))),$sum(X26,$sum(X28,X25))) | ~$less($sum(X27,X24),X28) (3:18:1) T [superposition 1945,956]
[SA] active: 65645. ~$less(3,X2) | $less(-5,X2) (9:6:1) [resolution 40184,2395]
[SA] active: 90199. $less($sum(X37,$sum(X35,$sum(X36,X34))),$sum(X36,$sum(X38,X35))) | ~$less($sum(X37,X34),X38) (3:18:1) T [superposition 1945,274]
[SA] active: 67255. $less(X0,6) | $less(-5,X0) (9:6:2) [resolution 40634,28837]
[SA] active: 90200. $less($sum(X42,$sum(X41,$sum(X39,X40))),$sum(X41,$sum(X43,X40))) | ~$less($sum(X42,X39),X43) (3:18:1) T [superposition 1945,274]
[SA] active: 70889. $less(-6,X1) | $less(X1,2) (9:6:2) [resolution 41136,3621]
[SA] active: 90284. $less($sum(X42,$sum(X44,$sum(X43,X41))),$sum(X44,$sum(X45,X43))) | ~$less($sum(X41,X42),X45) (3:18:1) T [superposition 1945,1011]
[SA] active: 70890. $less(-6,X2) | $less(X2,3) (9:6:2) [resolution 41136,6866]
[SA] active: 90285. $less($sum(X49,$sum(X46,$sum(X48,X47))),$sum(X49,$sum(X50,X48))) | ~$less($sum(X46,X47),X50) (3:18:1) T [superposition 1945,1011]
[SA] active: 70891. $less(-6,X3) | $less(X3,4) (9:6:2) [resolution 41136,11843]
[SA] active: 90286. $less($sum(X52,$sum(X53,$sum(X54,X51))),$sum(X54,$sum(X55,X53))) | ~$less($sum(X51,X52),X55) (3:18:1) T [superposition 1945,1012]
[SA] active: 70892. $less(-6,X4) | $less(X4,5) (9:6:2) [resolution 41136,33043]
[SA] active: 90287. $less($sum(X59,$sum(X56,$sum(X57,X58))),$sum(X59,$sum(X60,X58))) | ~$less($sum(X56,X57),X60) (3:18:1) T [superposition 1945,1012]
[SA] active: 72902. ~$less(6,X0) | $less(-4,X0) (9:6:1) [resolution 41435,2395]
[SA] active: 90294. $less($sum(X94,$sum(X92,$sum(X93,X91))),$sum(X94,$sum(X95,X93))) | ~$less($sum(X91,X92),X95) (3:18:1) T [superposition 1945,1095]
[SA] active: 72907. ~$less(X1,-4) | $less(X1,6) (9:6:1) [resolution 41435,2395]
[SA] active: 90296. $less($sum(X103,$sum(X104,$sum(X102,X101))),$sum(X104,$sum(X105,X103))) | ~$less($sum(X101,X102),X105) (3:18:1) T [superposition 1945,1096]
[SA] active: 74845. ~$less(6,X0) | $less(-3,X0) (9:6:1) [resolution 41436,2395]
[SA] active: 90297. $less($sum(X109,$sum(X108,$sum(X106,X107))),$sum(X109,$sum(X110,X108))) | ~$less($sum(X106,X107),X110) (3:18:1) T [superposition 1945,1096]
[SA] active: 74851. ~$less(X2,-3) | $less(X2,6) (9:6:1) [resolution 41436,2395]
[SA] active: 90300. $less($sum(X122,$sum(X121,$sum(X123,X124))),$sum(X124,$sum(X125,X123))) | ~$less($sum(X121,X122),X125) (3:18:1) T [superposition 1945,956]
[SA] active: 76027. ~$less(6,X0) | $less(-2,X0) (9:6:1) [resolution 41437,2395]
[SA] active: 90449. $less($sum(X27,$sum(X28,$sum(X26,X24))),$sum(X25,$sum(X24,X26))) | ~$less($sum(X27,X28),X25) (3:18:1) T [superposition 1945,956]
[SA] active: 76034. ~$less(X3,-2) | $less(X3,6) (9:6:1) [resolution 41437,2395]
[SA] active: 90553. $less($sum(X22,$sum(X23,$sum(X21,X20))),$sum(X20,$sum(X21,X24))) | ~$less($sum(X22,X23),X24) (3:18:1) T [forward demodulation 90552,17]
[SA] active: 76170. ~$less(6,X0) | $less(-1,X0) (9:6:1) [resolution 41438,2395]
[SA] active: 90780. $less($sum(X32,$sum(X33,$sum(X31,X34))),$sum(X34,$sum(X35,X33))) | ~$less($sum(X31,X32),X35) (3:18:1) T [forward demodulation 90282,17]
[SA] active: 76315. ~$less(6,X0) | $less(0,X0) (9:6:1) [resolution 41439,2395]
[SA] active: 90781. $less($sum(X39,$sum(X37,$sum(X36,X38))),$sum(X39,$sum(X40,X38))) | ~$less($sum(X36,X37),X40) (3:18:1) T [forward demodulation 90283,17]
[SA] active: 76326. $less(X8,6) | $less(-6,X8) (9:6:2) [resolution 41439,41136]
[SA] active: 90782. $less($sum(X63,$sum(X62,$sum(X64,X61))),$sum(X64,$sum(X65,X63))) | ~$less($sum(X61,X62),X65) (3:18:1) T [forward demodulation 90288,17]
[SA] active: 76399. ~$less(6,sK0) | $less(0,alength0) (9:6:1) [resolution 41464,2395]
[SA] active: 90138. $less($sum(X267,$sum(X268,i(X266))),$sum(-1,$sum(X269,i(s(X266))))) | ~$less($sum(X267,X268),X269) (3:19:1) [superposition 1945,1177]
[SA] active: 76468. ~$less(7,X0) | $less(1,X0) (9:6:1) [resolution 41512,2395]
[SA] active: 90224. $less($sum(X178,$sum(X177,i(s(X176)))),$sum(X177,$sum(X179,1))) | ~$less($sum(X178,i(X176)),X179) (3:19:1) [superposition 1945,1030]
[SA] active: 76471. ~$less(X0,2) | $less(X0,7) (9:6:1) [resolution 41512,511]
[SA] active: 90225. $less($sum(X182,$sum(i(s(X180)),X181)),$sum(X181,$sum(X183,1))) | ~$less($sum(X182,i(X180)),X183) (3:19:1) [superposition 1945,268]
[SA] active: 76477. ~$less(X6,-1) | $less(X6,7) (9:6:1) [resolution 41512,28201]
[SA] active: 90227. $less($sum(X191,$sum(i(s(X189)),X190)),$sum(1,$sum(X192,X190))) | ~$less($sum(X191,i(X189)),X192) (3:19:1) [superposition 1945,1152]
[SA] active: 76505. $less(sK0,7) | $less(1,alength0) (9:6:1) [resolution 41512,199]
[SA] active: 90228. $less($sum(X195,$sum(X194,i(s(X193)))),$sum(1,$sum(X196,X194))) | ~$less($sum(X195,i(X193)),X196) (3:19:1) [superposition 1945,1010]
[SA] active: 76765. ~$less(4,X7) | $less(-5,X7) (9:6:1) [evaluation 76764]
[SA] active: 90231. $less($sum(X209,$sum(i(X207),X208)),$sum(X208,$sum(X210,-1))) | ~$less($sum(X209,i(s(X207))),X210) (3:19:1) [superposition 1945,1178]
[SA] active: 76771. ~$less(0,X4) | $less(-8,X4) (9:6:1) [evaluation 76770]
[SA] active: 90233. $less($sum(X216,$sum(X215,i(s(X214)))),$sum(X215,$sum(X217,i(X214)))) | ~$less($sum(X216,1),X217) (3:19:1) [superposition 1945,1031]
[SA] active: 77965. $less(X0,8) | $less(-4,X0) (9:6:2) [resolution 42114,20290]
[SA] active: 90234. $less($sum(X220,$sum(i(s(X218)),X219)),$sum(X219,$sum(X221,i(X218)))) | ~$less($sum(X220,1),X221) (3:19:1) [superposition 1945,361]
[SA] active: 77966. $less(X1,8) | $less(-3,X1) (9:6:2) [resolution 42114,7880]
[SA] active: 90235. $less($sum(X224,$sum(i(s(X223)),X222)),$sum(i(X223),$sum(X225,X222))) | ~$less($sum(X224,1),X225) (3:19:1) [superposition 1945,1159]
[SA] active: 77967. $less(X2,8) | $less(-2,X2) (9:6:2) [resolution 42114,6345]
[SA] active: 90236. $less($sum(X228,$sum(X226,i(s(X227)))),$sum(i(X227),$sum(X229,X226))) | ~$less($sum(X228,1),X229) (3:19:1) [superposition 1945,1009]
[SA] active: 77968. $less(X3,8) | $less(-1,X3) (9:6:2) [resolution 42114,3590]
[SA] active: 90309. $less($sum(i(s(X163)),$sum(X164,X165)),$sum(1,$sum(X166,X165))) | ~$less($sum(i(X163),X164),X166) (3:19:1) [superposition 1945,1156]
[SA] active: 77969. $less(X4,8) | $less(0,X4) (9:6:2) [resolution 42114,378]
[SA] active: 90310. $less($sum(i(s(X167)),$sum(X169,X168)),$sum(X169,$sum(X170,1))) | ~$less($sum(i(X167),X168),X170) (3:19:1) [superposition 1945,1157]
[SA] active: 77995. $less(sK0,8) | $less(0,alength0) (9:6:1) [resolution 42114,393]
[SA] active: 90382. $less($sum(X225,$sum(X226,$sum(1,X227))),$sum(X227,i(s(X224)))) | ~$less($sum(X225,X226),i(X224)) (3:19:1) [superposition 1945,75]
[SA] active: 79011. ~$less(X0,-5) | $less(X0,4) (9:6:1) [resolution 44511,2395]
[SA] active: 90390. $less($sum(X267,$sum(X268,$sum(-1,X269))),$sum(X269,i(X266))) | ~$less($sum(X267,X268),i(s(X266))) (3:19:1) [superposition 1945,1177]
[SA] active: 79156. ~$less(X0,-2) | $less(X0,7) (9:6:1) [resolution 43482,41512]
[SA] active: 90393. $less($sum(X281,$sum(X282,$sum(i(X280),X283))),$sum(X283,i(s(X280)))) | ~$less($sum(X281,X282),1) (3:19:1) [superposition 1945,127]
[SA] active: 84826. $less(-6,X5) | $less(X5,7) (9:6:2) [resolution 49127,38065]
[SA] active: 90476. $less($sum(X178,$sum(X179,$sum(X177,i(X176)))),$sum(X177,i(s(X176)))) | ~$less($sum(X178,X179),1) (3:19:1) [superposition 1945,1030]
[SA] active: 84827. $less(-6,X6) | $less(X6,8) (9:6:2) [resolution 49127,42114]
[SA] active: 90477. $less($sum(X182,$sum(X183,$sum(X181,i(X180)))),$sum(i(s(X180)),X181)) | ~$less($sum(X182,X183),1) (3:19:1) [superposition 1945,268]
[SA] active: 79250. ~$less(X0,-6) | $less(X0,0) (9:6:1) T [resolution 44618,2395]
[SA] active: 90483. $less($sum(X209,$sum(X210,$sum(X208,i(s(X207))))),$sum(i(X207),X208)) | ~$less($sum(X209,X210),-1) (3:19:1) [superposition 1945,1178]
[SA] active: 79670. ~$less(X0,-6) | $less(X0,-1) (9:6:1) T [resolution 47397,2395]
[SA] active: 90485. $less($sum(X216,$sum(X217,$sum(X215,1))),$sum(X215,i(s(X214)))) | ~$less($sum(X216,X217),i(X214)) (3:19:1) [superposition 1945,1031]
[SA] active: 81518. ~$less(X0,-6) | $less(X0,-2) (9:6:1) T [resolution 47888,2395]
[SA] active: 90486. $less($sum(X220,$sum(X221,$sum(X219,1))),$sum(i(s(X218)),X219)) | ~$less($sum(X220,X221),i(X218)) (3:19:1) [superposition 1945,361]
[SA] active: 82902. ~$less(X0,-6) | $less(X0,-3) (9:6:1) T [resolution 48664,2395]
[SA] active: 90799. $less($sum(i(X171),$sum(X172,X173)),$sum(X173,$sum(X174,X172))) | ~$less($sum(-1,i(s(X171))),X174) (3:19:1) [forward demodulation 90311,16]
[SA] active: 84742. ~$less(X0,-6) | $less(X0,1) (9:6:1) T [resolution 49127,2395]
[SA] active: 90252. $less($sum(X319,$sum(X317,X318)),$sum(X318,$sum(X320,X316))) | ~$less($sum(X319,$sum($uminus(X316),X317)),X320) (3:19:1) T [superposition 1945,947]
[SA] active: 84901. ~$less(1,X7) | $less(-6,X7) (9:6:1) T [resolution 49127,2395]
[SA] active: 90253. $less($sum(X324,$sum(X323,X322)),$sum(X321,$sum(X325,X323))) | ~$less($sum(X324,$sum($uminus(X321),X322)),X325) (3:19:1) T [superposition 1945,1005]
[SA] active: 2891. ~$less(alength0,X26) | $less(-3,X26) (10:6:1) [evaluation 2890]
[SA] active: 90543. $less($sum(X37,$sum(X33,$sum(X34,$sum(X36,$uminus(X37))))),$sum(X35,X36)) | ~$less($sum(X33,X34),X35) (3:19:1) T [forward demodulation 90082,16]
[SA] active: 88607. $less(-5,X5) | $less(X5,7) (10:6:2) [resolution 25841,38065]
[SA] active: 90738. $less($sum(X7,$sum(X6,X5)),$sum(X5,$sum(X6,$sum(X8,$uminus(X4))))) | ~$less($sum(X7,X4),X8) (3:19:1) T [forward demodulation 90193,17]
[SA] active: 88608. $less(-5,X6) | $less(X6,8) (10:6:2) [resolution 25841,42114]
[SA] active: 88687. ~$less(1,X7) | $less(-5,X7) (10:6:1) [resolution 25841,2395]
[SA] active: 89696. ~$less(5,X0) | $less(-5,X0) (10:6:1) [resolution 54636,2395]
[SA] active: 89701. ~$less(X1,-5) | $less(X1,5) (10:6:1) [resolution 54636,2395]
[SA] active: 91297. ~$less(X0,3) | $less(X0,6) (10:6:1) [resolution 55801,512]
[SA] active: 91336. $less(sK0,6) | $less(2,alength0) (10:6:1) [resolution 55801,199]
[SA] active: 91507. ~$less(6,X0) | $less(3,X0) (10:6:1) [resolution 56905,2395]
[SA] active: 91549. $less(sK0,6) | $less(3,alength0) (10:6:1) [resolution 56905,199]
[SA] active: 91823. $less(X0,9) | $less(0,X0) (10:6:2) [resolution 57527,378]
[SA] active: 91824. $less(X1,9) | $less(-1,X1) (10:6:2) [resolution 57527,3590]
[SA] active: 91825. $less(X2,9) | $less(-2,X2) (10:6:2) [resolution 57527,6345]
% Refutation not found, non-redundant clauses discarded
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Refutation not found, non-redundant clauses discarded

% Memory used [KB]: 145328
% Time elapsed: 5.016 s
% ------------------------------
% ------------------------------
[SA] active: 67. ~$less(sK0,0) (0:3:1) [cnf transformation 57]
[SA] active: 88. $less(sK0,alength0) (0:3:1) [definition unfolding 66,83]
[SA] active: 96. ~$less(alength0,0) (0:3:1) [definition unfolding 84,83]
[SA] active: 21. ~$less(X0,X0) (0:3:1) T [theory axiom]
[SA] active: 85. i(zero) = 0 (0:4:1) [cnf transformation 3]
[SA] active: 87. 'Sub'(X0,s(X0)) (0:4:1) [cnf transformation 1]
[SA] active: 90. ~$less(i('$n'),alength0) (0:4:1) [definition unfolding 68,83]
[SA] active: 29. zero != s(X0) (0:4:1) T [term algebras distinctness]
[SA] active: 103. v0 != a('$n',sK0) (0:5:1) [backward demodulation 95,89]
[SA] active: 18. $sum(X0,0) = X0 (0:5:1) T [theory axiom]
[SA] active: 26. $uminus($uminus(X0)) = X0 (0:5:1) T [theory axiom]
[SA] active: 20. 0 = $sum(X0,$uminus(X0)) (0:6:1) T [theory axiom]
[SA] active: 109. 0 = $sum($uminus(X0),X0) (1:6:1) T [superposition 20,26]
[SA] active: 86. 'Sub'(X0,s(X1)) | ~'Sub'(X0,X1) (0:7:1) [cnf transformation 55]
[SA] active: 91. ~'Sub'(X0,'$n') | $less(i(X0),alength0) (0:7:1) [definition unfolding 69,83]
[SA] active: 95. v0 = a(s(X0),i(X0)) (0:7:1) [definition unfolding 82,77]
[SA] active: 113. v0 = a(s(zero),0) (1:6:1) [superposition 95,85]
[SA] active: 16. $sum(X0,X1) = $sum(X1,X0) (0:7:1) T [theory axiom]
[SA] active: 75. i(s(X0)) = $sum(i(X0),1) (0:8:1) [cnf transformation 50]
[SA] active: 78. ~'Sub'(X0,X1) | $less(i(X0),i(X1)) (0:8:1) [cnf transformation 51]
[SA] active: 131. i(s(zero)) = 1 (1:5:1) [evaluation 126]
[SA] active: 79. i(X0) != i(X1) | X0 = X1 (0:8:1) [cnf transformation 52]
[SA] active: 132. $less(i(X0),i(s(X0))) (1:6:1) [resolution 78,87]
[SA] active: 25. $less(X1,$sum(X0,1)) | $less(X0,X1) (0:8:1) T [theory axiom]
[SA] active: 148. $less(X0,$sum(X0,1)) (1:5:1) T [resolution 25,21]
[SA] active: 27. ~$less(X1,$sum(X0,1)) | ~$less(X0,X1) (0:8:1) T [theory axiom]
[SA] active: 154. $less(X0,$sum(1,X0)) (2:5:1) T [superposition 148,16]
[SA] active: 30. s(X0) != s(X1) | X0 = X1 (0:8:1) T [term algebras injectivity]
[SA] active: 136. i(s(s(zero))) = 2 (2:6:1) [evaluation 134]
[SA] active: 64. s(p(X0)) = X0 | zero = X0 (0:8:1) [cnf transformation 28]
[SA] active: 153. $less(X4,0) | $less(-1,X4) (2:6:2) T [evaluation 151]
[SA] active: 22. ~$less(X1,X2) | ~$less(X0,X1) | $less(X0,X2) (0:9:1) T [theory axiom]
[SA] active: 187. $less(-1,sK0) (3:3:1) [resolution 153,67]
[SA] active: 23. $less(X1,X0) | $less(X0,X1) | X0 = X1 (0:9:2) T [theory axiom]
[SA] active: 188. $less(-1,alength0) (3:3:1) [resolution 153,96]
[SA] active: 19. $sum($uminus(X1),$uminus(X0)) = $uminus($sum(X0,X1)) (0:10:1) T [theory axiom]
[SA] active: 199. ~$less(X11,sK0) | $less(X11,alength0) (1:6:1) [resolution 22,88]
[SA] active: 24. $less($sum(X0,X2),$sum(X1,X2)) | ~$less(X0,X1) (0:10:1) T [theory axiom]
[SA] active: 205. $less(0,sK0) | 0 = sK0 (1:6:1) [resolution 23,67]
[SA] active: 73. 'Sub'(sK2(X0),'$n') | a(zero,X0) = a('$n',X0) (0:11:1) [cnf transformation 61]
[SA] active: 206. $less(0,alength0) | alength0 = 0 (1:6:1) [resolution 23,96]
[SA] active: 108. ~$less(X0,i('$n')) | 'Sub'(sK3(X0),'$n') | $less(X0,0) (0:11:1) [backward demodulation 85,81]
[SA] active: 255. $less(0,alength0) | 0 = sK0 (2:6:1) [resolution 205,199]
[SA] active: 17. $sum(X0,$sum(X1,X2)) = $sum($sum(X0,X1),X2) (0:11:1) T [theory axiom]
[SA] active: 164. ~$less(-1,X4) | ~$less(X4,0) (2:6:1) T [evaluation 162]
[SA] active: 74. a(zero,X0) = a('$n',X0) | i(sK2(X0)) = X0 (0:12:1) [cnf transformation 61]
[SA] active: 201. ~$less(X0,-1) | $less(X0,sK0) (4:6:1) [resolution 187,22]
[SA] active: 76. a(X0,X1) = a(s(X0),X1) | i(X0) = X1 (0:12:1) [cnf transformation 50]
[SA] active: 215. ~$less(X0,-1) | $less(X0,alength0) (4:6:1) [resolution 188,22]
[SA] active: 107. ~$less(X0,i('$n')) | i(sK3(X0)) = X0 | $less(X0,0) (0:12:1) [backward demodulation 85,80]
[SA] active: 137. i(X0) != 0 | zero = X0 (1:7:1) [superposition 79,85]
[SA] active: 104. 'Sub'(sK1(X1,i(X1)),'$n') | v0 = a('$n',i(X1)) | ~'Sub'(X1,'$n') (0:15:1) [backward demodulation 95,102]
[SA] active: 184. 'Sub'(p(X7),X7) | zero = X7 (1:7:1) [superposition 87,64]
[SA] active: 105. sK1(X1,i(X1)) != X1 | v0 = a('$n',i(X1)) | ~'Sub'(X1,'$n') (0:15:1) [backward demodulation 95,100]
[SA] active: 135. v0 = a(s(s(zero)),1) (2:7:1) [superposition 95,131]
[SA] active: 106. ~'Sub'(X1,'$n') | i(X1) = i(sK1(X1,i(X1))) | v0 = a('$n',i(X1)) (0:17:1) [backward demodulation 95,98]
[SA] active: 177. i(s(s(s(zero)))) = 3 (3:7:1) [evaluation 175]
[SA] active: 127. i(s(X1)) = $sum(1,i(X1)) (1:8:1) [superposition 75,16]
[SA] active: 204. $less(alength0,i('$n')) | alength0 = i('$n') (1:8:1) [resolution 23,90]
[SA] active: 149. $less(X1,$sum(1,X0)) | $less(X0,X1) (1:8:1) T [superposition 25,16]
[SA] active: 378. $less(X2,1) | $less(0,X2) (2:6:2) T [superposition 149,18]
[SA] active: 160. ~$less(X1,$sum(1,X0)) | ~$less(X0,X1) (1:8:1) T [superposition 27,16]
[SA] active: 410. ~$less(0,X2) | ~$less(X2,1) (2:6:1) T [superposition 160,18]
[SA] active: 288. $sum(X2,$sum($uminus(X2),X3)) = X3 (1:8:1) T [evaluation 265]
[SA] active: 451. $less($sum(-1,X17),X17) (3:5:1) T [evaluation 449]
[SA] active: 133. ~'Sub'(X1,X2) | $less(i(X1),i(s(X2))) (1:9:1) [resolution 78,86]
[SA] active: 472. $less(-2,sK0) (5:3:1) [evaluation 465]
[SA] active: 152. $less(X6,i(s(X5))) | $less(i(X5),X6) (1:9:1) [superposition 25,75]
[SA] active: 473. $less(-2,alength0) (5:3:1) [evaluation 464]
[SA] active: 163. ~$less(X6,i(s(X5))) | ~$less(i(X5),X6) (1:9:1) [superposition 27,75]
[SA] active: 463. $less($sum(-1,sK0),alength0) (4:5:1) [resolution 451,199]
[SA] active: 240. $less(0,$sum(X3,$uminus(X2))) | ~$less(X2,X3) (1:9:1) T [superposition 24,20]
[SA] active: 466. $less($sum(X0,-1),X0) (4:5:1) T [superposition 451,16]
[SA] active: 247. $less($sum(X3,$uminus(X2)),0) | ~$less(X3,X2) (1:9:1) T [superposition 24,20]
[SA] active: 393. $less(sK0,1) | $less(0,alength0) (3:6:1) [resolution 378,199]
[SA] active: 183. ~'Sub'(X6,p(X5)) | 'Sub'(X6,X5) | zero = X5 (1:10:1) [superposition 86,64]
[SA] active: 419. ~$less(sK0,1) | 0 = sK0 (3:6:1) [resolution 410,205]
[SA] active: 185. v0 = a(X8,i(p(X8))) | zero = X8 (1:10:1) [superposition 95,64]
[SA] active: 420. ~$less(alength0,1) | 0 = sK0 (3:6:1) [resolution 410,255]
[SA] active: 303. v0 != a(zero,sK0) | i(sK2(sK0)) = sK0 (1:10:1) [superposition 103,74]
[SA] active: 421. ~$less(alength0,1) | alength0 = 0 (3:6:1) [resolution 410,206]
[SA] active: 216. $sum(X0,$uminus(X1)) = $uminus($sum(X1,$uminus(X0))) (1:10:1) T [superposition 19,26]
[SA] active: 496. $less(1,X1) | $less(X1,2) (3:6:2) [forward demodulation 493,131]
[SA] active: 217. $sum($uminus(X1),X0) = $uminus($sum($uminus(X0),X1)) (1:10:1) T [superposition 19,26]
[SA] active: 511. ~$less(1,X1) | ~$less(X1,2) (3:6:1) [forward demodulation 508,131]
[SA] active: 220. $sum($uminus(X5),$uminus(X4)) = $uminus($sum(X5,X4)) (1:10:1) T [superposition 19,16]
[SA] active: 497. $less(2,X2) | $less(X2,3) (4:6:2) [forward demodulation 494,136]
[SA] active: 241. $less($sum(X5,X4),$sum(X6,X5)) | ~$less(X4,X6) (1:10:1) T [superposition 24,16]
[SA] active: 512. ~$less(2,X2) | ~$less(X2,3) (4:6:1) [forward demodulation 509,136]
[SA] active: 248. $less($sum(X6,X5),$sum(X5,X4)) | ~$less(X6,X4) (1:10:1) T [superposition 24,16]
[SA] active: 614. $less(sK0,2) | $less(1,alength0) (4:6:1) [resolution 496,199]
[SA] active: 245. $less(i(s(X15)),$sum(X16,1)) | ~$less(i(X15),X16) (1:11:1) [superposition 24,75]
[SA] active: 729. $less(sK0,3) | $less(2,alength0) (5:6:1) [resolution 497,199]
[SA] active: 252. $less($sum(X16,1),i(s(X15))) | ~$less(X16,i(X15)) (1:11:1) [superposition 24,75]
[SA] active: 486. ~$less(X0,-2) | $less(X0,sK0) (6:6:1) [resolution 472,22]
[SA] active: 181. s(X2) != X1 | p(X1) = X2 | zero = X1 (1:11:1) [superposition 30,64]
[SA] active: 869. $less(-3,sK0) (7:3:1) [evaluation 864]
[SA] active: 194. ~$less(X0,X1) | $less(X0,$sum(X2,1)) | $less(X2,X1) (1:11:1) T [resolution 22,25]
[SA] active: 873. $less(-3,alength0) (8:3:1) [resolution 869,199]
[SA] active: 266. $sum(X4,$sum(X5,X6)) = $sum($sum(X5,X4),X6) (1:11:1) T [superposition 17,16]
[SA] active: 872. p(s(X0)) = X0 (2:5:1) [term algebras distinctness 871]
[SA] active: 274. $sum(X7,$sum(X8,X9)) = $sum(X9,$sum(X7,X8)) (1:11:1) T [superposition 17,16]
[SA] active: 921. $less(-2,X75) | $less(X75,alength0) (4:6:2) [evaluation 920]
[SA] active: 257. $less(i(sK2(X0)),alength0) | a(zero,X0) = a('$n',X0) (1:12:1) [resolution 73,91]
[SA] active: 1129. $less(-2,i('$n')) (5:4:1) [resolution 921,90]
[SA] active: 268. $sum(i(s(X10)),X11) = $sum(i(X10),$sum(1,X11)) (1:12:1) [superposition 17,75]
[SA] active: 923. $less(-2,X74) | $less(X74,sK0) (4:6:2) [evaluation 922]
[SA] active: 203. ~$less(X3,X2) | X1 = X2 | $less(X1,X2) | $less(X3,X1) (1:12:1) T [resolution 23,22]
[SA] active: 498. ~$less(X0,-2) | $less(X0,alength0) (6:6:1) [resolution 473,22]
[SA] active: 281. ~$less(X18,$sum(X16,$sum(X17,1))) | ~$less($sum(X16,X17),X18) (1:12:1) T [superposition 27,17]
[SA] active: 917. $less(-3,X77) | $less(X77,alength0) (6:6:2) [evaluation 916]
[SA] active: 283. $less(X23,$sum(X21,$sum(X22,1))) | $less($sum(X21,X22),X23) (1:12:1) T [superposition 25,17]
[SA] active: 1338. $less(-3,i('$n')) (7:4:1) [resolution 917,90]
[SA] active: 258. $less(i(sK2(X1)),i('$n')) | a(zero,X1) = a('$n',X1) (1:13:1) [resolution 73,78]
[SA] active: 919. $less(-3,X76) | $less(X76,sK0) (6:6:2) [evaluation 918]
[SA] active: 326. 'Sub'(sK1(zero,0),'$n') | v0 = a('$n',0) | ~'Sub'(zero,'$n') (1:13:1) [forward demodulation 323,85]
[SA] active: 874. ~$less(X0,-3) | $less(X0,sK0) (8:6:1) [resolution 869,22]
[SA] active: 334. zero != sK1(zero,0) | v0 = a('$n',0) | ~'Sub'(zero,'$n') (1:13:1) [forward demodulation 331,85]
[SA] active: 1443. $less(-4,sK0) (9:3:1) [evaluation 1433]
[SA] active: 207. $less($sum(X4,1),X5) | $sum(X4,1) = X5 | ~$less(X4,X5) (1:13:1) T [resolution 23,27]
[SA] active: 1444. $less(-4,alength0) (10:3:1) [resolution 1443,199]
[SA] active: 238. ~$less(X5,$sum(X3,X6)) | ~$less(X3,X4) | $less(X5,$sum(X4,X6)) (1:13:1) T [resolution 24,22]
[SA] active: 915. $less(-4,X78) | $less(X78,sK0) (8:6:2) [evaluation 914]
[SA] active: 244. $less($uminus($sum(X13,X12)),$sum(X14,$uminus(X13))) | ~$less($uminus(X12),X14) (1:13:1) T [superposition 24,19]
[SA] active: 941. ~$less(X1,-3) | $less(X1,alength0) (9:6:1) [resolution 873,22]
[SA] active: 251. $less($sum(X14,$uminus(X13)),$uminus($sum(X13,X12))) | ~$less(X14,$uminus(X12)) (1:13:1) T [superposition 24,19]
[SA] active: 943. $less(-4,X0) | $less(X0,alength0) (9:6:2) [evaluation 942]
[SA] active: 271. $sum($uminus($sum(X19,X18)),X20) = $sum($uminus(X18),$sum($uminus(X19),X20)) (1:14:1) T [superposition 17,19]
[SA] active: 1706. $less(-4,i('$n')) (10:4:1) [resolution 943,90]
[SA] active: 276. $less($sum(X3,X2),$sum(X0,$sum(X1,X2))) | ~$less(X3,$sum(X0,X1)) (1:14:1) T [superposition 24,17]
[SA] active: 1447. ~$less(X2,-4) | $less(X2,sK0) (10:6:1) [resolution 1443,22]
[SA] active: 277. $less($sum(X4,$sum(X5,X6)),$sum(X7,X6)) | ~$less($sum(X4,X5),X7) (1:14:1) T [superposition 24,17]
[SA] active: 1906. $less(-5,sK0) (11:3:1) [evaluation 1891]
[SA] active: 260. 'Sub'(sK3(X0),'$n') | $less(X0,0) | $less(i('$n'),X0) | i('$n') = X0 (1:15:1) [resolution 108,23]
[SA] active: 1970. $less(-5,alength0) (12:3:1) [resolution 1906,199]
[SA] active: 307. a(X0,X1) = a(p(X0),X1) | i(p(X0)) = X1 | zero = X0 (1:16:1) [superposition 76,64]
[SA] active: 1449. $less(-5,X1) | $less(X1,sK0) (10:6:2) [evaluation 1448]
[SA] active: 315. $less(i('$n'),X0) | $less(X0,0) | i(sK3(X0)) = X0 | i('$n') = X0 (1:16:1) [resolution 107,23]
[SA] active: 1494. ~$less(X2,-4) | $less(X2,alength0) (11:6:1) [resolution 1444,22]
[SA] active: 321. ~'Sub'(X0,'$n') | v0 = a('$n',i(X0)) | $less(i(sK1(X0,i(X0))),alength0) (1:16:1) [resolution 104,91]
[SA] active: 1496. $less(-5,X1) | $less(X1,alength0) (11:6:2) [evaluation 1495]
[SA] active: 322. ~'Sub'(X1,'$n') | v0 = a('$n',i(X1)) | $less(i(sK1(X1,i(X1))),i('$n')) (1:17:1) [resolution 104,78]
[SA] active: 2095. $less(-5,i('$n')) (12:4:1) [resolution 1496,90]
[SA] active: 339. i(sK2(X0)) = i(sK1(sK2(X0),i(sK2(X0)))) | v0 = a('$n',i(sK2(X0))) | a(zero,X0) = a('$n',X0) (1:25:1) [resolution 106,73]
[SA] active: 1973. ~$less(X2,-5) | $less(X2,sK0) (12:6:1) [resolution 1906,22]
[SA] active: 340. ~'Sub'(X1,'$n') | v0 = a('$n',i(sK1(X1,i(X1)))) | v0 = a('$n',i(X1)) | i(sK1(X1,i(X1))) = i(sK1(sK1(X1,i(X1)),i(sK1(X1,i(X1))))) (1:35:1) [resolution 106,104]
[SA] active: 2155. $less(-6,sK0) (13:3:1) [evaluation 2138]
[SA] active: 480. $less(i(X0),i(s(s(X0)))) (2:7:1) [resolution 133,87]
[SA] active: 2161. $less(-6,alength0) (14:3:1) [resolution 2155,199]
[SA] active: 514. ~$less($uminus(X0),0) | ~$less(X0,1) (2:7:1) T [resolution 240,160]
[SA] active: 2167. ~$less(i(s(X0)),i(X0)) (3:6:1) [resolution 480,163]
[SA] active: 138. i(X1) != 1 | s(zero) = X1 (2:8:1) [superposition 79,131]
[SA] active: 1975. $less(-6,X1) | $less(X1,sK0) (12:6:2) [evaluation 1974]
[SA] active: 329. $less(i(p('$n')),alength0) | zero = '$n' (2:8:1) [resolution 184,91]
[SA] active: 1982. ~$less(X2,-5) | $less(X2,alength0) (13:6:1) [resolution 1970,22]
[SA] active: 372. 'Sub'(sK3(alength0),'$n') | alength0 = i('$n') (2:8:1) [subsumption resolution 369,96]
[SA] active: 1984. $less(-6,X1) | $less(X1,alength0) (13:6:2) [evaluation 1983]
[SA] active: 903. $less(sK0,$sum(X66,1)) | $less(X66,alength0) (2:8:1) [resolution 194,88]
[SA] active: 2307. $less(-6,i('$n')) (14:4:1) [resolution 1984,90]
[SA] active: 1177. i(X15) = $sum(i(s(X15)),-1) (2:8:1) [evaluation 1158]
[SA] active: 2164. ~$less(X2,-6) | $less(X2,sK0) (14:6:1) [resolution 2155,22]
[SA] active: 195. $less(X3,$sum(X4,1)) | ~$less(X3,X4) (2:8:1) T [resolution 22,148]
[SA] active: 2394. $less(-7,sK0) (15:3:1) [evaluation 2375]
[SA] active: 287. $sum($uminus(X16),$sum(X16,X17)) = X17 (2:8:1) T [evaluation 270]
[SA] active: 2421. $less(-7,alength0) (16:3:1) [resolution 2394,199]
[SA] active: 407. ~$less(X9,$sum(X10,X9)) | ~$less(X10,1) (2:8:1) T [resolution 160,24]
[SA] active: 2398. ~$less($sum(X6,1),X6) (3:5:1) T [resolution 195,21]
[SA] active: 438. $sum(X3,$sum(X4,$uminus(X3))) = X4 (2:8:1) T [superposition 288,16]
[SA] active: 2621. $less(X9,$sum(2,X9)) (4:5:1) [forward demodulation 2620,1181]
[SA] active: 452. $less($sum(-1,X15),X16) | $less(X16,X15) (2:8:1) T [evaluation 448]
[SA] active: 2570. ~$less($sum(1,X5),X5) (4:5:1) T [resolution 2398,248]
[SA] active: 453. ~$less($sum(-1,X13),X14) | ~$less(X14,X13) (2:8:1) T [evaluation 447]
[SA] active: 2824. ~$less(alength0,sK0) (5:3:1) [resolution 453,463]
[SA] active: 744. ~$less(X3,$sum(X3,X2)) | ~$less(X2,1) (2:8:1) T [resolution 241,160]
[SA] active: 2713. $less(X0,$sum(X0,2)) (5:5:1) [superposition 2621,16]
[SA] active: 749. $less(X6,$sum(X7,X6)) | ~$less(0,X7) (2:8:1) T [superposition 241,18]
[SA] active: 2721. $less($sum(X7,-2),X7) (5:5:1) [evaluation 2719]
[SA] active: 804. $less($sum(X7,X6),X6) | ~$less(X7,0) (2:8:1) T [superposition 248,18]
[SA] active: 3063. $less(-8,sK0) (15:3:1) [evaluation 3057]
[SA] active: 1086. $sum($uminus(X19),$sum(X20,X19)) = X20 (2:8:1) T [evaluation 1008]
[SA] active: 3134. $less(-8,alength0) (16:3:1) [resolution 3063,199]
[SA] active: 1306. ~$less($sum(X18,-1),X17) | ~$less(X17,X18) (2:8:1) T [evaluation 1296]
[SA] active: 2722. $less($sum(-2,X2),X2) (5:5:1) [evaluation 2716]
[SA] active: 1374. $less($sum(X18,-1),X17) | $less(X17,X18) (2:8:1) T [evaluation 1364]
[SA] active: 2796. ~$less(X8,$sum(X8,-1)) (5:5:1) T [evaluation 2794]
[SA] active: 186. $less(i(p(X9)),i(X9)) | zero = X9 (2:9:1) [superposition 132,64]
[SA] active: 2797. ~$less(X3,$sum(-1,X3)) (5:5:1) T [evaluation 2791]
[SA] active: 198. $less(X9,i(s(X10))) | ~$less(X9,i(X10)) (2:9:1) [resolution 22,132]
[SA] active: 3088. $less($sum(-2,sK0),alength0) (6:5:1) [forward demodulation 3046,16]
[SA] active: 232. $less(sK0,X0) | $less(X0,alength0) | sK0 = X0 (2:9:2) [resolution 199,23]
[SA] active: 3536. ~$less(X0,0) | $less(X0,1) (3:6:1) [forward demodulation 3530,85]
[SA] active: 256. ~$less(X0,0) | 0 = sK0 | $less(X0,sK0) (2:9:1) [resolution 205,22]
[SA] active: 3537. ~$less(X1,1) | $less(X1,2) (3:6:1) [forward demodulation 3531,131]
[SA] active: 259. ~$less(X0,0) | alength0 = 0 | $less(X0,alength0) (2:9:1) [resolution 206,22]
[SA] active: 2395. ~$less(X1,X0) | ~$less(X0,X1) (3:6:1) T [resolution 195,27]
[SA] active: 371. alength0 = i(sK3(alength0)) | alength0 = i('$n') (2:9:1) [subsumption resolution 368,96]
[SA] active: 3811. ~$less(sK0,-1) (4:3:1) [resolution 2395,187]
[SA] active: 856. $less(0,i(s(X9))) | ~$less(-1,i(X9)) (2:9:1) [evaluation 845]
[SA] active: 3812. ~$less(alength0,-1) (4:3:1) [resolution 2395,188]
[SA] active: 931. $less(-1,X40) | $less(X40,sK0) | 0 = sK0 (2:9:2) [evaluation 930]
[SA] active: 3815. ~$less(alength0,-2) (6:3:1) [resolution 2395,473]
[SA] active: 935. $less(-1,X36) | $less(X36,alength0) | alength0 = 0 (2:9:2) [evaluation 934]
[SA] active: 3817. ~$less(sK0,-2) (6:3:1) [resolution 2395,472]
[SA] active: 1181. $sum(2,X1) = $sum(1,$sum(1,X1)) (2:9:1) [forward demodulation 1148,136]
[SA] active: 3822. ~$less(sK0,-3) (8:3:1) [resolution 2395,869]
[SA] active: 1248. $less(sK0,X72) | $less(X72,alength0) | alength0 = X72 (2:9:2) [resolution 203,88]
[SA] active: 3820. ~$less(alength0,-3) (9:3:1) [resolution 2395,873]
[SA] active: 2047. i(sK3(sK0)) = sK0 | i('$n') = sK0 (2:9:1) [subsumption resolution 2046,90]
[SA] active: 3827. ~$less(sK0,-4) (10:3:1) [resolution 2395,1443]
[SA] active: 243. $less(0,$sum(X11,X10)) | ~$less($uminus(X10),X11) (2:9:1) T [superposition 24,109]
[SA] active: 3825. ~$less(alength0,-4) (11:3:1) [resolution 2395,1444]
[SA] active: 250. $less($sum(X11,X10),0) | ~$less(X11,$uminus(X10)) (2:9:1) T [superposition 24,109]
[SA] active: 3832. ~$less(sK0,-5) (12:3:1) [resolution 2395,1906]
[SA] active: 282. $less($sum(X19,X20),$sum(X19,$sum(X20,1))) (2:9:1) T [superposition 148,17]
[SA] active: 3830. ~$less(alength0,-5) (13:3:1) [resolution 2395,1970]
[SA] active: 437. $sum(X1,$uminus($sum(X2,X1))) = $uminus(X2) (2:9:1) T [superposition 288,19]
[SA] active: 3837. ~$less(sK0,-6) (14:3:1) [resolution 2395,2155]
[SA] active: 519. $less(0,$sum($uminus(X2),X1)) | ~$less(X2,X1) (2:9:1) T [superposition 240,16]
[SA] active: 3835. ~$less(alength0,-6) (15:3:1) [resolution 2395,2161]
[SA] active: 557. $less($sum($uminus(X2),X1),0) | ~$less(X1,X2) (2:9:1) T [superposition 247,16]
[SA] active: 3838. ~$less(sK0,-7) (16:3:1) [resolution 2395,2394]
[SA] active: 656. $sum($uminus($sum(X8,X9)),X8) = $uminus(X9) (2:9:1) T [forward demodulation 655,19]
[SA] active: 3840. ~$less(sK0,-8) (16:3:1) [resolution 2395,3063]
[SA] active: 686. $uminus($sum(X2,X3)) = $uminus($sum(X3,X2)) (2:9:1) T [superposition 220,19]
[SA] active: 3839. ~$less(alength0,-7) (17:3:1) [resolution 2395,2421]
[SA] active: 694. $sum(X5,$uminus($sum(X5,X6))) = $uminus(X6) (2:9:1) T [superposition 288,220]
[SA] active: 3841. ~$less(alength0,-8) (17:3:1) [resolution 2395,3134]
[SA] active: 766. $less($sum(X21,X22),0) | ~$less(X22,$uminus(X21)) (2:9:1) T [superposition 241,109]
[SA] active: 3816. ~$less(i('$n'),-2) (6:4:1) [resolution 2395,1129]
[SA] active: 799. $less(0,$sum(X21,X22)) | ~$less($uminus(X21),X22) (2:9:1) T [superposition 248,109]
[SA] active: 3821. ~$less(i('$n'),-3) (8:4:1) [resolution 2395,1338]
[SA] active: 976. $less($sum(X48,X49),$sum(X49,$sum(X48,1))) (2:9:1) T [superposition 148,266]
[SA] active: 3826. ~$less(i('$n'),-4) (11:4:1) [resolution 2395,1706]
[SA] active: 1473. ~$less(-1,X12) | 0 = X12 | $less(0,X12) (2:9:1) T [evaluation 1472]
[SA] active: 3831. ~$less(i('$n'),-5) (13:4:1) [resolution 2395,2095]
[SA] active: 481. ~'Sub'(X1,X2) | $less(i(X1),i(s(s(X2)))) (2:10:1) [resolution 133,86]
[SA] active: 3836. ~$less(i('$n'),-6) (15:4:1) [resolution 2395,2307]
[SA] active: 484. $less(i(p(X5)),i(s(X5))) | zero = X5 (2:10:1) [resolution 133,184]
[SA] active: 3754. ~$less($sum(2,X17),X17) (5:5:1) [resolution 2395,2621]
[SA] active: 522. $less(0,$uminus($sum(X9,X8))) | ~$less(X9,$uminus(X8)) (2:10:1) T [superposition 240,19]
[SA] active: 3792. ~$less(alength0,$sum(-1,sK0)) (5:5:1) [resolution 2395,463]
[SA] active: 560. $less($uminus($sum(X9,X8)),0) | ~$less($uminus(X8),X9) (2:10:1) T [superposition 247,19]
[SA] active: 3761. ~$less($sum(X30,2),X30) (6:5:1) [resolution 2395,2713]
[SA] active: 586. $sum(X2,$uminus(X1)) = $uminus($sum($uminus(X2),X1)) (2:10:1) T [superposition 216,16]
[SA] active: 3786. ~$less(X52,$sum(X52,-2)) (6:5:1) [resolution 2395,2721]
[SA] active: 632. $sum($uminus(X4),X3) = $uminus($sum(X4,$uminus(X3))) (2:10:1) T [superposition 217,16]
[SA] active: 3787. ~$less(X53,$sum(-2,X53)) (6:5:1) [resolution 2395,2722]
[SA] active: 697. $less($uminus($sum(X10,X11)),0) | ~$less($uminus(X10),X11) (2:10:1) T [superposition 247,220]
[SA] active: 3793. ~$less(alength0,$sum(-2,sK0)) (7:5:1) [resolution 2395,3088]
[SA] active: 698. $less(0,$uminus($sum(X12,X13))) | ~$less(X13,$uminus(X12)) (2:10:1) T [superposition 240,220]
[SA] active: 2412. ~$less(X12,-1) | $less(X12,0) (3:6:1) T [evaluation 2411]
[SA] active: 758. $less($sum(X1,X2),$sum(X1,X0)) | ~$less(X2,X0) (2:10:1) T [superposition 241,16]
[SA] active: 3538. ~$less(X2,2) | $less(X2,3) (4:6:1) [forward demodulation 3532,136]
[SA] active: 881. $less(X16,$sum(X17,1)) | $less(X17,$sum(X16,1)) (2:10:2) T [resolution 194,148]
[SA] active: 6280. $less(X12,0) | $less(-2,X12) (3:6:2) T [evaluation 6279]
[SA] active: 1277. ~$less($sum(1,X1),X2) | $less($sum(X1,1),X2) (2:10:1) T [resolution 281,149]
[SA] active: 6505. $less(X42,-1) | $less(-2,X42) (3:6:2) T [evaluation 6504]
[SA] active: 1350. ~$less($sum(X3,1),X4) | $less($sum(1,X3),X4) (2:10:1) T [resolution 283,160]
[SA] active: 3590. $less(X0,1) | $less(-1,X0) (4:6:2) [resolution 3536,153]
[SA] active: 362. $less($sum(X4,i(X3)),i(s(X3))) | ~$less(X4,1) (2:11:1) [superposition 24,127]
[SA] active: 3621. $less(X7,2) | $less(0,X7) (4:6:2) [evaluation 3620]
[SA] active: 363. $less(i(s(X5)),$sum(X6,i(X5))) | ~$less(1,X6) (2:11:1) [superposition 24,127]
[SA] active: 3679. $less(sK0,2) | $less(0,alength0) (4:6:1) [resolution 3537,393]
[SA] active: 370. $less(X0,i('$n')) | ~$less(X0,alength0) | alength0 = i('$n') (2:11:1) [resolution 204,22]
[SA] active: 3696. $less(X5,3) | $less(1,X5) (4:6:2) [evaluation 3695]
[SA] active: 763. $less($sum(1,X14),i(s(X13))) | ~$less(X14,i(X13)) (2:11:1) [superposition 241,75]
[SA] active: 3767. ~$less(2,X36) | $less(1,X36) (4:6:1) [resolution 2395,496]
[SA] active: 764. $less($sum(i(X15),X16),i(s(X15))) | ~$less(X16,1) (2:11:1) [superposition 241,127]
[SA] active: 3805. ~$less(1,sK0) | $less(0,alength0) (4:6:1) [resolution 2395,393]
[SA] active: 796. $less(i(s(X13)),$sum(1,X14)) | ~$less(i(X13),X14) (2:11:1) [superposition 248,75]
[SA] active: 6345. $less(-2,X2) | $less(X2,1) (4:6:2) [resolution 6280,3536]
[SA] active: 797. $less(i(s(X15)),$sum(i(X15),X16)) | ~$less(1,X16) (2:11:1) [superposition 248,127]
[SA] active: 3751. ~$less(0,X14) | $less(-1,X14) (4:6:1) T [resolution 2395,153]
[SA] active: 887. $less(i(X28),$sum(X29,1)) | $less(X29,i(s(X28))) (2:11:1) [resolution 194,132]
[SA] active: 3752. ~$less(1,X15) | $less(0,X15) (4:6:1) T [resolution 2395,378]
[SA] active: 376. ~$less(X5,X4) | $less(X3,X4) | $less(X5,$sum(1,X3)) (2:11:1) T [resolution 149,22]
[SA] active: 6359. ~$less(0,X4) | $less(-2,X4) (4:6:1) T [resolution 6280,2395]
[SA] active: 445. $less($sum(X5,$sum($uminus(X3),X4)),X4) | ~$less(X5,X3) (2:11:1) T [superposition 24,288]
[SA] active: 6366. ~$less(X0,-2) | $less(X0,0) (4:6:1) T [resolution 6280,2395]
[SA] active: 446. $less(X7,$sum(X8,$sum($uminus(X6),X7))) | ~$less(X6,X8) (2:11:1) T [superposition 24,288]
[SA] active: 6404. ~$less(X13,-2) | $less(X13,-1) (4:6:1) T [evaluation 6403]
[SA] active: 593. 0 = $sum($sum(X6,$uminus(X7)),$sum(X7,$uminus(X6))) (2:11:1) T [superposition 20,216]
[SA] active: 6507. $less(X41,-1) | $less(-3,X41) (4:6:2) T [evaluation 6506]
[SA] active: 641. 0 = $sum($sum($uminus(X8),X9),$sum($uminus(X9),X8)) (2:11:1) T [superposition 20,217]
[SA] active: 6566. ~$less(-1,X3) | $less(-2,X3) (4:6:1) T [resolution 6505,2395]
[SA] active: 750. ~$less($sum($uminus(X8),X9),X10) | $less(X9,$sum(X10,X8)) (2:11:1) T [superposition 241,288]
[SA] active: 6624. $less(-3,X12) | $less(X12,-2) (4:6:2) T [evaluation 6623]
[SA] active: 770. $less($sum($uminus(X8),$sum(X9,X10)),X9) | ~$less(X10,X8) (2:11:1) T [forward demodulation 761,17]
[SA] active: 2873. ~$less(sK0,X49) | $less(-1,X49) (5:6:1) [evaluation 2872]
[SA] active: 805. ~$less(X10,$sum($uminus(X8),X9)) | $less($sum(X10,X8),X9) (2:11:1) T [superposition 248,288]
[SA] active: 2887. ~$less(alength0,X28) | $less(-1,X28) (5:6:1) [evaluation 2886]
[SA] active: 814. $less(X9,$sum($uminus(X8),$sum(X9,X10))) | ~$less(X8,X10) (2:11:1) T [forward demodulation 794,17]
[SA] active: 3768. ~$less(3,X37) | $less(2,X37) (5:6:1) [resolution 2395,497]
[SA] active: 956. $sum(X3,$sum(X4,X5)) = $sum(X4,$sum(X3,X5)) (2:11:1) T [superposition 266,17]
[SA] active: 3807. ~$less(2,sK0) | $less(1,alength0) (5:6:1) [resolution 2395,614]
[SA] active: 957. $sum(X7,$sum(X6,X8)) = $sum(X8,$sum(X6,X7)) (2:11:1) T [superposition 266,16]
[SA] active: 6189. $less(sK0,3) | $less(1,alength0) (5:6:1) [resolution 3538,614]
[SA] active: 1018. $sum(X10,$sum(X8,X9)) = $sum($sum(X9,X10),X8) (2:11:1) T [superposition 274,16]
[SA] active: 6210. $less(X5,4) | $less(2,X5) (5:6:2) [evaluation 6209]
[SA] active: 1498. $less(X3,$sum(X2,X4)) | ~$less(1,X2) | $less(X4,X3) (2:11:1) T [resolution 238,149]
[SA] active: 6746. $less(-1,X0) | $less(X0,2) (5:6:2) [resolution 3590,3537]
[SA] active: 1499. ~$less(X5,X6) | $less(X7,$sum(X6,1)) | $less(X5,X7) (2:11:1) T [resolution 238,25]
[SA] active: 6757. ~$less(1,X1) | $less(-1,X1) (5:6:1) [resolution 3590,2395]
[SA] active: 1566. ~$less($uminus(X5),$uminus($sum(X5,X4))) | ~$less($uminus(X4),1) (2:11:1) T [resolution 244,160]
[SA] active: 6768. ~$less(X2,-1) | $less(X2,1) (5:6:1) [resolution 3590,2395]
[SA] active: 1873. ~$less(X49,$sum(X50,$uminus(X48))) | $less($sum(X49,X48),X50) (2:11:1) T [evaluation 1859]
[SA] active: 6866. $less(0,X0) | $less(X0,3) (5:6:2) [resolution 3621,3538]
[SA] active: 1874. $less($sum(X26,$uminus(X25)),X27) | ~$less(X26,$sum(X27,X25)) (2:11:1) T [evaluation 1853]
[SA] active: 6877. ~$less(2,X1) | $less(0,X1) (5:6:1) [resolution 3621,2395]
[SA] active: 1957. ~$less($sum(X49,$uminus(X48)),X50) | $less(X49,$sum(X50,X48)) (2:11:1) T [evaluation 1926]
[SA] active: 6887. ~$less(X1,0) | $less(X1,2) (5:6:1) [resolution 3621,2395]
[SA] active: 1958. $less(X26,$sum(X27,$uminus(X25))) | ~$less($sum(X26,X25),X27) (2:11:1) T [evaluation 1920]
[SA] active: 6984. $less(sK0,3) | $less(0,alength0) (5:6:1) [resolution 3679,3538]
[SA] active: 361. $sum(i(s(X1)),X2) = $sum(1,$sum(i(X1),X2)) (2:12:1) [superposition 17,127]
[SA] active: 12324. $less(X8,$sum(3,X8)) (6:5:1) [backward demodulation 12321,5412]
[SA] active: 490. ~$less(X5,X4) | $less(i(X3),X4) | $less(X5,i(s(X3))) (2:12:1) [resolution 152,22]
[SA] active: 12325. ~$less($sum(3,X7),X7) (6:5:1) [backward demodulation 12321,5411]
[SA] active: 491. $less(i(p(X0)),X1) | $less(X1,i(X0)) | zero = X0 (2:12:1) [superposition 152,64]
[SA] active: 12661. $less(-9,sK0) (13:3:1) [evaluation 12660]
[SA] active: 506. ~$less(i(p(X0)),X1) | ~$less(X1,i(X0)) | zero = X0 (2:12:1) [superposition 163,64]
[SA] active: 12665. $less(-9,alength0) (14:3:1) [evaluation 12664]
[SA] active: 566. 'Sub'(p(p(X0)),X0) | zero = X0 | zero = p(X0) (2:12:1) [resolution 183,184]
[SA] active: 12749. ~$less(sK0,-9) (14:3:1) [resolution 12661,2395]
[SA] active: 825. $less(i(s(X5)),i(s(X4))) | ~$less(i(X5),i(X4)) (2:12:1) [superposition 245,75]
[SA] active: 12832. ~$less(alength0,-9) (15:3:1) [resolution 12665,2395]
[SA] active: 1009. $sum(X22,i(s(X21))) = $sum(1,$sum(X22,i(X21))) (2:12:1) [superposition 274,75]
[SA] active: 12404. $less($sum(X15,-3),X15) (7:5:1) [evaluation 12400]
[SA] active: 1010. $sum(X24,i(s(X23))) = $sum(i(X23),$sum(X24,1)) (2:12:1) [superposition 274,127]
[SA] active: 12405. $less($sum(-3,X2),X2) (7:5:1) [evaluation 12393]
[SA] active: 1030. $sum(X22,i(s(X21))) = $sum(i(X21),$sum(1,X22)) (2:12:1) [superposition 274,75]
[SA] active: 12417. $less(X19,$sum(X19,3)) (7:5:1) [forward demodulation 12416,1061]
[SA] active: 1031. $sum(X24,i(s(X23))) = $sum(1,$sum(i(X23),X24)) (2:12:1) [superposition 274,127]
[SA] active: 12592. ~$less($sum(X2,3),X2) (7:5:1) [resolution 12325,241]
[SA] active: 1152. $sum(i(s(X3)),X2) = $sum(i(X3),$sum(X2,1)) (2:12:1) [superposition 268,16]
[SA] active: 12650. ~$less(X15,$sum(X15,-3)) (7:5:1) [evaluation 12646]
[SA] active: 1159. $sum(i(s(X2)),X3) = $sum(1,$sum(X3,i(X2))) (2:12:1) [superposition 268,274]
[SA] active: 12651. ~$less(X2,$sum(-3,X2)) (7:5:1) [evaluation 12639]
[SA] active: 1161. $sum(i(s(X6)),X7) = $sum($sum(1,X7),i(X6)) (2:12:1) [superposition 268,16]
[SA] active: 12686. ~$less($sum(4,X13),X13) (7:5:1) [forward demodulation 12685,12327]
[SA] active: 1178. $sum(i(X8),X7) = $sum(i(s(X8)),$sum(-1,X7)) (2:12:1) [evaluation 1155]
[SA] active: 14396. $less(-10,sK0) (13:3:1) [evaluation 14395]
[SA] active: 1236. $less(0,X42) | $less(X42,alength0) | alength0 = X42 | alength0 = 0 (2:12:2) [resolution 203,206]
[SA] active: 14400. $less(-10,alength0) (14:3:1) [evaluation 14399]
[SA] active: 1238. $less(0,X46) | $less(X46,sK0) | sK0 = X46 | 0 = sK0 (2:12:2) [resolution 203,205]
[SA] active: 14613. ~$less(sK0,-10) (14:3:1) [resolution 14396,2395]
[SA] active: 516. $less(X5,$sum(X4,$uminus(X3))) | ~$less(X5,0) | ~$less(X3,X4) (2:12:1) T [resolution 240,22]
[SA] active: 14707. ~$less(alength0,-10) (15:3:1) [resolution 14400,2395]
[SA] active: 554. ~$less(X2,$sum(X0,$uminus(X1))) | ~$less(X0,X1) | $less(X2,0) (2:12:1) T [resolution 247,22]
[SA] active: 12688. $less(X14,$sum(4,X14)) (7:5:1) [forward demodulation 12687,12327]
[SA] active: 933. $less(X37,$sum(X38,$uminus(X39))) | $less(-1,X37) | ~$less(X39,X38) (2:12:1) T [evaluation 932]
[SA] active: 13082. $less($sum(-3,sK0),alength0) (8:5:1) [forward demodulation 13029,16]
[SA] active: 947. $sum(X9,X10) = $sum($sum($uminus(X8),X9),$sum(X8,X10)) (2:12:1) T [superposition 266,288]
[SA] active: 14327. ~$less($sum(X2,4),X2) (8:5:1) [resolution 12686,241]
[SA] active: 975. ~$less(X47,$sum(X46,$sum(X45,1))) | ~$less($sum(X45,X46),X47) (2:12:1) T [superposition 27,266]
[SA] active: 14385. ~$less(X15,$sum(X15,-4)) (8:5:1) [evaluation 14381]
[SA] active: 977. $less(X52,$sum(X51,$sum(X50,1))) | $less($sum(X50,X51),X52) (2:12:1) T [superposition 25,266]
[SA] active: 14386. ~$less(X2,$sum(-4,X2)) (8:5:1) [evaluation 14374]
[SA] active: 1005. $sum(X10,X9) = $sum($sum($uminus(X8),X9),$sum(X10,X8)) (2:12:1) T [superposition 274,288]
[SA] active: 15043. $less($sum(X15,-4),X15) (8:5:1) [evaluation 15039]
[SA] active: 1054. ~$less(X42,$sum(X41,$sum(1,X40))) | ~$less($sum(X40,X41),X42) (2:12:1) T [superposition 160,274]
[SA] active: 15044. $less($sum(-4,X2),X2) (8:5:1) [evaluation 15032]
[SA] active: 1055. $less(X45,$sum(X44,$sum(1,X43))) | $less($sum(X43,X44),X45) (2:12:1) T [superposition 149,274]
[SA] active: 15055. $less(X23,$sum(X23,4)) (8:5:1) [forward demodulation 15054,1061]
[SA] active: 1061. $sum(X62,X63) = $sum(X61,$sum(X63,$sum($uminus(X61),X62))) (2:12:1) T [superposition 288,274]
[SA] active: 15147. ~$less(alength0,$sum(-3,sK0)) (9:5:1) [resolution 13082,2395]
[SA] active: 1080. $sum(X62,X63) = $sum(X61,$sum(X62,$sum(X63,$uminus(X61)))) (2:12:1) T [superposition 288,274]
[SA] active: 16319. $less($sum(-4,sK0),alength0) (9:5:1) [forward demodulation 16264,16]
[SA] active: 1094. $sum(X10,X9) = $sum(X8,$sum($uminus(X8),$sum(X9,X10))) (2:12:1) T [forward demodulation 1026,17]
[SA] active: 17467. ~$less(alength0,$sum(-4,sK0)) (10:5:1) [resolution 16319,2395]
[SA] active: 1282. ~$less($sum(X13,X14),$sum(X15,$sum(X14,1))) | ~$less(X15,X13) (2:12:1) T [resolution 281,24]
[SA] active: 6985. ~$less(2,sK0) | $less(0,alength0) (5:6:1) [resolution 3679,2395]
[SA] active: 1284. ~$less($sum(X19,X20),$sum(X21,X19)) | ~$less(X21,$sum(X20,1)) (2:12:1) T [resolution 281,248]
[SA] active: 7026. ~$less(3,X0) | $less(1,X0) (5:6:1) [resolution 3696,2395]
[SA] active: 1291. ~$less(X1,$sum(X2,$sum(1,X0))) | ~$less($sum(X2,X0),X1) (2:12:1) T [superposition 281,16]
[SA] active: 7036. ~$less(X1,1) | $less(X1,3) (5:6:1) [resolution 3696,2395]
[SA] active: 1307. ~$less($sum(X16,X17),$sum(X17,$sum(1,X18))) | ~$less(X18,X16) (2:12:1) T [forward demodulation 1283,17]
[SA] active: 7285. ~$less(X0,-2) | $less(X0,1) (5:6:1) [resolution 6345,2395]
[SA] active: 1359. $less(X1,$sum(X2,$sum(1,X0))) | $less($sum(X2,X0),X1) (2:12:1) T [superposition 283,16]
[SA] active: 7323. $less(-2,X0) | $less(X0,2) (5:6:2) [resolution 6345,3537]
[SA] active: 1505. ~$less(X27,X25) | $less(0,$sum(X26,$uminus(X27))) | ~$less(X25,X26) (2:12:1) T [resolution 238,240]
[SA] active: 7335. ~$less(1,X1) | $less(-2,X1) (5:6:1) [resolution 6345,2395]
[SA] active: 1532. $less(X49,$sum(X50,X48)) | ~$less($uminus(X48),X50) | ~$less(X49,0) (2:12:1) T [superposition 238,109]
[SA] active: 7346. $less(-3,X10) | $less(X10,0) (5:6:2) [evaluation 7345]
[SA] active: 1827. ~$less($sum(X9,X10),$sum(X8,X10)) | ~$less(X8,$sum(1,X9)) (2:12:1) T [resolution 276,160]
[SA] active: 7876. ~$less(0,X23) | $less(-3,X23) (5:6:1) T [evaluation 7875]
[SA] active: 1910. ~$less(X14,$sum(X12,$sum(X13,X14))) | ~$less($sum(X12,X13),1) (2:12:1) T [resolution 277,160]
[SA] active: 7880. $less(-3,X21) | $less(X21,1) (5:6:2) T [evaluation 7879]
[SA] active: 886. $less(alength0,$sum(X27,1)) | $less(X27,i('$n')) | alength0 = i('$n') (2:13:1) [resolution 194,204]
[SA] active: 8422. ~$less(-1,X5) | $less(-3,X5) (5:6:1) T [resolution 6507,2395]
[SA] active: 1151. $sum(i(X1),i(s(X0))) = $sum(i(s(X1)),i(X0)) (2:13:1) [superposition 268,127]
[SA] active: 8432. ~$less(X2,-3) | $less(X2,-1) (5:6:1) T [resolution 6507,2395]
[SA] active: 1293. ~$less(X7,$sum(X8,i(s(X6)))) | ~$less($sum(X8,i(X6)),X7) (2:13:1) [superposition 281,75]
[SA] active: 8474. ~$less(X15,-3) | $less(X15,-2) (5:6:1) T [evaluation 8473]
[SA] active: 1361. $less(X7,$sum(X8,i(s(X6)))) | $less($sum(X8,i(X6)),X7) (2:13:1) [superposition 283,75]
[SA] active: 8488. $less(-4,X14) | $less(X14,-2) (5:6:2) T [evaluation 8487]
[SA] active: 1460. $less($sum(X21,1),alength0) | ~$less(X21,sK0) | $sum(X21,1) = sK0 (2:13:1) [resolution 207,199]
[SA] active: 9041. ~$less(-2,X6) | $less(-3,X6) (5:6:1) T [resolution 6624,2395]
[SA] active: 1599. $less($uminus(i(s(X26))),$sum(X27,-1)) | ~$less($uminus(i(X26)),X27) (2:13:1) [evaluation 1579]
[SA] active: 9053. $less(-4,X15) | $less(X15,-3) (5:6:2) T [evaluation 9052]
[SA] active: 1600. $less($uminus(i(s(X21))),$sum(X22,$uminus(i(X21)))) | ~$less(-1,X22) (2:13:1) [evaluation 1577]
[SA] active: 3808. ~$less(3,sK0) | $less(2,alength0) (6:6:1) [resolution 2395,729]
[SA] active: 1669. $less($sum(X27,-1),$uminus(i(s(X26)))) | ~$less(X27,$uminus(i(X26))) (2:13:1) [evaluation 1660]
[SA] active: 9669. $less(1,X6) | $less(X6,4) (6:6:2) [evaluation 9668]
[SA] active: 1670. $less($sum(X22,$uminus(i(X21))),$uminus(i(s(X21)))) | ~$less(X22,-1) (2:13:1) [evaluation 1658]
[SA] active: 10477. ~$less(3,sK0) | $less(1,alength0) (6:6:1) [resolution 6189,2395]
[SA] active: 404. $less($sum(1,X4),X5) | ~$less(X4,X5) | $sum(1,X4) = X5 (2:13:1) T [resolution 160,23]
[SA] active: 10761. ~$less(4,X2) | $less(2,X2) (6:6:1) [resolution 6210,2395]
[SA] active: 521. $less(0,$sum(X5,$sum(X6,$uminus(X7)))) | ~$less(X7,$sum(X5,X6)) (2:13:1) T [superposition 240,17]
[SA] active: 10771. ~$less(X1,3) | $less(X1,4) (6:6:1) [resolution 6210,512]
[SA] active: 559. $less($sum(X5,$sum(X6,$uminus(X7))),0) | ~$less($sum(X5,X6),X7) (2:13:1) T [superposition 247,17]
[SA] active: 10774. ~$less(X4,2) | $less(X4,4) (6:6:1) [resolution 6210,2395]
[SA] active: 606. $sum($sum(X18,$uminus(X19)),$sum(X19,$sum($uminus(X18),X20))) = X20 (2:13:1) T [forward demodulation 598,17]
[SA] active: 10796. $less(sK0,4) | $less(1,alength0) (6:6:1) [resolution 6210,3807]
[SA] active: 664. $sum($sum($uminus(X23),X24),$sum($uminus(X24),$sum(X23,X25))) = X25 (2:13:1) T [forward demodulation 647,17]
[SA] active: 10797. $less(sK0,4) | $less(2,alength0) (6:6:1) [resolution 6210,199]
[SA] active: 700. $less($sum(X19,$uminus(X18)),$uminus($sum(X17,X18))) | ~$less(X19,$uminus(X17)) (2:13:1) T [superposition 24,220]
[SA] active: 10834. ~$less(4,X9) | $less(3,X9) (6:6:1) [evaluation 10833]
[SA] active: 701. $less($uminus($sum(X20,X21)),$sum(X22,$uminus(X21))) | ~$less($uminus(X20),X22) (2:13:1) T [superposition 24,220]
[SA] active: 10915. ~$less(X5,-1) | $less(X5,2) (6:6:1) [resolution 6746,2395]
[SA] active: 746. ~$less(X7,$sum(X8,X5)) | ~$less(X5,X6) | $less(X7,$sum(X6,X8)) (2:13:1) T [resolution 241,22]
[SA] active: 10953. $less(-1,X0) | $less(X0,3) (6:6:2) [resolution 6746,3538]
[SA] active: 767. $less($sum($uminus(X24),X25),$uminus($sum(X23,X24))) | ~$less(X25,$uminus(X23)) (2:13:1) T [superposition 241,220]
[SA] active: 10973. ~$less(2,X3) | $less(-1,X3) (6:6:1) [resolution 6746,2395]
[SA] active: 768. $less($sum($uminus(X27),X28),$uminus($sum(X27,X26))) | ~$less(X28,$uminus(X26)) (2:13:1) T [superposition 241,19]
[SA] active: 11642. $less(X0,3) | $less(-2,X0) (6:6:2) [resolution 6866,6359]
[SA] active: 790. ~$less(X7,$sum(X5,X8)) | ~$less(X5,X6) | $less(X7,$sum(X8,X6)) (2:13:1) T [resolution 248,22]
[SA] active: 11648. ~$less(X7,0) | $less(X7,3) (6:6:1) [resolution 6866,2395]
[SA] active: 800. $less($uminus($sum(X23,X24)),$sum($uminus(X24),X25)) | ~$less($uminus(X23),X25) (2:13:1) T [superposition 248,220]
[SA] active: 11709. ~$less(3,X4) | $less(0,X4) (6:6:1) [resolution 6866,2395]
[SA] active: 801. $less($uminus($sum(X27,X26)),$sum($uminus(X27),X28)) | ~$less($uminus(X26),X28) (2:13:1) T [superposition 248,19]
[SA] active: 11843. $less(0,X0) | $less(X0,4) (6:6:2) [resolution 6877,6210]
[SA] active: 880. $less(X13,$sum(X14,1)) | $less(X14,$sum(X15,1)) | $less(X15,X13) (2:13:3) T [resolution 194,25]
[SA] active: 12205. ~$less(3,sK0) | $less(0,alength0) (6:6:1) [resolution 6984,2395]
[SA] active: 883. $less(X20,$sum(X21,1)) | $less(X21,$sum(1,X22)) | $less(X22,X20) (2:13:3) T [resolution 194,149]
[SA] active: 18033. $less(sK0,4) | $less(0,alength0) (6:6:1) [resolution 6985,6210]
[SA] active: 981. $less($sum(X62,$sum(X61,$uminus(X63))),0) | ~$less($sum(X61,X62),X63) (2:13:1) T [superposition 247,266]
[SA] active: 18976. ~$less(X2,-2) | $less(X2,2) (6:6:1) [resolution 7323,2395]
[SA] active: 982. $less(0,$sum(X65,$sum(X64,$uminus(X66)))) | ~$less(X66,$sum(X64,X65)) (2:13:1) T [superposition 240,266]
[SA] active: 19048. ~$less(2,X3) | $less(-2,X3) (6:6:1) [resolution 7323,2395]
[SA] active: 1466. $less($sum(1,X0),X1) | $sum(X0,1) = X1 | ~$less(X0,X1) (2:13:1) T [superposition 207,16]
[SA] active: 19066. ~$less(X4,1) | $less(X4,4) (6:6:1) [evaluation 19065]
[SA] active: 1508. ~$less(X36,X34) | $less($sum(X36,X37),$sum(X35,X37)) | ~$less(X34,X35) (2:13:1) T [resolution 238,24]
[SA] active: 19488. ~$less(X2,-3) | $less(X2,0) (6:6:1) [resolution 7346,2395]
[SA] active: 1509. ~$less(X41,X38) | $less($sum(X40,X41),$sum(X39,X40)) | ~$less(X38,X39) (2:13:1) T [resolution 238,241]
[SA] active: 19533. $less(-3,X0) | $less(X0,2) (6:6:2) [resolution 7346,6887]
[SA] active: 1510. $less($sum(X44,X42),$sum(X43,X45)) | ~$less(X42,X43) | ~$less(X44,X45) (2:13:1) T [resolution 238,248]
[SA] active: 19748. $less(-3,X0) | $less(X0,3) (6:6:2) [resolution 7876,6866]
[SA] active: 1603. $less($uminus(X9),$sum(X10,$uminus(X8))) | ~$less($sum($uminus(X9),X8),X10) (2:13:1) T [forward demodulation 1573,217]
[SA] active: 20703. $less(-4,X0) | $less(X0,1) (6:6:2) [resolution 8488,7285]
[SA] active: 1679. $less($sum(X10,$uminus(X8)),$uminus(X9)) | ~$less(X10,$sum($uminus(X9),X8)) (2:13:1) T [forward demodulation 1654,217]
[SA] active: 19470. ~$less(1,X2) | ~$less(X2,0) (6:6:1) T [evaluation 19356]
[SA] active: 1742. $sum($uminus(X6),X7) = $sum(X5,$sum($uminus($sum(X6,X5)),X7)) (2:13:1) T [superposition 288,271]
[SA] active: 19823. ~$less(0,X26) | $less(-4,X26) (6:6:1) T [evaluation 19822]
[SA] active: 1783. $sum($uminus(X22),X21) = $sum($uminus($sum(X20,X22)),$sum(X20,X21)) (2:13:1) T [forward demodulation 1730,26]
[SA] active: 19825. $less(-4,X21) | $less(X21,0) (6:6:2) T [evaluation 19824]
[SA] active: 1843. $less(0,$sum(X49,$sum(X50,X48))) | ~$less($uminus(X48),$sum(X49,X50)) (2:13:1) T [superposition 276,109]
[SA] active: 20033. ~$less(X2,-3) | $less(X2,1) (6:6:1) T [resolution 7880,2395]
[SA] active: 1952. $less($sum(X49,$sum(X50,X48)),0) | ~$less($sum(X49,X50),$uminus(X48)) (2:13:1) T [superposition 277,109]
[SA] active: 20226. ~$less(-1,X15) | $less(-4,X15) (6:6:1) T [evaluation 20225]
[SA] active: 482. $less(i(sK2(X3)),i(s('$n'))) | a(zero,X3) = a('$n',X3) (2:14:1) [resolution 133,73]
[SA] active: 20228. $less(-4,X10) | $less(X10,-1) (6:6:2) T [evaluation 20227]
[SA] active: 500. $less(i(s(X2)),X3) | ~$less(i(X2),X3) | i(s(X2)) = X3 (2:14:1) [resolution 163,23]
[SA] active: 20654. ~$less(X2,-4) | $less(X2,-2) (6:6:1) T [resolution 8488,2395]
[SA] active: 818. ~$less(X5,i(s(X3))) | ~$less(i(X3),X4) | $less(X5,$sum(X4,1)) (2:14:1) [resolution 245,22]
[SA] active: 20740. ~$less(-2,X7) | $less(-4,X7) (6:6:1) T [resolution 8488,2395]
[SA] active: 819. $less(i(X0),$sum(X1,1)) | ~$less(i(p(X0)),X1) | zero = X0 (2:14:1) [superposition 245,64]
[SA] active: 20760. ~$less(X8,-4) | $less(X8,-3) (6:6:1) T [evaluation 20759]
[SA] active: 840. ~$less(X4,$sum(X2,1)) | ~$less(X2,i(X3)) | $less(X4,i(s(X3))) (2:14:1) [resolution 252,22]
[SA] active: 21009. ~$less(-3,X6) | $less(-4,X6) (6:6:1) T [resolution 9053,2395]
[SA] active: 846. $less($sum(X1,1),i(X0)) | ~$less(X1,i(p(X0))) | zero = X0 (2:14:1) [superposition 252,64]
[SA] active: 2875. ~$less(sK0,X48) | $less(-2,X48) (7:6:1) [evaluation 2874]
[SA] active: 877. $less(X6,$sum(X7,1)) | $less(X7,i(s(X8))) | $less(i(X8),X6) (2:14:3) [resolution 194,152]
[SA] active: 2889. ~$less(alength0,X27) | $less(-2,X27) (7:6:1) [evaluation 2888]
[SA] active: 1426. ~'Sub'(zero,'$n') | v0 = a('$n',0) | $less(i(sK1(zero,0)),alength0) (2:14:1) [resolution 326,91]
[SA] active: 19469. ~$less(2,X3) | ~$less(X3,0) (7:6:1) [evaluation 19357]
[SA] active: 1504. ~$less(i(X24),X22) | $less(i(s(X24)),$sum(X23,1)) | ~$less(X22,X23) (2:14:1) [resolution 238,245]
[SA] active: 21142. $less(X0,4) | $less(-2,X0) (7:6:2) [resolution 9669,7335]
[SA] active: 1529. $less(X36,$sum(X37,i(X35))) | ~$less(1,X37) | ~$less(X36,i(s(X35))) (2:14:1) [superposition 238,127]
[SA] active: 21143. $less(X1,4) | $less(-1,X1) (7:6:2) [resolution 9669,6757]
[SA] active: 596. $less(0,$sum(X14,$sum(X13,$uminus(X12)))) | ~$less($sum(X12,$uminus(X13)),X14) (2:14:1) T [superposition 240,216]
[SA] active: 21228. ~$less(4,X2) | $less(1,X2) (7:6:1) [resolution 9669,2395]
[SA] active: 597. $less($sum(X17,$sum(X16,$uminus(X15))),0) | ~$less(X17,$sum(X15,$uminus(X16))) (2:14:1) T [superposition 247,216]
[SA] active: 21673. $less(X6,5) | $less(3,X6) (7:6:2) [evaluation 21672]
[SA] active: 601. $sum(X2,$uminus($sum(X4,X3))) = $uminus($sum(X4,$sum(X3,$uminus(X2)))) (2:14:1) T [forward demodulation 600,19]
[SA] active: 21849. $less(X7,5) | $less(2,X7) (7:6:2) [evaluation 21848]
[SA] active: 605. $sum(X4,$uminus($sum(X5,X3))) = $sum($sum(X4,$uminus(X3)),$uminus(X5)) (2:14:1) T [forward demodulation 592,601]
[SA] active: 22234. ~$less(4,sK0) | $less(1,alength0) (7:6:1) [resolution 10796,2395]
[SA] active: 645. $less(0,$sum(X19,$sum($uminus(X18),X17))) | ~$less($sum($uminus(X17),X18),X19) (2:14:1) T [superposition 240,217]
[SA] active: 22661. ~$less(4,sK0) | $less(2,alength0) (7:6:1) [resolution 10797,2395]
[SA] active: 646. $less($sum(X22,$sum($uminus(X21),X20)),0) | ~$less(X22,$sum($uminus(X20),X21)) (2:14:1) T [superposition 247,217]
[SA] active: 23216. ~$less(X6,-1) | $less(X6,3) (7:6:1) [resolution 10953,2395]
[SA] active: 699. $sum($uminus($sum(X14,X15)),X16) = $sum($uminus(X14),$sum($uminus(X15),X16)) (2:14:1) T [superposition 17,220]
[SA] active: 23301. ~$less(3,X3) | $less(-1,X3) (7:6:1) [resolution 10953,2395]
[SA] active: 706. $sum(X3,$uminus($sum(X4,X2))) = $uminus($sum(X2,$sum($uminus(X3),X4))) (2:14:1) T [forward demodulation 705,19]
[SA] active: 23780. ~$less(3,X3) | $less(-2,X3) (7:6:1) [resolution 11642,2395]
[SA] active: 707. $sum(X1,$uminus($sum(X2,X0))) = $sum($uminus(X2),$sum(X1,$uminus(X0))) (2:14:1) T [backward demodulation 706,604]
[SA] active: 23790. ~$less(X3,-2) | $less(X3,3) (7:6:1) [resolution 11642,2395]
[SA] active: 709. $sum($uminus($sum(X6,X7)),X5) = $sum($sum($uminus(X6),X5),$uminus(X7)) (2:14:1) T [forward demodulation 708,217]
[SA] active: 23849. ~$less(X5,1) | $less(X5,5) (7:6:1) [evaluation 23848]
[SA] active: 711. $sum(X5,$uminus($sum(X6,X7))) = $sum($uminus(X7),$sum($uminus(X6),X5)) (2:14:1) T [forward demodulation 685,706]
[SA] active: 24542. $less(X0,4) | $less(-3,X0) (7:6:2) [resolution 11843,7876]
[SA] active: 754. $less($sum(X17,$sum(X18,X19)),$sum(X20,$sum(X17,X18))) | ~$less(X19,X20) (2:14:1) T [superposition 241,17]
[SA] active: 24548. ~$less(X6,0) | $less(X6,4) (7:6:1) [resolution 11843,2395]
[SA] active: 765. $less($sum(X19,X20),$sum(X17,$sum(X18,X19))) | ~$less(X20,$sum(X17,X18)) (2:14:1) T [superposition 241,17]
[SA] active: 24635. ~$less(4,X2) | $less(0,X2) (7:6:1) [resolution 11843,2395]
[SA] active: 798. $less($sum(X17,$sum(X18,X19)),$sum(X19,X20)) | ~$less($sum(X17,X18),X20) (2:14:1) T [superposition 248,17]
[SA] active: 25151. ~$less(4,sK0) | $less(0,alength0) (7:6:1) [resolution 18033,2395]
[SA] active: 809. $less($sum(X20,$sum(X17,X18)),$sum(X17,$sum(X18,X19))) | ~$less(X20,X19) (2:14:1) T [superposition 248,17]
[SA] active: 25510. ~$less(2,X13) | $less(-3,X13) (7:6:1) [evaluation 25509]
[SA] active: 899. $less($sum(X51,$uminus(X52)),$sum(X53,1)) | $less(X53,0) | ~$less(X51,X52) (2:14:1) T [resolution 194,247]
[SA] active: 25708. $less(X7,5) | $less(1,X7) (7:6:2) [evaluation 25707]
[SA] active: 966. $less($sum(X16,X15),$sum(X14,$sum(X13,X15))) | ~$less(X16,$sum(X13,X14)) (2:14:1) T [superposition 24,266]
[SA] active: 26084. ~$less(X0,-3) | $less(X0,2) (7:6:1) [resolution 19533,2395]
[SA] active: 967. $less($sum(X18,$sum(X17,X19)),$sum(X20,X19)) | ~$less($sum(X17,X18),X20) (2:14:1) T [superposition 24,266]
[SA] active: 26345. ~$less(X0,-3) | $less(X0,3) (7:6:1) [resolution 19748,2395]
[SA] active: 968. $less($sum(X22,$sum(X21,X23)),$sum(X24,$sum(X21,X22))) | ~$less(X23,X24) (2:14:1) T [superposition 241,266]
[SA] active: 26431. ~$less(3,X1) | $less(-3,X1) (7:6:1) [resolution 19748,2395]
[SA] active: 969. $less($sum(X27,X28),$sum(X26,$sum(X25,X27))) | ~$less(X28,$sum(X25,X26)) (2:14:1) T [superposition 241,266]
[SA] active: 26449. ~$less(X2,1) | $less(X2,6) (7:6:1) [evaluation 26448]
[SA] active: 970. $less($sum(X32,$sum(X29,X30)),$sum(X30,$sum(X29,X31))) | ~$less(X32,X31) (2:14:1) T [superposition 248,266]
[SA] active: 26524. ~$less(X0,-4) | $less(X0,1) (7:6:1) [resolution 20703,2395]
[SA] active: 971. $less($sum(X34,$sum(X33,X35)),$sum(X35,X36)) | ~$less($sum(X33,X34),X36) (2:14:1) T [superposition 248,266]
[SA] active: 26579. $less(-4,X0) | $less(X0,4) (7:6:2) [resolution 20703,19066]
[SA] active: 980. $sum(X60,$uminus($sum(X58,X59))) = $uminus($sum(X59,$sum(X58,$uminus(X60)))) (2:14:1) T [superposition 216,266]
[SA] active: 26580. $less(-4,X1) | $less(X1,3) (7:6:2) [resolution 20703,7036]
[SA] active: 1014. $sum(X37,$uminus($sum(X35,X36))) = $sum($uminus(X36),$sum(X37,$uminus(X35))) (2:14:1) T [superposition 274,220]
[SA] active: 26629. ~$less(X4,-4) | $less(X4,0) (7:6:1) [evaluation 26628]
[SA] active: 1035. $sum(X37,$uminus($sum(X35,X36))) = $sum($uminus(X35),$sum($uminus(X36),X37)) (2:14:1) T [superposition 274,220]
[SA] active: 26765. ~$less(1,X8) | ~$less(X8,-1) (7:6:1) T [evaluation 26764]
[SA] active: 1047. $less($sum(X15,$sum(X13,X14)),$sum(X14,$sum(X12,X13))) | ~$less(X15,X12) (2:14:1) T [superposition 24,274]
[SA] active: 26771. ~$less(0,X5) | ~$less(X5,-1) (7:6:1) T [evaluation 26770]
[SA] active: 1048. $less($sum(X18,$sum(X16,X17)),$sum(X19,$sum(X17,X18))) | ~$less(X16,X19) (2:14:1) T [superposition 24,274]
[SA] active: 27357. ~$less(0,X28) | $less(-5,X28) (7:6:1) T [evaluation 27356]
[SA] active: 1049. $less($sum(X22,$sum(X20,X21)),$sum(X23,X20)) | ~$less($sum(X21,X22),X23) (2:14:1) T [superposition 241,274]
[SA] active: 27361. $less(-5,X26) | $less(X26,0) (7:6:2) T [evaluation 27360]
[SA] active: 1051. $less($sum(X31,X28),$sum(X30,$sum(X28,X29))) | ~$less(X31,$sum(X29,X30)) (2:14:1) T [superposition 248,274]
[SA] active: 28196. ~$less(X4,-4) | $less(X4,-1) (7:6:1) T [evaluation 28195]
[SA] active: 1060. $sum($uminus($sum(X59,X60)),X58) = $uminus($sum(X60,$sum($uminus(X58),X59))) (2:14:1) T [superposition 217,274]
[SA] active: 28724. ~$less(-1,X9) | $less(-5,X9) (7:6:1) T [evaluation 28723]
[SA] active: 1079. $sum($uminus($sum(X59,X60)),X58) = $uminus($sum(X59,$sum(X60,$uminus(X58)))) (2:14:1) T [superposition 217,274]
[SA] active: 28728. $less(-5,X7) | $less(X7,-1) (7:6:2) T [evaluation 28727]
[SA] active: 1101. $less($sum(X25,$sum(X26,X27)),$sum(X26,$sum(X24,X25))) | ~$less(X27,X24) (2:14:1) T [forward demodulation 1050,17]
[SA] active: 29098. ~$less(-2,X12) | $less(-5,X12) (7:6:1) T [evaluation 29097]
[SA] active: 1102. $less($sum(X34,$sum(X32,X33)),$sum(X33,$sum(X34,X35))) | ~$less(X32,X35) (2:14:1) T [forward demodulation 1052,17]
[SA] active: 29102. $less(-5,X10) | $less(X10,-2) (7:6:2) T [evaluation 29101]
[SA] active: 1112. $less($sum(X25,$sum(X26,X27)),$sum(X25,$sum(X26,X24))) | ~$less(X27,X24) (2:14:1) T [forward demodulation 1069,17]
[SA] active: 29370. ~$less(-3,X13) | $less(-5,X13) (7:6:1) T [evaluation 29369]
[SA] active: 1523. $less(X13,$sum(X14,$sum($uminus(X11),X12))) | ~$less(X11,X14) | ~$less(X13,X12) (2:14:1) T [superposition 238,288]
[SA] active: 29374. $less(-5,X11) | $less(X11,-3) (7:6:2) T [evaluation 29373]
[SA] active: 1595. $less($uminus($sum(X17,X18)),$uminus($sum(X16,X17))) | ~$less($uminus(X18),$uminus(X16)) (2:14:1) T [superposition 244,220]
[SA] active: 13896. ~$less(1,X0) | $less(-3,X0) (8:6:1) [resolution 12650,1498]
[SA] active: 1596. $less($uminus($sum(X20,X21)),$uminus($sum(X20,X19))) | ~$less($uminus(X21),$uminus(X19)) (2:14:1) T [superposition 244,19]
[SA] active: 19468. ~$less(3,X4) | ~$less(X4,0) (8:6:1) [evaluation 19358]
[SA] active: 1648. $less($uminus($sum(X16,X17)),$uminus($sum(X17,X18))) | ~$less($uminus(X16),$uminus(X18)) (2:14:1) T [superposition 251,220]
[SA] active: 29837. ~$less(2,X7) | ~$less(X7,-1) (8:6:1) [evaluation 29836]
[SA] active: 1727. $sum($uminus($sum(X12,X14)),X13) = $sum($uminus(X14),$sum(X13,$uminus(X12))) (2:14:1) T [superposition 271,16]
[SA] active: 29843. ~$less(0,X4) | ~$less(X4,-2) (8:6:1) [evaluation 29842]
[SA] active: 1736. $sum($uminus($sum(X8,X7)),X9) = $sum($uminus(X8),$sum(X9,$uminus(X7))) (2:14:1) T [superposition 271,274]
[SA] active: 29976. ~$less(4,X0) | $less(-2,X0) (8:6:1) [resolution 21142,2395]
[SA] active: 1780. $sum($uminus(X4),$uminus($sum(X2,X3))) = $uminus($sum(X3,$sum(X2,X4))) (2:14:1) T [forward demodulation 1724,19]
[SA] active: 29986. ~$less(X2,-2) | $less(X2,4) (8:6:1) [resolution 21142,2395]
[SA] active: 1781. $sum($uminus(X7),$uminus($sum(X6,X5))) = $uminus($sum(X6,$sum(X5,X7))) (2:14:1) T [forward demodulation 1725,19]
[SA] active: 30052. ~$less(4,X1) | $less(-1,X1) (8:6:1) [evaluation 30051]
[SA] active: 1847. $less($sum(X2,X1),$sum(X3,$sum(X1,X0))) | ~$less(X2,$sum(X3,X0)) (2:14:1) T [superposition 276,16]
[SA] active: 30175. ~$less(X4,-1) | $less(X4,4) (8:6:1) [resolution 21143,2395]
[SA] active: 1914. $less($sum(X2,$sum(X1,X0)),$sum(X3,X1)) | ~$less($sum(X2,X0),X3) (2:14:1) T [superposition 277,16]
[SA] active: 30606. ~$less(5,X0) | $less(3,X0) (8:6:1) [resolution 21673,2395]
[SA] active: 826. $less(i(s(X8)),$sum(X6,$sum(X7,1))) | ~$less(i(X8),$sum(X6,X7)) (2:15:1) [superposition 245,17]
[SA] active: 30613. $less(X0,5) | $less(0,X0) (8:6:2) [resolution 21673,11709]
[SA] active: 844. $less($sum(X6,$sum(X7,1)),i(s(X8))) | ~$less($sum(X6,X7),i(X8)) (2:15:1) [superposition 252,17]
[SA] active: 30617. ~$less(X3,3) | $less(X3,5) (8:6:1) [resolution 21673,2395]
[SA] active: 973. $less($sum(X40,$sum(X39,1)),i(s(X41))) | ~$less($sum(X39,X40),i(X41)) (2:15:1) [superposition 252,266]
[SA] active: 30644. $less(sK0,5) | $less(0,alength0) (8:6:1) [resolution 21673,12205]
[SA] active: 974. $less(i(s(X44)),$sum(X43,$sum(X42,1))) | ~$less(i(X44),$sum(X42,X43)) (2:15:1) [superposition 245,266]
[SA] active: 30645. $less(sK0,5) | $less(1,alength0) (8:6:1) [resolution 21673,10477]
[SA] active: 1168. $less($sum(X13,$sum(1,X12)),$sum(i(s(X11)),X12)) | ~$less(X13,i(X11)) (2:15:1) [superposition 24,268]
[SA] active: 30646. $less(sK0,5) | $less(2,alength0) (8:6:1) [resolution 21673,3808]
[SA] active: 1169. $less($sum(i(s(X14)),X15),$sum(X16,$sum(1,X15))) | ~$less(i(X14),X16) (2:15:1) [superposition 24,268]
[SA] active: 30647. $less(sK0,5) | $less(3,alength0) (8:6:1) [resolution 21673,199]
[SA] active: 1170. $less($sum(i(s(X17)),X18),$sum(X19,i(X17))) | ~$less($sum(1,X18),X19) (2:15:1) [superposition 241,268]
[SA] active: 30686. ~$less(5,X1) | $less(4,X1) (8:6:1) [evaluation 30685]
[SA] active: 1172. $less($sum(X25,i(X23)),$sum(i(s(X23)),X24)) | ~$less(X25,$sum(1,X24)) (2:15:1) [superposition 248,268]
[SA] active: 31099. ~$less(5,X0) | $less(2,X0) (8:6:1) [resolution 21849,2395]
[SA] active: 1190. $less($sum(1,$sum(X21,X22)),$sum(i(s(X20)),X21)) | ~$less(X22,i(X20)) (2:15:1) [forward demodulation 1171,17]
[SA] active: 31106. ~$less(X0,0) | $less(X0,5) (8:6:1) [resolution 21849,19469]
[SA] active: 1191. $less($sum(i(s(X26)),X27),$sum(1,$sum(X27,X28))) | ~$less(i(X26),X28) (2:15:1) [forward demodulation 1173,17]
[SA] active: 31107. $less(X1,5) | $less(-2,X1) (8:6:2) [resolution 21849,19048]
[SA] active: 1840. $less(i(s(X35)),$sum(X36,$sum(X37,i(X35)))) | ~$less(1,$sum(X36,X37)) (2:15:1) [superposition 276,127]
[SA] active: 31108. $less(X2,5) | $less(-1,X2) (8:6:2) [resolution 21849,10973]
[SA] active: 1854. $less($sum(X29,1),$sum(X30,i(s(X28)))) | ~$less(X29,$sum(X30,i(X28))) (2:15:1) [superposition 276,75]
[SA] active: 31113. ~$less(X6,2) | $less(X6,5) (8:6:1) [resolution 21849,2395]
[SA] active: 1856. $less($sum(X36,i(X35)),$sum(X37,i(s(X35)))) | ~$less(X36,$sum(X37,1)) (2:15:1) [superposition 276,127]
[SA] active: 33982. $less(X0,5) | $less(-4,X0) (8:6:2) [resolution 23849,20703]
[SA] active: 1921. $less($sum(X29,i(s(X28))),$sum(X30,1)) | ~$less($sum(X29,i(X28)),X30) (2:15:1) [superposition 277,75]
[SA] active: 33983. $less(X1,5) | $less(-3,X1) (8:6:2) [resolution 23849,7880]
[SA] active: 1923. $less($sum(X36,i(s(X35))),$sum(X37,i(X35))) | ~$less($sum(X36,1),X37) (2:15:1) [superposition 277,127]
[SA] active: 34048. $less(X6,6) | $less(1,X6) (8:6:2) [evaluation 34047]
[SA] active: 1949. $less($sum(X36,$sum(X37,i(X35))),i(s(X35))) | ~$less($sum(X36,X37),1) (2:15:1) [superposition 277,127]
[SA] active: 34624. ~$less(4,X0) | $less(-3,X0) (8:6:1) [resolution 24542,2395]
[SA] active: 989. $sum(X19,$sum(X17,$sum(X18,X20))) = $sum($sum(X18,$sum(X17,X19)),X20) (2:15:1) T [forward demodulation 951,17]
[SA] active: 34633. ~$less(X1,-3) | $less(X1,4) (8:6:1) [resolution 24542,2395]
[SA] active: 991. $sum(X23,$sum(X21,$sum(X22,X24))) = $sum(X23,$sum(X22,$sum(X21,X24))) (2:15:1) T [forward demodulation 990,17]
[SA] active: 34698. ~$less(X3,1) | $less(X3,7) (8:6:1) [evaluation 34697]
[SA] active: 1006. $sum(X14,$sum(X12,$sum(X13,X11))) = $sum($sum(X12,X13),$sum(X14,X11)) (2:15:1) T [superposition 274,274]
[SA] active: 35671. ~$less(2,X8) | $less(-4,X8) (8:6:1) [evaluation 35670]
[SA] active: 1007. $sum(X18,$sum(X17,$sum(X15,X16))) = $sum($sum(X16,X17),$sum(X18,X15)) (2:15:1) T [superposition 274,274]
[SA] active: 35724. ~$less(5,X0) | $less(1,X0) (8:6:1) [resolution 25708,2395]
[SA] active: 1011. $sum(X27,$sum(X28,$sum(X25,X26))) = $sum(X28,$sum(X26,$sum(X25,X27))) (2:15:1) T [superposition 274,266]
[SA] active: 36916. $less(X0,6) | $less(-4,X0) (8:6:2) [resolution 26449,20703]
[SA] active: 1012. $sum(X31,$sum(X32,$sum(X29,X30))) = $sum(X32,$sum(X29,$sum(X30,X31))) (2:15:1) T [superposition 274,17]
[SA] active: 36917. $less(X1,6) | $less(-3,X1) (8:6:2) [resolution 26449,7880]
[SA] active: 1032. $sum(X28,$sum(X26,$sum(X25,X27))) = $sum($sum(X25,X26),$sum(X27,X28)) (2:15:1) T [superposition 274,266]
[SA] active: 36918. $less(X2,6) | $less(-2,X2) (8:6:2) [resolution 26449,6345]
[SA] active: 1033. $sum(X32,$sum(X29,$sum(X30,X31))) = $sum($sum(X29,X30),$sum(X31,X32)) (2:15:1) T [superposition 274,17]
[SA] active: 36919. $less(X3,6) | $less(-1,X3) (8:6:2) [resolution 26449,3590]
[SA] active: 1041. $sum(X18,$sum(X19,$sum(X20,X21))) = $sum(X20,$sum(X21,$sum(X18,X19))) (2:15:1) T [superposition 274,17]
[SA] active: 36920. $less(X4,6) | $less(0,X4) (8:6:2) [resolution 26449,378]
[SA] active: 1090. $sum(X15,$sum(X14,$sum(X16,X17))) = $sum(X17,$sum(X14,$sum(X15,X16))) (2:15:1) T [forward demodulation 1020,17]
[SA] active: 36942. $less(sK0,6) | $less(0,alength0) (8:6:1) [resolution 26449,393]
[SA] active: 1095. $sum(X11,$sum(X12,$sum(X13,X14))) = $sum(X14,$sum(X12,$sum(X13,X11))) (2:15:1) T [forward demodulation 1027,17]
[SA] active: 37709. ~$less(X0,-4) | $less(X0,4) (8:6:1) [resolution 26579,2395]
[SA] active: 1096. $sum(X15,$sum(X16,$sum(X17,X18))) = $sum(X18,$sum(X17,$sum(X15,X16))) (2:15:1) T [forward demodulation 1028,17]
[SA] active: 37812. ~$less(4,X0) | $less(-4,X0) (8:6:1) [resolution 26579,2395]
[SA] active: 1100. $sum(X8,$sum(X9,$sum(X10,X11))) = $sum($sum(X10,$sum(X8,X9)),X11) (2:15:1) T [forward demodulation 1046,17]
[SA] active: 38277. ~$less(X0,-4) | $less(X0,3) (8:6:1) [resolution 26580,2395]
[SA] active: 1834. $less(X12,$sum(X13,$sum(X14,$sum($uminus(X11),X12)))) | ~$less(X11,$sum(X13,X14)) (2:15:1) T [superposition 276,288]
[SA] active: 38381. ~$less(3,X1) | $less(-4,X1) (8:6:1) [resolution 26580,2395]
[SA] active: 1850. $less($sum(X13,$sum($uminus(X11),X12)),$sum(X14,X12)) | ~$less(X13,$sum(X14,X11)) (2:15:1) T [superposition 276,288]
[SA] active: 39584. ~$less(X0,-1) | $less(X0,5) (8:6:1) [resolution 26765,25708]
[SA] active: 1917. $less($sum(X13,X12),$sum(X14,$sum($uminus(X11),X12))) | ~$less($sum(X13,X11),X14) (2:15:1) T [superposition 277,288]
[SA] active: 40599. $less(-5,X0) | $less(X0,4) (8:6:2) [resolution 27357,11843]
[SA] active: 1943. $less($sum(X13,$sum(X14,$sum($uminus(X11),X12))),X12) | ~$less($sum(X13,X14),X11) (2:15:1) T [superposition 277,288]
[SA] active: 40600. $less(-5,X1) | $less(X1,3) (8:6:2) [resolution 27357,6866]
[SA] active: 1156. $sum(i(s(X11)),$sum(X9,X10)) = $sum(i(X11),$sum(X9,$sum(X10,1))) (2:16:1) [superposition 268,274]
[SA] active: 40601. $less(-5,X2) | $less(X2,2) (8:6:2) [resolution 27357,3621]
[SA] active: 1157. $sum(i(s(X14)),$sum(X12,X13)) = $sum(i(X14),$sum(X13,$sum(1,X12))) (2:16:1) [superposition 268,274]
[SA] active: 40996. ~$less(X0,-5) | $less(X0,0) (8:6:1) T [resolution 27361,2395]
[SA] active: 1176. $sum(X37,$sum(i(s(X35)),X36)) = $sum($sum(1,X36),$sum(X37,i(X35))) (2:16:1) [superposition 274,268]
[SA] active: 42735. ~$less(X0,-5) | $less(X0,-1) (8:6:1) T [resolution 28728,2395]
[SA] active: 1192. $sum(i(s(X29)),$sum(X30,X31)) = $sum($sum(1,X30),$sum(i(X29),X31)) (2:16:1) [forward demodulation 1174,17]
[SA] active: 43674. ~$less(X0,-5) | $less(X0,-2) (8:6:1) T [resolution 29102,2395]
[SA] active: 1193. $sum(X34,$sum(i(s(X32)),X33)) = $sum(i(X32),$sum(1,$sum(X33,X34))) (2:16:1) [forward demodulation 1175,17]
[SA] active: 44401. ~$less(X0,-5) | $less(X0,-3) (8:6:1) T [resolution 29374,2395]
[SA] active: 1908. ~$less($sum(X5,X6),$sum(X3,$sum(X4,$sum(X6,1)))) | ~$less($sum(X3,X4),X5) (2:16:1) T [resolution 277,281]
[SA] active: 2877. ~$less(sK0,X47) | $less(-3,X47) (9:6:1) [evaluation 2876]
[SA] active: 1580. $less($uminus($sum(X29,$sum(X28,X30))),$sum(X31,$uminus($sum(X28,X29)))) | ~$less($uminus(X30),X31) (2:17:1) T [superposition 244,266]
[SA] active: 15682. ~$less(1,X0) | $less(-4,X0) (9:6:1) [resolution 14385,1498]
[SA] active: 1581. $less($uminus($sum(X32,$sum(X33,X34))),$sum(X35,$uminus($sum(X32,X33)))) | ~$less($uminus(X34),X35) (2:17:1) T [superposition 244,17]
[SA] active: 19467. ~$less(4,X5) | ~$less(X5,0) (9:6:1) [evaluation 19359]
[SA] active: 1614. $less($sum(X4,$uminus($sum(X5,X3))),$sum(X6,$sum(X4,$uminus(X3)))) | ~$less($uminus(X5),X6) (2:17:1) T [forward demodulation 1613,706]
[SA] active: 24822. $less(X27,2) | $less(-4,X27) (9:6:2) [evaluation 24703]
[SA] active: 1616. $less($sum($uminus($sum(X8,X9)),X7),$sum(X10,$sum($uminus(X8),X7))) | ~$less($uminus(X9),X10) (2:17:1) T [forward demodulation 1615,217]
[SA] active: 24834. $less(-5,X27) | $less(X27,1) (9:6:2) [evaluation 24833]
[SA] active: 1661. $less($sum(X31,$uminus($sum(X28,X29))),$uminus($sum(X29,$sum(X28,X30)))) | ~$less(X31,$uminus(X30)) (2:17:1) T [superposition 251,266]
[SA] active: 49391. ~$less(5,X0) | $less(0,X0) (9:6:1) [resolution 30613,2395]
[SA] active: 1662. $less($sum(X35,$uminus($sum(X32,X33))),$uminus($sum(X32,$sum(X33,X34)))) | ~$less(X35,$uminus(X34)) (2:17:1) T [superposition 251,17]
[SA] active: 49400. $less(X6,5) | $less(-5,X6) (9:6:2) [resolution 30613,27357]
[SA] active: 1675. $less($sum(X5,$sum(X4,$uminus(X3))),$sum(X4,$uminus($sum(X6,X3)))) | ~$less(X5,$uminus(X6)) (2:17:1) T [forward demodulation 1674,706]
[SA] active: 49401. ~$less(X7,-2) | $less(X7,5) (9:6:1) [resolution 30613,29843]
[SA] active: 1677. $less($sum(X9,$sum($uminus(X8),X7)),$sum($uminus($sum(X8,X10)),X7)) | ~$less(X9,$uminus(X10)) (2:17:1) T [forward demodulation 1676,217]
[SA] active: 49634. ~$less(5,sK0) | $less(0,alength0) (9:6:1) [resolution 30644,2395]
[SA] active: 1808. $less($sum($uminus(X50),$sum(X51,X52)),$sum($uminus($sum(X50,X49)),X51)) | ~$less(X52,$uminus(X49)) (2:17:1) T [forward demodulation 1755,17]
[SA] active: 49736. ~$less(5,sK0) | $less(1,alength0) (9:6:1) [resolution 30645,2395]
[SA] active: 1812. $less($sum($uminus($sum(X62,X61)),X63),$sum($uminus(X62),$sum(X63,X64))) | ~$less($uminus(X61),X64) (2:17:1) T [forward demodulation 1758,17]
[SA] active: 49893. ~$less(5,sK0) | $less(2,alength0) (9:6:1) [resolution 30646,2395]
[SA] active: 1731. $sum($uminus($sum(X23,X26)),$sum(X24,X25)) = $sum($uminus(X26),$sum(X24,$sum(X25,$uminus(X23)))) (2:18:1) T [superposition 271,274]
[SA] active: 50051. ~$less(5,sK0) | $less(3,alength0) (9:6:1) [resolution 30647,2395]
[SA] active: 1732. $sum($uminus($sum(X27,X30)),$sum(X28,X29)) = $sum($uminus(X30),$sum(X29,$sum($uminus(X27),X28))) (2:18:1) T [superposition 271,274]
[SA] active: 50929. ~$less(5,X0) | $less(-2,X0) (9:6:1) [resolution 31107,2395]
[SA] active: 1762. $sum(X80,$sum($uminus($sum(X78,X77)),X79)) = $sum($sum($uminus(X78),X79),$sum(X80,$uminus(X77))) (2:18:1) T [superposition 274,271]
[SA] active: 51183. ~$less(5,X0) | $less(-1,X0) (9:6:1) [resolution 31108,2395]
[SA] active: 1768. $sum(X4,$sum($uminus($sum(X5,X3)),X6)) = $sum($sum(X4,$uminus(X3)),$sum($uminus(X5),X6)) (2:18:1) T [forward demodulation 1767,17]
[SA] active: 51522. ~$less(5,X0) | $less(-4,X0) (9:6:1) [resolution 33982,2395]
[SA] active: 1770. $sum(X7,$sum($uminus($sum(X8,X9)),X10)) = $sum($sum($uminus(X8),X7),$sum($uminus(X9),X10)) (2:18:1) T [forward demodulation 1769,17]
[SA] active: 51526. ~$less(X0,-4) | $less(X0,5) (9:6:1) [resolution 33982,2395]
[SA] active: 1775. $sum(X4,$sum($uminus($sum(X5,X3)),X6)) = $sum($uminus(X5),$sum(X4,$sum($uminus(X3),X6))) (2:18:1) T [forward demodulation 1774,17]
[SA] active: 51703. ~$less(5,X0) | $less(-3,X0) (9:6:1) [resolution 33983,2395]
[SA] active: 1782. $sum($uminus($sum(X9,$sum(X8,X11))),X10) = $sum($uminus(X11),$sum($uminus($sum(X9,X8)),X10)) (2:18:1) T [forward demodulation 1726,271]
[SA] active: 51709. ~$less(X2,-3) | $less(X2,5) (9:6:1) [resolution 33983,2395]
[SA] active: 1817. $sum(X76,$sum($uminus($sum(X74,X73)),X75)) = $sum($uminus(X73),$sum($uminus(X74),$sum(X75,X76))) (2:18:1) T [forward demodulation 1761,17]
[SA] active: 51905. ~$less(6,X0) | $less(1,X0) (9:6:1) [resolution 34048,2395]
[SA] active: 1835. $less($sum(X16,$sum(X17,X15)),$sum(X18,$sum(X19,$sum(X16,X17)))) | ~$less(X15,$sum(X18,X19)) (2:18:1) T [superposition 276,274]
[SA] active: 51908. ~$less(X0,2) | $less(X0,6) (9:6:1) [resolution 34048,511]
[SA] active: 1836. $less($sum(X22,$sum(X20,X21)),$sum(X23,$sum(X24,$sum(X21,X22)))) | ~$less(X20,$sum(X23,X24)) (2:18:1) T [superposition 276,274]
[SA] active: 51913. ~$less(X5,0) | $less(X5,6) (9:6:1) [resolution 34048,19470]
[SA] active: 1851. $less($sum(X18,$sum(X16,X17)),$sum(X19,$sum(X16,$sum(X17,X15)))) | ~$less(X18,$sum(X19,X15)) (2:18:1) T [superposition 276,274]
[SA] active: 51914. ~$less(X6,-1) | $less(X6,6) (9:6:1) [resolution 34048,26765]
[SA] active: 1852. $less($sum(X23,$sum(X21,X22)),$sum(X24,$sum(X22,$sum(X20,X21)))) | ~$less(X23,$sum(X24,X20)) (2:18:1) T [superposition 276,274]
[SA] active: 51941. $less(sK0,6) | $less(1,alength0) (9:6:1) [resolution 34048,199]
[SA] active: 1918. $less($sum(X18,$sum(X16,$sum(X17,X15))),$sum(X19,$sum(X16,X17))) | ~$less($sum(X18,X15),X19) (2:18:1) T [superposition 277,274]
[SA] active: 53010. ~$less(X5,0) | $less(X5,7) (9:6:1) [evaluation 53009]
[SA] active: 1919. $less($sum(X23,$sum(X22,$sum(X20,X21))),$sum(X24,$sum(X21,X22))) | ~$less($sum(X23,X20),X24) (2:18:1) T [superposition 277,274]
[SA] active: 54408. $less(X0,7) | $less(-4,X0) (9:6:2) [resolution 34698,20703]
[SA] active: 1944. $less($sum(X18,$sum(X19,$sum(X16,X17))),$sum(X16,$sum(X17,X15))) | ~$less($sum(X18,X19),X15) (2:18:1) T [superposition 277,274]
[SA] active: 54409. $less(X1,7) | $less(-3,X1) (9:6:2) [resolution 34698,7880]
[SA] active: 1945. $less($sum(X23,$sum(X24,$sum(X21,X22))),$sum(X22,$sum(X20,X21))) | ~$less($sum(X23,X24),X20) (2:18:1) T [superposition 277,274]
% Refutation not found, non-redundant clauses discarded
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Refutation not found, non-redundant clauses discarded

% Memory used [KB]: 65244
% Time elapsed: 2.014 s
% ------------------------------
% ------------------------------
