preprocessing started
[PP] input: 19. ~$lesseq(1,2) [negated conjecture 18]
[PP] input: 17. ! [X0 : $int] : X0 = X0 [input]
[PP] input: 16. $lesseq(0,1) [input]
[PP] input: 15. ! [X0 : $int] : X0 = X0 [input]
[PP] input: 14. ! [X0 : $int,X2 : $int,X1 : $int] : $sum($sum(X0,X1),X2) = $sum($sum(X0,X2),X1) [input]
[PP] input: 13. ! [X0 : 'Nat()'] : test(l9(zero)) = test(l9(X0)) [input]
[PP] input: 12. ! [X0 : 'Nat()'] : (('Sub'(s(X0),nl9) & test(l9(s(X0))) = $sum(test(l9(X0)),1)) => ! [X1 : 'Nat()',X2 : 'Nat()'] : (test(l9(X1)) = test(l9(X2)) => X1 = X2)) [input]
[PP] input: 11. ! [X0 : 'Nat()'] : (('Sub'(s(X0),nl9) & i(l9(s(X0))) = $sum(i(l9(X0)),1)) => ! [X1 : 'Nat()',X2 : 'Nat()'] : (i(l9(X1)) = i(l9(X2)) => X1 = X2)) [input]
[PP] input: 10. ! [X0 : $int,X1 : 'Nat()',X2 : $int] : (($lesseq(a(l9(zero),X2),X0) & $less(X0,a(l9(nl9),X2)) & a(l9(s(X1)),X2) = $sum(a(l9(X1),X2),1)) => ? [X3 : 'Nat()'] : (a(l9(X3),X2) = X0 & 'Sub'(X3,nl9))) [input]
[PP] input: 9. ! [X0 : $int,X1 : 'Nat()'] : (($lesseq(test(l9(zero)),X0) & $less(X0,test(l9(nl9))) & test(l9(s(X1))) = $sum(test(l9(X1)),1)) => ? [X2 : 'Nat()'] : (test(l9(X2)) = X0 & 'Sub'(X2,nl9))) [input]
[PP] input: 8. ! [X0 : $int,X1 : 'Nat()'] : (($lesseq(i(l9(zero)),X0) & $less(X0,i(l9(nl9))) & i(l9(s(X1))) = $sum(i(l9(X1)),1)) => ? [X2 : 'Nat()'] : (i(l9(X2)) = X0 & 'Sub'(X2,nl9))) [input]
[PP] input: 7. $less(i(l9(zero)),alength) => ? [X0 : 'Nat()'] : s(X0) = nl9 [input]
[PP] input: 6. ! [X0 : $int] : (! [X1 : 'Nat()'] : ('Sub'(X1,nl9) => a(l9(X1),X0) = a(l9(s(X1)),X0)) => a(l9(zero),X0) = a(l9(nl9),X0)) [input]
[PP] input: 5. ! [X0 : 'Nat()'] : ('Sub'(X0,nl9) => test(l9(X0)) = test(l9(s(X0)))) => test(l9(zero)) = test(l9(nl9)) [input]
[PP] input: 4. ! [X0 : 'Nat()'] : ('Sub'(X0,nl9) => i(l9(X0)) = i(l9(s(X0)))) => i(l9(zero)) = i(l9(nl9)) [input]
[PP] input: 3. (test(l8) = 1 & ! [X8 : $int] : a(l7,X8) = a(l8,X8)) & (i(l9(zero)) = 0 & ! [X7 : $int] : a(l8,X7) = a(l9(zero),X7) & test(l8) = test(l9(zero))) & (! [X5 : 'Nat()'] : (i(l9(X5)) = i(l11(X5)) & test(l9(X5)) = test(l11(X5)) & ! [X6 : $int] : a(l9(X5),X6) = a(l11(X5),X6)) & ! [X2 : 'Nat()'] : ((v = a(l12(X2),i(l11(X2))) & ! [X4 : $int] : (~i(l11(X2)) = X4 => a(l11(X2),X4) = a(l12(X2),X4)) & test(l11(X2)) = test(l12(X2)) & i(l11(X2)) = i(l12(X2))) & (i(l9(s(X2))) = $sum(i(l12(X2)),1) & ! [X3 : $int] : a(l12(X2),X3) = a(l9(s(X2)),X3) & test(l12(X2)) = test(l9(s(X2))))) & ! [X1 : 'Nat()'] : ('Sub'(X1,nl9) => $less(i(l9(X1)),alength)) & ~$less(i(l9(nl9)),alength) & (i(main_end) = i(l9(nl9)) & test(main_end) = test(l9(nl9)) & ! [X0 : $int] : a(main_end,X0) = a(l9(nl9),X0))) [input]
[PP] input: 2. ! [X0 : 'Nat()',X1 : 'Nat()'] : ('Sub'(X0,X1) => 'Sub'(X0,s(X1))) [input]
[PP] input: 1. ! [X0 : 'Nat()'] : 'Sub'(X0,s(X0)) [input]
adding theory axioms
preprocess1 (rectify, simplify false true, flatten)
[PP] flatten in: 20. ~~$less(2,1) [evaluation 19]
[PP] flatten out: 40. $less(2,1) [flattening 20]
[PP] flatten in: 42. (test(l8) = 1 & ! [X0 : $int] : a(l7,X0) = a(l8,X0)) & (i(l9(zero)) = 0 & ! [X1 : $int] : a(l8,X1) = a(l9(zero),X1) & test(l8) = test(l9(zero))) & (! [X2 : 'Nat()'] : (i(l9(X2)) = i(l11(X2)) & test(l9(X2)) = test(l11(X2)) & ! [X3 : $int] : a(l9(X2),X3) = a(l11(X2),X3)) & ! [X4 : 'Nat()'] : ((v = a(l12(X4),i(l11(X4))) & ! [X5 : $int] : (~i(l11(X4)) = X5 => a(l11(X4),X5) = a(l12(X4),X5)) & test(l11(X4)) = test(l12(X4)) & i(l11(X4)) = i(l12(X4))) & (i(l9(s(X4))) = $sum(i(l12(X4)),1) & ! [X6 : $int] : a(l12(X4),X6) = a(l9(s(X4)),X6) & test(l12(X4)) = test(l9(s(X4))))) & ! [X7 : 'Nat()'] : ('Sub'(X7,nl9) => $less(i(l9(X7)),alength)) & ~$less(i(l9(nl9)),alength) & (i(main_end) = i(l9(nl9)) & test(main_end) = test(l9(nl9)) & ! [X8 : $int] : a(main_end,X8) = a(l9(nl9),X8))) [rectify 3]
[PP] flatten out: 43. test(l8) = 1 & ! [X0 : $int] : a(l7,X0) = a(l8,X0) & i(l9(zero)) = 0 & ! [X1 : $int] : a(l8,X1) = a(l9(zero),X1) & test(l8) = test(l9(zero)) & ! [X2 : 'Nat()'] : (i(l9(X2)) = i(l11(X2)) & test(l9(X2)) = test(l11(X2)) & ! [X3 : $int] : a(l9(X2),X3) = a(l11(X2),X3)) & ! [X4 : 'Nat()'] : (v = a(l12(X4),i(l11(X4))) & ! [X5 : $int] : (i(l11(X4)) != X5 => a(l11(X4),X5) = a(l12(X4),X5)) & test(l11(X4)) = test(l12(X4)) & i(l11(X4)) = i(l12(X4)) & i(l9(s(X4))) = $sum(i(l12(X4)),1) & ! [X6 : $int] : a(l12(X4),X6) = a(l9(s(X4)),X6) & test(l12(X4)) = test(l9(s(X4)))) & ! [X7 : 'Nat()'] : ('Sub'(X7,nl9) => $less(i(l9(X7)),alength)) & ~$less(i(l9(nl9)),alength) & i(main_end) = i(l9(nl9)) & test(main_end) = test(l9(nl9)) & ! [X8 : $int] : a(main_end,X8) = a(l9(nl9),X8) [flattening 42]
unused predicate definition removal
[PP] pred marked as built-in: =
[PP] pred marked as built-in: $less
[PP] pred marked as built-in: $lesseq
=: +(42) -(12) 0(0)
'Sub': +(8) -(4) 0(0)
$less: +(11) -(12) 0(0)
=: +(42) -(12) 0(0)
'Sub': +(8) -(4) 0(0)
$less: +(11) -(12) 0(0)
preprocess 2 (ennf,flatten)
[PP] ennf in: 12. ! [X0 : 'Nat()'] : (('Sub'(s(X0),nl9) & test(l9(s(X0))) = $sum(test(l9(X0)),1)) => ! [X1 : 'Nat()',X2 : 'Nat()'] : (test(l9(X1)) = test(l9(X2)) => X1 = X2)) [input]
[PP] ennf out: 44. ! [X0 : 'Nat()'] : (! [X1 : 'Nat()',X2 : 'Nat()'] : (X1 = X2 | test(l9(X1)) != test(l9(X2))) | (~'Sub'(s(X0),nl9) | test(l9(s(X0))) != $sum(test(l9(X0)),1))) [ennf transformation 12]
[PP] flatten in: 44. ! [X0 : 'Nat()'] : (! [X1 : 'Nat()',X2 : 'Nat()'] : (X1 = X2 | test(l9(X1)) != test(l9(X2))) | (~'Sub'(s(X0),nl9) | test(l9(s(X0))) != $sum(test(l9(X0)),1))) [ennf transformation 12]
[PP] flatten out: 45. ! [X0 : 'Nat()'] : (! [X1 : 'Nat()',X2 : 'Nat()'] : (X1 = X2 | test(l9(X1)) != test(l9(X2))) | ~'Sub'(s(X0),nl9) | test(l9(s(X0))) != $sum(test(l9(X0)),1)) [flattening 44]
[PP] ennf in: 11. ! [X0 : 'Nat()'] : (('Sub'(s(X0),nl9) & i(l9(s(X0))) = $sum(i(l9(X0)),1)) => ! [X1 : 'Nat()',X2 : 'Nat()'] : (i(l9(X1)) = i(l9(X2)) => X1 = X2)) [input]
[PP] ennf out: 46. ! [X0 : 'Nat()'] : (! [X1 : 'Nat()',X2 : 'Nat()'] : (X1 = X2 | i(l9(X1)) != i(l9(X2))) | (~'Sub'(s(X0),nl9) | i(l9(s(X0))) != $sum(i(l9(X0)),1))) [ennf transformation 11]
[PP] flatten in: 46. ! [X0 : 'Nat()'] : (! [X1 : 'Nat()',X2 : 'Nat()'] : (X1 = X2 | i(l9(X1)) != i(l9(X2))) | (~'Sub'(s(X0),nl9) | i(l9(s(X0))) != $sum(i(l9(X0)),1))) [ennf transformation 11]
[PP] flatten out: 47. ! [X0 : 'Nat()'] : (! [X1 : 'Nat()',X2 : 'Nat()'] : (X1 = X2 | i(l9(X1)) != i(l9(X2))) | ~'Sub'(s(X0),nl9) | i(l9(s(X0))) != $sum(i(l9(X0)),1)) [flattening 46]
[PP] ennf in: 22. ! [X0 : $int,X1 : 'Nat()',X2 : $int] : ((~$less(X0,a(l9(zero),X2)) & $less(X0,a(l9(nl9),X2)) & a(l9(s(X1)),X2) = $sum(a(l9(X1),X2),1)) => ? [X3 : 'Nat()'] : (a(l9(X3),X2) = X0 & 'Sub'(X3,nl9))) [evaluation 10]
[PP] ennf out: 48. ! [X0 : $int,X1 : 'Nat()',X2 : $int] : (? [X3 : 'Nat()'] : (a(l9(X3),X2) = X0 & 'Sub'(X3,nl9)) | ($less(X0,a(l9(zero),X2)) | ~$less(X0,a(l9(nl9),X2)) | a(l9(s(X1)),X2) != $sum(a(l9(X1),X2),1))) [ennf transformation 22]
[PP] flatten in: 48. ! [X0 : $int,X1 : 'Nat()',X2 : $int] : (? [X3 : 'Nat()'] : (a(l9(X3),X2) = X0 & 'Sub'(X3,nl9)) | ($less(X0,a(l9(zero),X2)) | ~$less(X0,a(l9(nl9),X2)) | a(l9(s(X1)),X2) != $sum(a(l9(X1),X2),1))) [ennf transformation 22]
[PP] flatten out: 49. ! [X0 : $int,X1 : 'Nat()',X2 : $int] : (? [X3 : 'Nat()'] : (a(l9(X3),X2) = X0 & 'Sub'(X3,nl9)) | $less(X0,a(l9(zero),X2)) | ~$less(X0,a(l9(nl9),X2)) | a(l9(s(X1)),X2) != $sum(a(l9(X1),X2),1)) [flattening 48]
[PP] ennf in: 23. ! [X0 : $int,X1 : 'Nat()'] : ((~$less(X0,test(l9(zero))) & $less(X0,test(l9(nl9))) & test(l9(s(X1))) = $sum(test(l9(X1)),1)) => ? [X2 : 'Nat()'] : (test(l9(X2)) = X0 & 'Sub'(X2,nl9))) [evaluation 9]
[PP] ennf out: 50. ! [X0 : $int,X1 : 'Nat()'] : (? [X2 : 'Nat()'] : (test(l9(X2)) = X0 & 'Sub'(X2,nl9)) | ($less(X0,test(l9(zero))) | ~$less(X0,test(l9(nl9))) | test(l9(s(X1))) != $sum(test(l9(X1)),1))) [ennf transformation 23]
[PP] flatten in: 50. ! [X0 : $int,X1 : 'Nat()'] : (? [X2 : 'Nat()'] : (test(l9(X2)) = X0 & 'Sub'(X2,nl9)) | ($less(X0,test(l9(zero))) | ~$less(X0,test(l9(nl9))) | test(l9(s(X1))) != $sum(test(l9(X1)),1))) [ennf transformation 23]
[PP] flatten out: 51. ! [X0 : $int,X1 : 'Nat()'] : (? [X2 : 'Nat()'] : (test(l9(X2)) = X0 & 'Sub'(X2,nl9)) | $less(X0,test(l9(zero))) | ~$less(X0,test(l9(nl9))) | test(l9(s(X1))) != $sum(test(l9(X1)),1)) [flattening 50]
[PP] ennf in: 24. ! [X0 : $int,X1 : 'Nat()'] : ((~$less(X0,i(l9(zero))) & $less(X0,i(l9(nl9))) & i(l9(s(X1))) = $sum(i(l9(X1)),1)) => ? [X2 : 'Nat()'] : (i(l9(X2)) = X0 & 'Sub'(X2,nl9))) [evaluation 8]
[PP] ennf out: 52. ! [X0 : $int,X1 : 'Nat()'] : (? [X2 : 'Nat()'] : (i(l9(X2)) = X0 & 'Sub'(X2,nl9)) | ($less(X0,i(l9(zero))) | ~$less(X0,i(l9(nl9))) | i(l9(s(X1))) != $sum(i(l9(X1)),1))) [ennf transformation 24]
[PP] flatten in: 52. ! [X0 : $int,X1 : 'Nat()'] : (? [X2 : 'Nat()'] : (i(l9(X2)) = X0 & 'Sub'(X2,nl9)) | ($less(X0,i(l9(zero))) | ~$less(X0,i(l9(nl9))) | i(l9(s(X1))) != $sum(i(l9(X1)),1))) [ennf transformation 24]
[PP] flatten out: 53. ! [X0 : $int,X1 : 'Nat()'] : (? [X2 : 'Nat()'] : (i(l9(X2)) = X0 & 'Sub'(X2,nl9)) | $less(X0,i(l9(zero))) | ~$less(X0,i(l9(nl9))) | i(l9(s(X1))) != $sum(i(l9(X1)),1)) [flattening 52]
[PP] ennf in: 7. $less(i(l9(zero)),alength) => ? [X0 : 'Nat()'] : s(X0) = nl9 [input]
[PP] ennf out: 54. ? [X0 : 'Nat()'] : s(X0) = nl9 | ~$less(i(l9(zero)),alength) [ennf transformation 7]
[PP] ennf in: 6. ! [X0 : $int] : (! [X1 : 'Nat()'] : ('Sub'(X1,nl9) => a(l9(X1),X0) = a(l9(s(X1)),X0)) => a(l9(zero),X0) = a(l9(nl9),X0)) [input]
[PP] ennf out: 55. ! [X0 : $int] : (a(l9(zero),X0) = a(l9(nl9),X0) | ? [X1 : 'Nat()'] : (a(l9(X1),X0) != a(l9(s(X1)),X0) & 'Sub'(X1,nl9))) [ennf transformation 6]
[PP] ennf in: 5. ! [X0 : 'Nat()'] : ('Sub'(X0,nl9) => test(l9(X0)) = test(l9(s(X0)))) => test(l9(zero)) = test(l9(nl9)) [input]
[PP] ennf out: 56. test(l9(zero)) = test(l9(nl9)) | ? [X0 : 'Nat()'] : (test(l9(X0)) != test(l9(s(X0))) & 'Sub'(X0,nl9)) [ennf transformation 5]
[PP] ennf in: 4. ! [X0 : 'Nat()'] : ('Sub'(X0,nl9) => i(l9(X0)) = i(l9(s(X0)))) => i(l9(zero)) = i(l9(nl9)) [input]
[PP] ennf out: 57. i(l9(zero)) = i(l9(nl9)) | ? [X0 : 'Nat()'] : (i(l9(X0)) != i(l9(s(X0))) & 'Sub'(X0,nl9)) [ennf transformation 4]
[PP] ennf in: 43. test(l8) = 1 & ! [X0 : $int] : a(l7,X0) = a(l8,X0) & i(l9(zero)) = 0 & ! [X1 : $int] : a(l8,X1) = a(l9(zero),X1) & test(l8) = test(l9(zero)) & ! [X2 : 'Nat()'] : (i(l9(X2)) = i(l11(X2)) & test(l9(X2)) = test(l11(X2)) & ! [X3 : $int] : a(l9(X2),X3) = a(l11(X2),X3)) & ! [X4 : 'Nat()'] : (v = a(l12(X4),i(l11(X4))) & ! [X5 : $int] : (i(l11(X4)) != X5 => a(l11(X4),X5) = a(l12(X4),X5)) & test(l11(X4)) = test(l12(X4)) & i(l11(X4)) = i(l12(X4)) & i(l9(s(X4))) = $sum(i(l12(X4)),1) & ! [X6 : $int] : a(l12(X4),X6) = a(l9(s(X4)),X6) & test(l12(X4)) = test(l9(s(X4)))) & ! [X7 : 'Nat()'] : ('Sub'(X7,nl9) => $less(i(l9(X7)),alength)) & ~$less(i(l9(nl9)),alength) & i(main_end) = i(l9(nl9)) & test(main_end) = test(l9(nl9)) & ! [X8 : $int] : a(main_end,X8) = a(l9(nl9),X8) [flattening 42]
[PP] ennf out: 58. test(l8) = 1 & ! [X0 : $int] : a(l7,X0) = a(l8,X0) & i(l9(zero)) = 0 & ! [X1 : $int] : a(l8,X1) = a(l9(zero),X1) & test(l8) = test(l9(zero)) & ! [X2 : 'Nat()'] : (i(l9(X2)) = i(l11(X2)) & test(l9(X2)) = test(l11(X2)) & ! [X3 : $int] : a(l9(X2),X3) = a(l11(X2),X3)) & ! [X4 : 'Nat()'] : (v = a(l12(X4),i(l11(X4))) & ! [X5 : $int] : (a(l11(X4),X5) = a(l12(X4),X5) | i(l11(X4)) = X5) & test(l11(X4)) = test(l12(X4)) & i(l11(X4)) = i(l12(X4)) & i(l9(s(X4))) = $sum(i(l12(X4)),1) & ! [X6 : $int] : a(l12(X4),X6) = a(l9(s(X4)),X6) & test(l12(X4)) = test(l9(s(X4)))) & ! [X7 : 'Nat()'] : ($less(i(l9(X7)),alength) | ~'Sub'(X7,nl9)) & ~$less(i(l9(nl9)),alength) & i(main_end) = i(l9(nl9)) & test(main_end) = test(l9(nl9)) & ! [X8 : $int] : a(main_end,X8) = a(l9(nl9),X8) [ennf transformation 43]
[PP] ennf in: 2. ! [X0 : 'Nat()',X1 : 'Nat()'] : ('Sub'(X0,X1) => 'Sub'(X0,s(X1))) [input]
[PP] ennf out: 59. ! [X0 : 'Nat()',X1 : 'Nat()'] : ('Sub'(X0,s(X1)) | ~'Sub'(X0,X1)) [ennf transformation 2]
naming
[PP] naming args: 37. ! [X0 : 'Nat()'] : (zero = X0 | s(p(X0)) = X0) [term algebras exhaustiveness]
[PP] naming args: 40. $less(2,1) [flattening 20]
[PP] naming args: 17. ! [X0 : $int] : X0 = X0 [input]
[PP] naming args: 21. ~$less(1,0) [evaluation 16]
[PP] naming args: 15. ! [X0 : $int] : X0 = X0 [input]
[PP] naming args: 41. ! [X0 : $int,X1 : $int,X2 : $int] : $sum($sum(X0,X1),X2) = $sum($sum(X0,X2),X1) [rectify 14]
[PP] naming args: 13. ! [X0 : 'Nat()'] : test(l9(zero)) = test(l9(X0)) [input]
[PP] naming args: 45. ! [X0 : 'Nat()'] : (! [X1 : 'Nat()',X2 : 'Nat()'] : (X1 = X2 | test(l9(X1)) != test(l9(X2))) | ~'Sub'(s(X0),nl9) | test(l9(s(X0))) != $sum(test(l9(X0)),1)) [flattening 44]
[PP] naming args: 47. ! [X0 : 'Nat()'] : (! [X1 : 'Nat()',X2 : 'Nat()'] : (X1 = X2 | i(l9(X1)) != i(l9(X2))) | ~'Sub'(s(X0),nl9) | i(l9(s(X0))) != $sum(i(l9(X0)),1)) [flattening 46]
[PP] naming args: 49. ! [X0 : $int,X1 : 'Nat()',X2 : $int] : (? [X3 : 'Nat()'] : (a(l9(X3),X2) = X0 & 'Sub'(X3,nl9)) | $less(X0,a(l9(zero),X2)) | ~$less(X0,a(l9(nl9),X2)) | a(l9(s(X1)),X2) != $sum(a(l9(X1),X2),1)) [flattening 48]
[PP] naming args: 51. ! [X0 : $int,X1 : 'Nat()'] : (? [X2 : 'Nat()'] : (test(l9(X2)) = X0 & 'Sub'(X2,nl9)) | $less(X0,test(l9(zero))) | ~$less(X0,test(l9(nl9))) | test(l9(s(X1))) != $sum(test(l9(X1)),1)) [flattening 50]
[PP] naming args: 53. ! [X0 : $int,X1 : 'Nat()'] : (? [X2 : 'Nat()'] : (i(l9(X2)) = X0 & 'Sub'(X2,nl9)) | $less(X0,i(l9(zero))) | ~$less(X0,i(l9(nl9))) | i(l9(s(X1))) != $sum(i(l9(X1)),1)) [flattening 52]
[PP] naming args: 54. ? [X0 : 'Nat()'] : s(X0) = nl9 | ~$less(i(l9(zero)),alength) [ennf transformation 7]
[PP] naming args: 55. ! [X0 : $int] : (a(l9(zero),X0) = a(l9(nl9),X0) | ? [X1 : 'Nat()'] : (a(l9(X1),X0) != a(l9(s(X1)),X0) & 'Sub'(X1,nl9))) [ennf transformation 6]
[PP] naming args: 56. test(l9(zero)) = test(l9(nl9)) | ? [X0 : 'Nat()'] : (test(l9(X0)) != test(l9(s(X0))) & 'Sub'(X0,nl9)) [ennf transformation 5]
[PP] naming args: 57. i(l9(zero)) = i(l9(nl9)) | ? [X0 : 'Nat()'] : (i(l9(X0)) != i(l9(s(X0))) & 'Sub'(X0,nl9)) [ennf transformation 4]
[PP] naming args: 58. test(l8) = 1 & ! [X0 : $int] : a(l7,X0) = a(l8,X0) & i(l9(zero)) = 0 & ! [X1 : $int] : a(l8,X1) = a(l9(zero),X1) & test(l8) = test(l9(zero)) & ! [X2 : 'Nat()'] : (i(l9(X2)) = i(l11(X2)) & test(l9(X2)) = test(l11(X2)) & ! [X3 : $int] : a(l9(X2),X3) = a(l11(X2),X3)) & ! [X4 : 'Nat()'] : (v = a(l12(X4),i(l11(X4))) & ! [X5 : $int] : (a(l11(X4),X5) = a(l12(X4),X5) | i(l11(X4)) = X5) & test(l11(X4)) = test(l12(X4)) & i(l11(X4)) = i(l12(X4)) & i(l9(s(X4))) = $sum(i(l12(X4)),1) & ! [X6 : $int] : a(l12(X4),X6) = a(l9(s(X4)),X6) & test(l12(X4)) = test(l9(s(X4)))) & ! [X7 : 'Nat()'] : ($less(i(l9(X7)),alength) | ~'Sub'(X7,nl9)) & ~$less(i(l9(nl9)),alength) & i(main_end) = i(l9(nl9)) & test(main_end) = test(l9(nl9)) & ! [X8 : $int] : a(main_end,X8) = a(l9(nl9),X8) [ennf transformation 43]
[PP] naming args: 59. ! [X0 : 'Nat()',X1 : 'Nat()'] : ('Sub'(X0,s(X1)) | ~'Sub'(X0,X1)) [ennf transformation 2]
[PP] naming args: 1. ! [X0 : 'Nat()'] : 'Sub'(X0,s(X0)) [input]
preprocess3 (nnf, flatten, skolemize)
clausify
[PP] clausify: 39. s(X0) != s(X1) | X0 = X1 [term algebras injectivity] {a:0,w:8,tD:1,inD:0}
[PP] clausify: 38. zero != s(X0) [term algebras distinctness] {a:0,w:4,tD:1,inD:0}
[PP] clausify: 37. ! [X0 : 'Nat()'] : (zero = X0 | s(p(X0)) = X0) [term algebras exhaustiveness]
[PP] clausify: 36. ~$less(X0,X1) | ~$less(X1,$sum(X0,1)) [theory axiom] {a:0,w:8,tD:1,inD:0}
[PP] clausify: 35. $uminus($uminus(X0)) = X0 [theory axiom] {a:0,w:5,tD:1,inD:0}
[PP] clausify: 34. $less(X0,X1) | $less(X1,$sum(X0,1)) [theory axiom] {a:0,w:8,tD:1,inD:0}
[PP] clausify: 33. ~$less(X0,X1) | $less($sum(X0,X2),$sum(X1,X2)) [theory axiom] {a:0,w:10,tD:1,inD:0}
[PP] clausify: 32. $less(X0,X1) | $less(X1,X0) | X0 = X1 [theory axiom] {a:0,w:9,tD:1,inD:0}
[PP] clausify: 31. ~$less(X0,X1) | ~$less(X1,X2) | $less(X0,X2) [theory axiom] {a:0,w:9,tD:1,inD:0}
[PP] clausify: 30. ~$less(X0,X0) [theory axiom] {a:0,w:3,tD:1,inD:0}
[PP] clausify: 29. $sum(X0,$uminus(X0)) = 0 [theory axiom] {a:0,w:6,tD:1,inD:0}
[PP] clausify: 28. $sum($uminus(X1),$uminus(X0)) = $uminus($sum(X0,X1)) [theory axiom] {a:0,w:10,tD:1,inD:0}
[PP] clausify: 27. $sum(X0,0) = X0 [theory axiom] {a:0,w:5,tD:1,inD:0}
[PP] clausify: 26. $sum(X0,$sum(X1,X2)) = $sum($sum(X0,X1),X2) [theory axiom] {a:0,w:11,tD:1,inD:0}
[PP] clausify: 25. $sum(X0,X1) = $sum(X1,X0) [theory axiom] {a:0,w:7,tD:1,inD:0}
[PP] clausify: 40. $less(2,1) [flattening 20]
[PP] clausify: 17. ! [X0 : $int] : X0 = X0 [input]
[PP] clausify: 21. ~$less(1,0) [evaluation 16]
[PP] clausify: 15. ! [X0 : $int] : X0 = X0 [input]
[PP] clausify: 41. ! [X0 : $int,X1 : $int,X2 : $int] : $sum($sum(X0,X1),X2) = $sum($sum(X0,X2),X1) [rectify 14]
[PP] clausify: 13. ! [X0 : 'Nat()'] : test(l9(zero)) = test(l9(X0)) [input]
[PP] clausify: 45. ! [X0 : 'Nat()'] : (! [X1 : 'Nat()',X2 : 'Nat()'] : (X1 = X2 | test(l9(X1)) != test(l9(X2))) | ~'Sub'(s(X0),nl9) | test(l9(s(X0))) != $sum(test(l9(X0)),1)) [flattening 44]
[PP] clausify: 47. ! [X0 : 'Nat()'] : (! [X1 : 'Nat()',X2 : 'Nat()'] : (X1 = X2 | i(l9(X1)) != i(l9(X2))) | ~'Sub'(s(X0),nl9) | i(l9(s(X0))) != $sum(i(l9(X0)),1)) [flattening 46]
[PP] clausify: 61. ! [X0 : $int,X1 : 'Nat()',X2 : $int] : ((a(l9(sK0(X0,X2)),X2) = X0 & 'Sub'(sK0(X0,X2),nl9)) | $less(X0,a(l9(zero),X2)) | ~$less(X0,a(l9(nl9),X2)) | a(l9(s(X1)),X2) != $sum(a(l9(X1),X2),1)) [skolemisation 49,60]
[PP] clausify: 63. ! [X0 : $int,X1 : 'Nat()'] : ((test(l9(sK1(X0))) = X0 & 'Sub'(sK1(X0),nl9)) | $less(X0,test(l9(zero))) | ~$less(X0,test(l9(nl9))) | test(l9(s(X1))) != $sum(test(l9(X1)),1)) [skolemisation 51,62]
[PP] clausify: 65. ! [X0 : $int,X1 : 'Nat()'] : ((i(l9(sK2(X0))) = X0 & 'Sub'(sK2(X0),nl9)) | $less(X0,i(l9(zero))) | ~$less(X0,i(l9(nl9))) | i(l9(s(X1))) != $sum(i(l9(X1)),1)) [skolemisation 53,64]
[PP] clausify: 67. s(sK3) = nl9 | ~$less(i(l9(zero)),alength) [skolemisation 54,66]
[PP] clausify: 69. ! [X0 : $int] : (a(l9(zero),X0) = a(l9(nl9),X0) | (a(l9(sK4(X0)),X0) != a(l9(s(sK4(X0))),X0) & 'Sub'(sK4(X0),nl9))) [skolemisation 55,68]
[PP] clausify: 71. test(l9(zero)) = test(l9(nl9)) | (test(l9(sK5)) != test(l9(s(sK5))) & 'Sub'(sK5,nl9)) [skolemisation 56,70]
[PP] clausify: 73. i(l9(zero)) = i(l9(nl9)) | (i(l9(sK6)) != i(l9(s(sK6))) & 'Sub'(sK6,nl9)) [skolemisation 57,72]
[PP] clausify: 58. test(l8) = 1 & ! [X0 : $int] : a(l7,X0) = a(l8,X0) & i(l9(zero)) = 0 & ! [X1 : $int] : a(l8,X1) = a(l9(zero),X1) & test(l8) = test(l9(zero)) & ! [X2 : 'Nat()'] : (i(l9(X2)) = i(l11(X2)) & test(l9(X2)) = test(l11(X2)) & ! [X3 : $int] : a(l9(X2),X3) = a(l11(X2),X3)) & ! [X4 : 'Nat()'] : (v = a(l12(X4),i(l11(X4))) & ! [X5 : $int] : (a(l11(X4),X5) = a(l12(X4),X5) | i(l11(X4)) = X5) & test(l11(X4)) = test(l12(X4)) & i(l11(X4)) = i(l12(X4)) & i(l9(s(X4))) = $sum(i(l12(X4)),1) & ! [X6 : $int] : a(l12(X4),X6) = a(l9(s(X4)),X6) & test(l12(X4)) = test(l9(s(X4)))) & ! [X7 : 'Nat()'] : ($less(i(l9(X7)),alength) | ~'Sub'(X7,nl9)) & ~$less(i(l9(nl9)),alength) & i(main_end) = i(l9(nl9)) & test(main_end) = test(l9(nl9)) & ! [X8 : $int] : a(main_end,X8) = a(l9(nl9),X8) [ennf transformation 43]
[PP] clausify: 59. ! [X0 : 'Nat()',X1 : 'Nat()'] : ('Sub'(X0,s(X1)) | ~'Sub'(X0,X1)) [ennf transformation 2]
[PP] clausify: 1. ! [X0 : 'Nat()'] : 'Sub'(X0,s(X0)) [input]
[PP] final: 39. s(X0) != s(X1) | X0 = X1 [term algebras injectivity] {a:0,w:8,tD:1,inD:0}
[PP] final: 38. zero != s(X0) [term algebras distinctness] {a:0,w:4,tD:1,inD:0}
[PP] final: 74. zero = X0 | s(p(X0)) = X0 [cnf transformation 37] {a:0,w:8,inD:0}
[PP] final: 36. ~$less(X0,X1) | ~$less(X1,$sum(X0,1)) [theory axiom] {a:0,w:8,tD:1,inD:0}
[PP] final: 35. $uminus($uminus(X0)) = X0 [theory axiom] {a:0,w:5,tD:1,inD:0}
[PP] final: 34. $less(X0,X1) | $less(X1,$sum(X0,1)) [theory axiom] {a:0,w:8,tD:1,inD:0}
[PP] final: 33. ~$less(X0,X1) | $less($sum(X0,X2),$sum(X1,X2)) [theory axiom] {a:0,w:10,tD:1,inD:0}
[PP] final: 32. $less(X0,X1) | $less(X1,X0) | X0 = X1 [theory axiom] {a:0,w:9,tD:1,inD:0}
[PP] final: 31. ~$less(X0,X1) | ~$less(X1,X2) | $less(X0,X2) [theory axiom] {a:0,w:9,tD:1,inD:0}
[PP] final: 30. ~$less(X0,X0) [theory axiom] {a:0,w:3,tD:1,inD:0}
[PP] final: 29. $sum(X0,$uminus(X0)) = 0 [theory axiom] {a:0,w:6,tD:1,inD:0}
[PP] final: 28. $sum($uminus(X1),$uminus(X0)) = $uminus($sum(X0,X1)) [theory axiom] {a:0,w:10,tD:1,inD:0}
[PP] final: 27. $sum(X0,0) = X0 [theory axiom] {a:0,w:5,tD:1,inD:0}
[PP] final: 26. $sum(X0,$sum(X1,X2)) = $sum($sum(X0,X1),X2) [theory axiom] {a:0,w:11,tD:1,inD:0}
[PP] final: 25. $sum(X0,X1) = $sum(X1,X0) [theory axiom] {a:0,w:7,tD:1,inD:0}
[PP] final: 75. $less(2,1) [cnf transformation 40] {a:0,w:3,goal:1,inD:0}
[PP] final: 76. X0 = X0 [cnf transformation 17] {a:0,w:3,inD:0}
[PP] final: 77. ~$less(1,0) [cnf transformation 21] {a:0,w:3,inD:0}
[PP] final: 78. X0 = X0 [cnf transformation 15] {a:0,w:3,inD:0}
[PP] final: 79. $sum($sum(X0,X1),X2) = $sum($sum(X0,X2),X1) [cnf transformation 41] {a:0,w:11,inD:0}
[PP] final: 80. test(l9(zero)) = test(l9(X0)) [cnf transformation 13] {a:0,w:7,inD:0}
[PP] final: 81. X1 = X2 | test(l9(X1)) != test(l9(X2)) | ~'Sub'(s(X0),nl9) | test(l9(s(X0))) != $sum(test(l9(X0)),1) [cnf transformation 45] {a:0,w:24,inD:0}
[PP] final: 82. X1 = X2 | i(l9(X1)) != i(l9(X2)) | ~'Sub'(s(X0),nl9) | i(l9(s(X0))) != $sum(i(l9(X0)),1) [cnf transformation 47] {a:0,w:24,inD:0}
[PP] final: 84. a(l9(sK0(X0,X2)),X2) = X0 | $less(X0,a(l9(zero),X2)) | ~$less(X0,a(l9(nl9),X2)) | a(l9(s(X1)),X2) != $sum(a(l9(X1),X2),1) [cnf transformation 61] {a:0,w:32,inD:0}
[PP] final: 83. 'Sub'(sK0(X0,X2),nl9) | $less(X0,a(l9(zero),X2)) | ~$less(X0,a(l9(nl9),X2)) | a(l9(s(X1)),X2) != $sum(a(l9(X1),X2),1) [cnf transformation 61] {a:0,w:29,inD:0}
[PP] final: 86. test(l9(sK1(X0))) = X0 | $less(X0,test(l9(zero))) | ~$less(X0,test(l9(nl9))) | test(l9(s(X1))) != $sum(test(l9(X1)),1) [cnf transformation 63] {a:0,w:26,inD:0}
[PP] final: 85. 'Sub'(sK1(X0),nl9) | $less(X0,test(l9(zero))) | ~$less(X0,test(l9(nl9))) | test(l9(s(X1))) != $sum(test(l9(X1)),1) [cnf transformation 63] {a:0,w:24,inD:0}
[PP] final: 88. i(l9(sK2(X0))) = X0 | $less(X0,i(l9(zero))) | ~$less(X0,i(l9(nl9))) | i(l9(s(X1))) != $sum(i(l9(X1)),1) [cnf transformation 65] {a:0,w:26,inD:0}
[PP] final: 87. 'Sub'(sK2(X0),nl9) | $less(X0,i(l9(zero))) | ~$less(X0,i(l9(nl9))) | i(l9(s(X1))) != $sum(i(l9(X1)),1) [cnf transformation 65] {a:0,w:24,inD:0}
[PP] final: 89. s(sK3) = nl9 | ~$less(i(l9(zero)),alength) [cnf transformation 67] {a:0,w:9,inD:0}
[PP] final: 91. a(l9(zero),X0) = a(l9(nl9),X0) | a(l9(sK4(X0)),X0) != a(l9(s(sK4(X0))),X0) [cnf transformation 69] {a:0,w:21,inD:0}
[PP] final: 90. a(l9(zero),X0) = a(l9(nl9),X0) | 'Sub'(sK4(X0),nl9) [cnf transformation 69] {a:0,w:13,inD:0}
[PP] final: 93. test(l9(zero)) = test(l9(nl9)) | test(l9(sK5)) != test(l9(s(sK5))) [cnf transformation 71] {a:0,w:15,inD:0}
[PP] final: 92. test(l9(zero)) = test(l9(nl9)) | 'Sub'(sK5,nl9) [cnf transformation 71] {a:0,w:10,inD:0}
[PP] final: 95. i(l9(zero)) = i(l9(nl9)) | i(l9(sK6)) != i(l9(s(sK6))) [cnf transformation 73] {a:0,w:15,inD:0}
[PP] final: 94. i(l9(zero)) = i(l9(nl9)) | 'Sub'(sK6,nl9) [cnf transformation 73] {a:0,w:10,inD:0}
[PP] final: 115. test(l8) = 1 [cnf transformation 58] {a:0,w:4,inD:0}
[PP] final: 114. a(l7,X0) = a(l8,X0) [cnf transformation 58] {a:0,w:7,inD:0}
[PP] final: 113. i(l9(zero)) = 0 [cnf transformation 58] {a:0,w:5,inD:0}
[PP] final: 112. a(l8,X1) = a(l9(zero),X1) [cnf transformation 58] {a:0,w:8,inD:0}
[PP] final: 111. test(l8) = test(l9(zero)) [cnf transformation 58] {a:0,w:6,inD:0}
[PP] final: 110. i(l9(X2)) = i(l11(X2)) [cnf transformation 58] {a:0,w:7,inD:0}
[PP] final: 109. test(l9(X2)) = test(l11(X2)) [cnf transformation 58] {a:0,w:7,inD:0}
[PP] final: 108. a(l9(X2),X3) = a(l11(X2),X3) [cnf transformation 58] {a:0,w:9,inD:0}
[PP] final: 107. v = a(l12(X4),i(l11(X4))) [cnf transformation 58] {a:0,w:8,inD:0}
[PP] final: 106. a(l11(X4),X5) = a(l12(X4),X5) | i(l11(X4)) = X5 [cnf transformation 58] {a:0,w:14,inD:0}
[PP] final: 105. test(l11(X4)) = test(l12(X4)) [cnf transformation 58] {a:0,w:7,inD:0}
[PP] final: 104. i(l11(X4)) = i(l12(X4)) [cnf transformation 58] {a:0,w:7,inD:0}
[PP] final: 103. i(l9(s(X4))) = $sum(i(l12(X4)),1) [cnf transformation 58] {a:0,w:10,inD:0}
[PP] final: 102. a(l12(X4),X6) = a(l9(s(X4)),X6) [cnf transformation 58] {a:0,w:10,inD:0}
[PP] final: 101. test(l12(X4)) = test(l9(s(X4))) [cnf transformation 58] {a:0,w:8,inD:0}
[PP] final: 100. $less(i(l9(X7)),alength) | ~'Sub'(X7,nl9) [cnf transformation 58] {a:0,w:8,inD:0}
[PP] final: 99. ~$less(i(l9(nl9)),alength) [cnf transformation 58] {a:0,w:5,inD:0}
[PP] final: 98. i(main_end) = i(l9(nl9)) [cnf transformation 58] {a:0,w:6,inD:0}
[PP] final: 97. test(main_end) = test(l9(nl9)) [cnf transformation 58] {a:0,w:6,inD:0}
[PP] final: 96. a(main_end,X8) = a(l9(nl9),X8) [cnf transformation 58] {a:0,w:8,inD:0}
[PP] final: 116. 'Sub'(X0,s(X1)) | ~'Sub'(X0,X1) [cnf transformation 59] {a:0,w:7,inD:0}
[PP] final: 117. 'Sub'(X0,s(X0)) [cnf transformation 1] {a:0,w:4,inD:0}
preprocessing finished
[PP] onPreprocessingEnd(), Proving Helper
1. ! [X0 : 'Nat()'] : 'Sub'(X0,s(X0)) [input]
2. ! [X0 : 'Nat()',X1 : 'Nat()'] : ('Sub'(X0,X1) => 'Sub'(X0,s(X1))) [input]
3. (test(l8) = 1 & ! [X8 : $int] : a(l7,X8) = a(l8,X8)) & (i(l9(zero)) = 0 & ! [X7 : $int] : a(l8,X7) = a(l9(zero),X7) & test(l8) = test(l9(zero))) & (! [X5 : 'Nat()'] : (i(l9(X5)) = i(l11(X5)) & test(l9(X5)) = test(l11(X5)) & ! [X6 : $int] : a(l9(X5),X6) = a(l11(X5),X6)) & ! [X2 : 'Nat()'] : ((v = a(l12(X2),i(l11(X2))) & ! [X4 : $int] : (~i(l11(X2)) = X4 => a(l11(X2),X4) = a(l12(X2),X4)) & test(l11(X2)) = test(l12(X2)) & i(l11(X2)) = i(l12(X2))) & (i(l9(s(X2))) = $sum(i(l12(X2)),1) & ! [X3 : $int] : a(l12(X2),X3) = a(l9(s(X2)),X3) & test(l12(X2)) = test(l9(s(X2))))) & ! [X1 : 'Nat()'] : ('Sub'(X1,nl9) => $less(i(l9(X1)),alength)) & ~$less(i(l9(nl9)),alength) & (i(main_end) = i(l9(nl9)) & test(main_end) = test(l9(nl9)) & ! [X0 : $int] : a(main_end,X0) = a(l9(nl9),X0))) [input]
4. ! [X0 : 'Nat()'] : ('Sub'(X0,nl9) => i(l9(X0)) = i(l9(s(X0)))) => i(l9(zero)) = i(l9(nl9)) [input]
5. ! [X0 : 'Nat()'] : ('Sub'(X0,nl9) => test(l9(X0)) = test(l9(s(X0)))) => test(l9(zero)) = test(l9(nl9)) [input]
6. ! [X0 : $int] : (! [X1 : 'Nat()'] : ('Sub'(X1,nl9) => a(l9(X1),X0) = a(l9(s(X1)),X0)) => a(l9(zero),X0) = a(l9(nl9),X0)) [input]
7. $less(i(l9(zero)),alength) => ? [X0 : 'Nat()'] : s(X0) = nl9 [input]
8. ! [X0 : $int,X1 : 'Nat()'] : (($lesseq(i(l9(zero)),X0) & $less(X0,i(l9(nl9))) & i(l9(s(X1))) = $sum(i(l9(X1)),1)) => ? [X2 : 'Nat()'] : (i(l9(X2)) = X0 & 'Sub'(X2,nl9))) [input]
9. ! [X0 : $int,X1 : 'Nat()'] : (($lesseq(test(l9(zero)),X0) & $less(X0,test(l9(nl9))) & test(l9(s(X1))) = $sum(test(l9(X1)),1)) => ? [X2 : 'Nat()'] : (test(l9(X2)) = X0 & 'Sub'(X2,nl9))) [input]
10. ! [X0 : $int,X1 : 'Nat()',X2 : $int] : (($lesseq(a(l9(zero),X2),X0) & $less(X0,a(l9(nl9),X2)) & a(l9(s(X1)),X2) = $sum(a(l9(X1),X2),1)) => ? [X3 : 'Nat()'] : (a(l9(X3),X2) = X0 & 'Sub'(X3,nl9))) [input]
11. ! [X0 : 'Nat()'] : (('Sub'(s(X0),nl9) & i(l9(s(X0))) = $sum(i(l9(X0)),1)) => ! [X1 : 'Nat()',X2 : 'Nat()'] : (i(l9(X1)) = i(l9(X2)) => X1 = X2)) [input]
12. ! [X0 : 'Nat()'] : (('Sub'(s(X0),nl9) & test(l9(s(X0))) = $sum(test(l9(X0)),1)) => ! [X1 : 'Nat()',X2 : 'Nat()'] : (test(l9(X1)) = test(l9(X2)) => X1 = X2)) [input]
13. ! [X0 : 'Nat()'] : test(l9(zero)) = test(l9(X0)) [input]
14. ! [X0 : $int,X2 : $int,X1 : $int] : $sum($sum(X0,X1),X2) = $sum($sum(X0,X2),X1) [input]
15. ! [X0 : $int] : X0 = X0 [input]
16. $lesseq(0,1) [input]
17. ! [X0 : $int] : X0 = X0 [input]
18. $lesseq(1,2) [input]
19. ~$lesseq(1,2) [negated conjecture 18]
20. ~~$less(2,1) [evaluation 19]
21. ~$less(1,0) [evaluation 16]
22. ! [X0 : $int,X1 : 'Nat()',X2 : $int] : ((~$less(X0,a(l9(zero),X2)) & $less(X0,a(l9(nl9),X2)) & a(l9(s(X1)),X2) = $sum(a(l9(X1),X2),1)) => ? [X3 : 'Nat()'] : (a(l9(X3),X2) = X0 & 'Sub'(X3,nl9))) [evaluation 10]
23. ! [X0 : $int,X1 : 'Nat()'] : ((~$less(X0,test(l9(zero))) & $less(X0,test(l9(nl9))) & test(l9(s(X1))) = $sum(test(l9(X1)),1)) => ? [X2 : 'Nat()'] : (test(l9(X2)) = X0 & 'Sub'(X2,nl9))) [evaluation 9]
24. ! [X0 : $int,X1 : 'Nat()'] : ((~$less(X0,i(l9(zero))) & $less(X0,i(l9(nl9))) & i(l9(s(X1))) = $sum(i(l9(X1)),1)) => ? [X2 : 'Nat()'] : (i(l9(X2)) = X0 & 'Sub'(X2,nl9))) [evaluation 8]
25. $sum(X0,X1) = $sum(X1,X0) [theory axiom] {a:0,w:7,tD:1,inD:0}
26. $sum(X0,$sum(X1,X2)) = $sum($sum(X0,X1),X2) [theory axiom] {a:0,w:11,tD:1,inD:0}
27. $sum(X0,0) = X0 [theory axiom] {a:0,w:5,tD:1,inD:0}
28. $sum($uminus(X1),$uminus(X0)) = $uminus($sum(X0,X1)) [theory axiom] {a:0,w:10,tD:1,inD:0}
29. $sum(X0,$uminus(X0)) = 0 [theory axiom] {a:0,w:6,tD:1,inD:0}
30. ~$less(X0,X0) [theory axiom] {a:0,w:3,tD:1,inD:0}
31. ~$less(X0,X1) | ~$less(X1,X2) | $less(X0,X2) [theory axiom] {a:0,w:9,tD:1,inD:0}
32. $less(X0,X1) | $less(X1,X0) | X0 = X1 [theory axiom] {a:0,w:9,tD:1,inD:0}
33. ~$less(X0,X1) | $less($sum(X0,X2),$sum(X1,X2)) [theory axiom] {a:0,w:10,tD:1,inD:0}
34. $less(X0,X1) | $less(X1,$sum(X0,1)) [theory axiom] {a:0,w:8,tD:1,inD:0}
35. $uminus($uminus(X0)) = X0 [theory axiom] {a:0,w:5,tD:1,inD:0}
36. ~$less(X0,X1) | ~$less(X1,$sum(X0,1)) [theory axiom] {a:0,w:8,tD:1,inD:0}
37. ! [X0 : 'Nat()'] : (zero = X0 | s(p(X0)) = X0) [term algebras exhaustiveness]
38. zero != s(X0) [term algebras distinctness] {a:0,w:4,tD:1,inD:0}
39. s(X0) != s(X1) | X0 = X1 [term algebras injectivity] {a:0,w:8,tD:1,inD:0}
40. $less(2,1) [flattening 20]
41. ! [X0 : $int,X1 : $int,X2 : $int] : $sum($sum(X0,X1),X2) = $sum($sum(X0,X2),X1) [rectify 14]
42. (test(l8) = 1 & ! [X0 : $int] : a(l7,X0) = a(l8,X0)) & (i(l9(zero)) = 0 & ! [X1 : $int] : a(l8,X1) = a(l9(zero),X1) & test(l8) = test(l9(zero))) & (! [X2 : 'Nat()'] : (i(l9(X2)) = i(l11(X2)) & test(l9(X2)) = test(l11(X2)) & ! [X3 : $int] : a(l9(X2),X3) = a(l11(X2),X3)) & ! [X4 : 'Nat()'] : ((v = a(l12(X4),i(l11(X4))) & ! [X5 : $int] : (~i(l11(X4)) = X5 => a(l11(X4),X5) = a(l12(X4),X5)) & test(l11(X4)) = test(l12(X4)) & i(l11(X4)) = i(l12(X4))) & (i(l9(s(X4))) = $sum(i(l12(X4)),1) & ! [X6 : $int] : a(l12(X4),X6) = a(l9(s(X4)),X6) & test(l12(X4)) = test(l9(s(X4))))) & ! [X7 : 'Nat()'] : ('Sub'(X7,nl9) => $less(i(l9(X7)),alength)) & ~$less(i(l9(nl9)),alength) & (i(main_end) = i(l9(nl9)) & test(main_end) = test(l9(nl9)) & ! [X8 : $int] : a(main_end,X8) = a(l9(nl9),X8))) [rectify 3]
43. test(l8) = 1 & ! [X0 : $int] : a(l7,X0) = a(l8,X0) & i(l9(zero)) = 0 & ! [X1 : $int] : a(l8,X1) = a(l9(zero),X1) & test(l8) = test(l9(zero)) & ! [X2 : 'Nat()'] : (i(l9(X2)) = i(l11(X2)) & test(l9(X2)) = test(l11(X2)) & ! [X3 : $int] : a(l9(X2),X3) = a(l11(X2),X3)) & ! [X4 : 'Nat()'] : (v = a(l12(X4),i(l11(X4))) & ! [X5 : $int] : (i(l11(X4)) != X5 => a(l11(X4),X5) = a(l12(X4),X5)) & test(l11(X4)) = test(l12(X4)) & i(l11(X4)) = i(l12(X4)) & i(l9(s(X4))) = $sum(i(l12(X4)),1) & ! [X6 : $int] : a(l12(X4),X6) = a(l9(s(X4)),X6) & test(l12(X4)) = test(l9(s(X4)))) & ! [X7 : 'Nat()'] : ('Sub'(X7,nl9) => $less(i(l9(X7)),alength)) & ~$less(i(l9(nl9)),alength) & i(main_end) = i(l9(nl9)) & test(main_end) = test(l9(nl9)) & ! [X8 : $int] : a(main_end,X8) = a(l9(nl9),X8) [flattening 42]
44. ! [X0 : 'Nat()'] : (! [X1 : 'Nat()',X2 : 'Nat()'] : (X1 = X2 | test(l9(X1)) != test(l9(X2))) | (~'Sub'(s(X0),nl9) | test(l9(s(X0))) != $sum(test(l9(X0)),1))) [ennf transformation 12]
45. ! [X0 : 'Nat()'] : (! [X1 : 'Nat()',X2 : 'Nat()'] : (X1 = X2 | test(l9(X1)) != test(l9(X2))) | ~'Sub'(s(X0),nl9) | test(l9(s(X0))) != $sum(test(l9(X0)),1)) [flattening 44]
46. ! [X0 : 'Nat()'] : (! [X1 : 'Nat()',X2 : 'Nat()'] : (X1 = X2 | i(l9(X1)) != i(l9(X2))) | (~'Sub'(s(X0),nl9) | i(l9(s(X0))) != $sum(i(l9(X0)),1))) [ennf transformation 11]
47. ! [X0 : 'Nat()'] : (! [X1 : 'Nat()',X2 : 'Nat()'] : (X1 = X2 | i(l9(X1)) != i(l9(X2))) | ~'Sub'(s(X0),nl9) | i(l9(s(X0))) != $sum(i(l9(X0)),1)) [flattening 46]
48. ! [X0 : $int,X1 : 'Nat()',X2 : $int] : (? [X3 : 'Nat()'] : (a(l9(X3),X2) = X0 & 'Sub'(X3,nl9)) | ($less(X0,a(l9(zero),X2)) | ~$less(X0,a(l9(nl9),X2)) | a(l9(s(X1)),X2) != $sum(a(l9(X1),X2),1))) [ennf transformation 22]
49. ! [X0 : $int,X1 : 'Nat()',X2 : $int] : (? [X3 : 'Nat()'] : (a(l9(X3),X2) = X0 & 'Sub'(X3,nl9)) | $less(X0,a(l9(zero),X2)) | ~$less(X0,a(l9(nl9),X2)) | a(l9(s(X1)),X2) != $sum(a(l9(X1),X2),1)) [flattening 48]
50. ! [X0 : $int,X1 : 'Nat()'] : (? [X2 : 'Nat()'] : (test(l9(X2)) = X0 & 'Sub'(X2,nl9)) | ($less(X0,test(l9(zero))) | ~$less(X0,test(l9(nl9))) | test(l9(s(X1))) != $sum(test(l9(X1)),1))) [ennf transformation 23]
51. ! [X0 : $int,X1 : 'Nat()'] : (? [X2 : 'Nat()'] : (test(l9(X2)) = X0 & 'Sub'(X2,nl9)) | $less(X0,test(l9(zero))) | ~$less(X0,test(l9(nl9))) | test(l9(s(X1))) != $sum(test(l9(X1)),1)) [flattening 50]
52. ! [X0 : $int,X1 : 'Nat()'] : (? [X2 : 'Nat()'] : (i(l9(X2)) = X0 & 'Sub'(X2,nl9)) | ($less(X0,i(l9(zero))) | ~$less(X0,i(l9(nl9))) | i(l9(s(X1))) != $sum(i(l9(X1)),1))) [ennf transformation 24]
53. ! [X0 : $int,X1 : 'Nat()'] : (? [X2 : 'Nat()'] : (i(l9(X2)) = X0 & 'Sub'(X2,nl9)) | $less(X0,i(l9(zero))) | ~$less(X0,i(l9(nl9))) | i(l9(s(X1))) != $sum(i(l9(X1)),1)) [flattening 52]
54. ? [X0 : 'Nat()'] : s(X0) = nl9 | ~$less(i(l9(zero)),alength) [ennf transformation 7]
55. ! [X0 : $int] : (a(l9(zero),X0) = a(l9(nl9),X0) | ? [X1 : 'Nat()'] : (a(l9(X1),X0) != a(l9(s(X1)),X0) & 'Sub'(X1,nl9))) [ennf transformation 6]
56. test(l9(zero)) = test(l9(nl9)) | ? [X0 : 'Nat()'] : (test(l9(X0)) != test(l9(s(X0))) & 'Sub'(X0,nl9)) [ennf transformation 5]
57. i(l9(zero)) = i(l9(nl9)) | ? [X0 : 'Nat()'] : (i(l9(X0)) != i(l9(s(X0))) & 'Sub'(X0,nl9)) [ennf transformation 4]
58. test(l8) = 1 & ! [X0 : $int] : a(l7,X0) = a(l8,X0) & i(l9(zero)) = 0 & ! [X1 : $int] : a(l8,X1) = a(l9(zero),X1) & test(l8) = test(l9(zero)) & ! [X2 : 'Nat()'] : (i(l9(X2)) = i(l11(X2)) & test(l9(X2)) = test(l11(X2)) & ! [X3 : $int] : a(l9(X2),X3) = a(l11(X2),X3)) & ! [X4 : 'Nat()'] : (v = a(l12(X4),i(l11(X4))) & ! [X5 : $int] : (a(l11(X4),X5) = a(l12(X4),X5) | i(l11(X4)) = X5) & test(l11(X4)) = test(l12(X4)) & i(l11(X4)) = i(l12(X4)) & i(l9(s(X4))) = $sum(i(l12(X4)),1) & ! [X6 : $int] : a(l12(X4),X6) = a(l9(s(X4)),X6) & test(l12(X4)) = test(l9(s(X4)))) & ! [X7 : 'Nat()'] : ($less(i(l9(X7)),alength) | ~'Sub'(X7,nl9)) & ~$less(i(l9(nl9)),alength) & i(main_end) = i(l9(nl9)) & test(main_end) = test(l9(nl9)) & ! [X8 : $int] : a(main_end,X8) = a(l9(nl9),X8) [ennf transformation 43]
59. ! [X0 : 'Nat()',X1 : 'Nat()'] : ('Sub'(X0,s(X1)) | ~'Sub'(X0,X1)) [ennf transformation 2]
60. ! [X2 : $int,X0 : $int] : (? [X3 : 'Nat()'] : (a(l9(X3),X2) = X0 & 'Sub'(X3,nl9)) => (a(l9(sK0(X0,X2)),X2) = X0 & 'Sub'(sK0(X0,X2),nl9))) [choice axiom]
61. ! [X0 : $int,X1 : 'Nat()',X2 : $int] : ((a(l9(sK0(X0,X2)),X2) = X0 & 'Sub'(sK0(X0,X2),nl9)) | $less(X0,a(l9(zero),X2)) | ~$less(X0,a(l9(nl9),X2)) | a(l9(s(X1)),X2) != $sum(a(l9(X1),X2),1)) [skolemisation 49,60]
62. ! [X0 : $int] : (? [X2 : 'Nat()'] : (test(l9(X2)) = X0 & 'Sub'(X2,nl9)) => (test(l9(sK1(X0))) = X0 & 'Sub'(sK1(X0),nl9))) [choice axiom]
63. ! [X0 : $int,X1 : 'Nat()'] : ((test(l9(sK1(X0))) = X0 & 'Sub'(sK1(X0),nl9)) | $less(X0,test(l9(zero))) | ~$less(X0,test(l9(nl9))) | test(l9(s(X1))) != $sum(test(l9(X1)),1)) [skolemisation 51,62]
64. ! [X0 : $int] : (? [X2 : 'Nat()'] : (i(l9(X2)) = X0 & 'Sub'(X2,nl9)) => (i(l9(sK2(X0))) = X0 & 'Sub'(sK2(X0),nl9))) [choice axiom]
65. ! [X0 : $int,X1 : 'Nat()'] : ((i(l9(sK2(X0))) = X0 & 'Sub'(sK2(X0),nl9)) | $less(X0,i(l9(zero))) | ~$less(X0,i(l9(nl9))) | i(l9(s(X1))) != $sum(i(l9(X1)),1)) [skolemisation 53,64]
66. ? [X0 : 'Nat()'] : s(X0) = nl9 => s(sK3) = nl9 [choice axiom]
67. s(sK3) = nl9 | ~$less(i(l9(zero)),alength) [skolemisation 54,66]
68. ! [X0 : $int] : (? [X1 : 'Nat()'] : (a(l9(X1),X0) != a(l9(s(X1)),X0) & 'Sub'(X1,nl9)) => (a(l9(sK4(X0)),X0) != a(l9(s(sK4(X0))),X0) & 'Sub'(sK4(X0),nl9))) [choice axiom]
69. ! [X0 : $int] : (a(l9(zero),X0) = a(l9(nl9),X0) | (a(l9(sK4(X0)),X0) != a(l9(s(sK4(X0))),X0) & 'Sub'(sK4(X0),nl9))) [skolemisation 55,68]
70. ? [X0 : 'Nat()'] : (test(l9(X0)) != test(l9(s(X0))) & 'Sub'(X0,nl9)) => (test(l9(sK5)) != test(l9(s(sK5))) & 'Sub'(sK5,nl9)) [choice axiom]
71. test(l9(zero)) = test(l9(nl9)) | (test(l9(sK5)) != test(l9(s(sK5))) & 'Sub'(sK5,nl9)) [skolemisation 56,70]
72. ? [X0 : 'Nat()'] : (i(l9(X0)) != i(l9(s(X0))) & 'Sub'(X0,nl9)) => (i(l9(sK6)) != i(l9(s(sK6))) & 'Sub'(sK6,nl9)) [choice axiom]
73. i(l9(zero)) = i(l9(nl9)) | (i(l9(sK6)) != i(l9(s(sK6))) & 'Sub'(sK6,nl9)) [skolemisation 57,72]
74. zero = X0 | s(p(X0)) = X0 [cnf transformation 37] {a:0,w:8,inD:0}
75. $less(2,1) [cnf transformation 40] {a:0,w:3,goal:1,inD:0}
76. X0 = X0 [cnf transformation 17] {a:0,w:3,inD:0}
77. ~$less(1,0) [cnf transformation 21] {a:0,w:3,inD:0}
78. X0 = X0 [cnf transformation 15] {a:0,w:3,inD:0}
79. $sum($sum(X0,X1),X2) = $sum($sum(X0,X2),X1) [cnf transformation 41] {a:0,w:11,inD:0}
80. test(l9(zero)) = test(l9(X0)) [cnf transformation 13] {a:0,w:7,inD:0}
81. X1 = X2 | test(l9(X1)) != test(l9(X2)) | ~'Sub'(s(X0),nl9) | test(l9(s(X0))) != $sum(test(l9(X0)),1) [cnf transformation 45] {a:0,w:24,inD:0}
82. X1 = X2 | i(l9(X1)) != i(l9(X2)) | ~'Sub'(s(X0),nl9) | i(l9(s(X0))) != $sum(i(l9(X0)),1) [cnf transformation 47] {a:0,w:24,inD:0}
83. 'Sub'(sK0(X0,X2),nl9) | $less(X0,a(l9(zero),X2)) | ~$less(X0,a(l9(nl9),X2)) | a(l9(s(X1)),X2) != $sum(a(l9(X1),X2),1) [cnf transformation 61] {a:0,w:29,inD:0}
84. a(l9(sK0(X0,X2)),X2) = X0 | $less(X0,a(l9(zero),X2)) | ~$less(X0,a(l9(nl9),X2)) | a(l9(s(X1)),X2) != $sum(a(l9(X1),X2),1) [cnf transformation 61] {a:0,w:32,inD:0}
85. 'Sub'(sK1(X0),nl9) | $less(X0,test(l9(zero))) | ~$less(X0,test(l9(nl9))) | test(l9(s(X1))) != $sum(test(l9(X1)),1) [cnf transformation 63] {a:0,w:24,inD:0}
86. test(l9(sK1(X0))) = X0 | $less(X0,test(l9(zero))) | ~$less(X0,test(l9(nl9))) | test(l9(s(X1))) != $sum(test(l9(X1)),1) [cnf transformation 63] {a:0,w:26,inD:0}
87. 'Sub'(sK2(X0),nl9) | $less(X0,i(l9(zero))) | ~$less(X0,i(l9(nl9))) | i(l9(s(X1))) != $sum(i(l9(X1)),1) [cnf transformation 65] {a:0,w:24,inD:0}
88. i(l9(sK2(X0))) = X0 | $less(X0,i(l9(zero))) | ~$less(X0,i(l9(nl9))) | i(l9(s(X1))) != $sum(i(l9(X1)),1) [cnf transformation 65] {a:0,w:26,inD:0}
89. s(sK3) = nl9 | ~$less(i(l9(zero)),alength) [cnf transformation 67] {a:0,w:9,inD:0}
90. a(l9(zero),X0) = a(l9(nl9),X0) | 'Sub'(sK4(X0),nl9) [cnf transformation 69] {a:0,w:13,inD:0}
91. a(l9(zero),X0) = a(l9(nl9),X0) | a(l9(sK4(X0)),X0) != a(l9(s(sK4(X0))),X0) [cnf transformation 69] {a:0,w:21,inD:0}
92. test(l9(zero)) = test(l9(nl9)) | 'Sub'(sK5,nl9) [cnf transformation 71] {a:0,w:10,inD:0}
93. test(l9(zero)) = test(l9(nl9)) | test(l9(sK5)) != test(l9(s(sK5))) [cnf transformation 71] {a:0,w:15,inD:0}
94. i(l9(zero)) = i(l9(nl9)) | 'Sub'(sK6,nl9) [cnf transformation 73] {a:0,w:10,inD:0}
95. i(l9(zero)) = i(l9(nl9)) | i(l9(sK6)) != i(l9(s(sK6))) [cnf transformation 73] {a:0,w:15,inD:0}
96. a(main_end,X8) = a(l9(nl9),X8) [cnf transformation 58] {a:0,w:8,inD:0}
97. test(main_end) = test(l9(nl9)) [cnf transformation 58] {a:0,w:6,inD:0}
98. i(main_end) = i(l9(nl9)) [cnf transformation 58] {a:0,w:6,inD:0}
99. ~$less(i(l9(nl9)),alength) [cnf transformation 58] {a:0,w:5,inD:0}
100. $less(i(l9(X7)),alength) | ~'Sub'(X7,nl9) [cnf transformation 58] {a:0,w:8,inD:0}
101. test(l12(X4)) = test(l9(s(X4))) [cnf transformation 58] {a:0,w:8,inD:0}
102. a(l12(X4),X6) = a(l9(s(X4)),X6) [cnf transformation 58] {a:0,w:10,inD:0}
103. i(l9(s(X4))) = $sum(i(l12(X4)),1) [cnf transformation 58] {a:0,w:10,inD:0}
104. i(l11(X4)) = i(l12(X4)) [cnf transformation 58] {a:0,w:7,inD:0}
105. test(l11(X4)) = test(l12(X4)) [cnf transformation 58] {a:0,w:7,inD:0}
106. a(l11(X4),X5) = a(l12(X4),X5) | i(l11(X4)) = X5 [cnf transformation 58] {a:0,w:14,inD:0}
107. v = a(l12(X4),i(l11(X4))) [cnf transformation 58] {a:0,w:8,inD:0}
108. a(l9(X2),X3) = a(l11(X2),X3) [cnf transformation 58] {a:0,w:9,inD:0}
109. test(l9(X2)) = test(l11(X2)) [cnf transformation 58] {a:0,w:7,inD:0}
110. i(l9(X2)) = i(l11(X2)) [cnf transformation 58] {a:0,w:7,inD:0}
111. test(l8) = test(l9(zero)) [cnf transformation 58] {a:0,w:6,inD:0}
112. a(l8,X1) = a(l9(zero),X1) [cnf transformation 58] {a:0,w:8,inD:0}
113. i(l9(zero)) = 0 [cnf transformation 58] {a:0,w:5,inD:0}
114. a(l7,X0) = a(l8,X0) [cnf transformation 58] {a:0,w:7,inD:0}
115. test(l8) = 1 [cnf transformation 58] {a:0,w:4,inD:0}
116. 'Sub'(X0,s(X1)) | ~'Sub'(X0,X1) [cnf transformation 59] {a:0,w:7,inD:0}
117. 'Sub'(X0,s(X0)) [cnf transformation 1] {a:0,w:4,inD:0}
[SA] new: 39. s(X0) != s(X1) | X0 = X1 [term algebras injectivity] {a:0,w:8,tD:1,inD:0}
[SA] new: 38. zero != s(X0) [term algebras distinctness] {a:0,w:4,tD:1,inD:0}
[SA] new: 74. zero = X0 | s(p(X0)) = X0 [cnf transformation 37] {a:0,w:8,inD:0}
[SA] new: 36. ~$less(X0,X1) | ~$less(X1,$sum(X0,1)) [theory axiom] {a:0,w:8,tD:1,inD:0}
[SA] new: 35. $uminus($uminus(X0)) = X0 [theory axiom] {a:0,w:5,tD:1,inD:0}
[SA] new: 34. $less(X0,X1) | $less(X1,$sum(X0,1)) [theory axiom] {a:0,w:8,tD:1,inD:0}
[SA] new: 33. ~$less(X0,X1) | $less($sum(X0,X2),$sum(X1,X2)) [theory axiom] {a:0,w:10,tD:1,inD:0}
[SA] new: 32. $less(X0,X1) | $less(X1,X0) | X0 = X1 [theory axiom] {a:0,w:9,tD:1,inD:0}
[SA] new: 31. ~$less(X0,X1) | ~$less(X1,X2) | $less(X0,X2) [theory axiom] {a:0,w:9,tD:1,inD:0}
[SA] new: 30. ~$less(X0,X0) [theory axiom] {a:0,w:3,tD:1,inD:0}
[SA] new: 29. $sum(X0,$uminus(X0)) = 0 [theory axiom] {a:0,w:6,tD:1,inD:0}
[SA] new: 28. $sum($uminus(X1),$uminus(X0)) = $uminus($sum(X0,X1)) [theory axiom] {a:0,w:10,tD:1,inD:0}
[SA] new: 27. $sum(X0,0) = X0 [theory axiom] {a:0,w:5,tD:1,inD:0}
[SA] new: 26. $sum(X0,$sum(X1,X2)) = $sum($sum(X0,X1),X2) [theory axiom] {a:0,w:11,tD:1,inD:0}
[SA] new: 25. $sum(X0,X1) = $sum(X1,X0) [theory axiom] {a:0,w:7,tD:1,inD:0}
[SA] new: 75. $less(2,1) [cnf transformation 40] {a:0,w:3,goal:1,inD:0}
[SA] new: 76. X0 = X0 [cnf transformation 17] {a:0,w:3,inD:0}
[SA] new: 77. ~$less(1,0) [cnf transformation 21] {a:0,w:3,inD:0}
[SA] new: 78. X0 = X0 [cnf transformation 15] {a:0,w:3,inD:0}
[SA] new: 79. $sum($sum(X0,X1),X2) = $sum($sum(X0,X2),X1) [cnf transformation 41] {a:0,w:11,inD:0}
[SA] new: 80. test(l9(zero)) = test(l9(X0)) [cnf transformation 13] {a:0,w:7,inD:0}
[SA] new: 81. X1 = X2 | test(l9(X1)) != test(l9(X2)) | ~'Sub'(s(X0),nl9) | test(l9(s(X0))) != $sum(test(l9(X0)),1) [cnf transformation 45] {a:0,w:24,inD:0}
[SA] new: 82. X1 = X2 | i(l9(X1)) != i(l9(X2)) | ~'Sub'(s(X0),nl9) | i(l9(s(X0))) != $sum(i(l9(X0)),1) [cnf transformation 47] {a:0,w:24,inD:0}
[SA] new: 84. a(l9(sK0(X0,X2)),X2) = X0 | $less(X0,a(l9(zero),X2)) | ~$less(X0,a(l9(nl9),X2)) | a(l9(s(X1)),X2) != $sum(a(l9(X1),X2),1) [cnf transformation 61] {a:0,w:32,inD:0}
[SA] new: 83. 'Sub'(sK0(X0,X2),nl9) | $less(X0,a(l9(zero),X2)) | ~$less(X0,a(l9(nl9),X2)) | a(l9(s(X1)),X2) != $sum(a(l9(X1),X2),1) [cnf transformation 61] {a:0,w:29,inD:0}
[SA] new: 86. test(l9(sK1(X0))) = X0 | $less(X0,test(l9(zero))) | ~$less(X0,test(l9(nl9))) | test(l9(s(X1))) != $sum(test(l9(X1)),1) [cnf transformation 63] {a:0,w:26,inD:0}
[SA] new: 85. 'Sub'(sK1(X0),nl9) | $less(X0,test(l9(zero))) | ~$less(X0,test(l9(nl9))) | test(l9(s(X1))) != $sum(test(l9(X1)),1) [cnf transformation 63] {a:0,w:24,inD:0}
[SA] new: 88. i(l9(sK2(X0))) = X0 | $less(X0,i(l9(zero))) | ~$less(X0,i(l9(nl9))) | i(l9(s(X1))) != $sum(i(l9(X1)),1) [cnf transformation 65] {a:0,w:26,inD:0}
[SA] new: 87. 'Sub'(sK2(X0),nl9) | $less(X0,i(l9(zero))) | ~$less(X0,i(l9(nl9))) | i(l9(s(X1))) != $sum(i(l9(X1)),1) [cnf transformation 65] {a:0,w:24,inD:0}
[SA] new: 89. s(sK3) = nl9 | ~$less(i(l9(zero)),alength) [cnf transformation 67] {a:0,w:9,inD:0}
[SA] new: 91. a(l9(zero),X0) = a(l9(nl9),X0) | a(l9(sK4(X0)),X0) != a(l9(s(sK4(X0))),X0) [cnf transformation 69] {a:0,w:21,inD:0}
[SA] new: 90. a(l9(zero),X0) = a(l9(nl9),X0) | 'Sub'(sK4(X0),nl9) [cnf transformation 69] {a:0,w:13,inD:0}
[SA] new: 93. test(l9(zero)) = test(l9(nl9)) | test(l9(sK5)) != test(l9(s(sK5))) [cnf transformation 71] {a:0,w:15,inD:0}
[SA] new: 92. test(l9(zero)) = test(l9(nl9)) | 'Sub'(sK5,nl9) [cnf transformation 71] {a:0,w:10,inD:0}
[SA] new: 95. i(l9(zero)) = i(l9(nl9)) | i(l9(sK6)) != i(l9(s(sK6))) [cnf transformation 73] {a:0,w:15,inD:0}
[SA] new: 94. i(l9(zero)) = i(l9(nl9)) | 'Sub'(sK6,nl9) [cnf transformation 73] {a:0,w:10,inD:0}
[SA] new: 115. test(l8) = 1 [cnf transformation 58] {a:0,w:4,inD:0}
[SA] new: 114. a(l7,X0) = a(l8,X0) [cnf transformation 58] {a:0,w:7,inD:0}
[SA] new: 113. i(l9(zero)) = 0 [cnf transformation 58] {a:0,w:5,inD:0}
[SA] new: 112. a(l8,X1) = a(l9(zero),X1) [cnf transformation 58] {a:0,w:8,inD:0}
[SA] new: 111. test(l8) = test(l9(zero)) [cnf transformation 58] {a:0,w:6,inD:0}
[SA] new: 110. i(l9(X2)) = i(l11(X2)) [cnf transformation 58] {a:0,w:7,inD:0}
[SA] new: 109. test(l9(X2)) = test(l11(X2)) [cnf transformation 58] {a:0,w:7,inD:0}
[SA] new: 108. a(l9(X2),X3) = a(l11(X2),X3) [cnf transformation 58] {a:0,w:9,inD:0}
[SA] new: 107. v = a(l12(X4),i(l11(X4))) [cnf transformation 58] {a:0,w:8,inD:0}
[SA] new: 106. a(l11(X4),X5) = a(l12(X4),X5) | i(l11(X4)) = X5 [cnf transformation 58] {a:0,w:14,inD:0}
[SA] new: 105. test(l11(X4)) = test(l12(X4)) [cnf transformation 58] {a:0,w:7,inD:0}
[SA] new: 104. i(l11(X4)) = i(l12(X4)) [cnf transformation 58] {a:0,w:7,inD:0}
[SA] new: 103. i(l9(s(X4))) = $sum(i(l12(X4)),1) [cnf transformation 58] {a:0,w:10,inD:0}
[SA] new: 102. a(l12(X4),X6) = a(l9(s(X4)),X6) [cnf transformation 58] {a:0,w:10,inD:0}
[SA] new: 101. test(l12(X4)) = test(l9(s(X4))) [cnf transformation 58] {a:0,w:8,inD:0}
[SA] new: 100. $less(i(l9(X7)),alength) | ~'Sub'(X7,nl9) [cnf transformation 58] {a:0,w:8,inD:0}
[SA] new: 99. ~$less(i(l9(nl9)),alength) [cnf transformation 58] {a:0,w:5,inD:0}
[SA] new: 98. i(main_end) = i(l9(nl9)) [cnf transformation 58] {a:0,w:6,inD:0}
[SA] new: 97. test(main_end) = test(l9(nl9)) [cnf transformation 58] {a:0,w:6,inD:0}
[SA] new: 96. a(main_end,X8) = a(l9(nl9),X8) [cnf transformation 58] {a:0,w:8,inD:0}
[SA] new: 116. 'Sub'(X0,s(X1)) | ~'Sub'(X0,X1) [cnf transformation 59] {a:0,w:7,inD:0}
[SA] new: 117. 'Sub'(X0,s(X0)) [cnf transformation 1] {a:0,w:4,inD:0}
[SA] passive: 39. s(X0) != s(X1) | X0 = X1 [term algebras injectivity] {a:0,w:8,tD:1,inD:0}
[SA] passive: 38. zero != s(X0) [term algebras distinctness] {a:0,w:4,tD:1,inD:0}
[SA] passive: 74. zero = X0 | s(p(X0)) = X0 [cnf transformation 37] {a:0,w:8,inD:0}
[SA] passive: 36. ~$less(X0,X1) | ~$less(X1,$sum(X0,1)) [theory axiom] {a:0,w:8,tD:1,inD:0}
[SA] passive: 35. $uminus($uminus(X0)) = X0 [theory axiom] {a:0,w:5,tD:1,inD:0}
[SA] passive: 34. $less(X0,X1) | $less(X1,$sum(X0,1)) [theory axiom] {a:0,w:8,tD:1,inD:0}
[SA] passive: 33. ~$less(X0,X1) | $less($sum(X0,X2),$sum(X1,X2)) [theory axiom] {a:0,w:10,tD:1,inD:0}
[SA] passive: 32. $less(X0,X1) | $less(X1,X0) | X0 = X1 [theory axiom] {a:0,w:9,tD:1,inD:0}
[SA] passive: 31. ~$less(X0,X1) | ~$less(X1,X2) | $less(X0,X2) [theory axiom] {a:0,w:9,tD:1,inD:0}
[SA] passive: 30. ~$less(X0,X0) [theory axiom] {a:0,w:3,tD:1,inD:0}
[SA] passive: 29. $sum(X0,$uminus(X0)) = 0 [theory axiom] {a:0,w:6,tD:1,inD:0}
[SA] passive: 28. $sum($uminus(X1),$uminus(X0)) = $uminus($sum(X0,X1)) [theory axiom] {a:0,w:10,tD:1,inD:0}
[SA] passive: 27. $sum(X0,0) = X0 [theory axiom] {a:0,w:5,tD:1,inD:0}
[SA] passive: 26. $sum(X0,$sum(X1,X2)) = $sum($sum(X0,X1),X2) [theory axiom] {a:0,w:11,tD:1,inD:0}
[SA] passive: 25. $sum(X0,X1) = $sum(X1,X0) [theory axiom] {a:0,w:7,tD:1,inD:0}
[SA] passive: 75. $less(2,1) [cnf transformation 40] {a:0,w:3,goal:1,inD:0}
[SA] passive: 77. ~$less(1,0) [cnf transformation 21] {a:0,w:3,inD:0}
[SA] new: 118. $sum(X0,$sum(X2,X1)) = $sum($sum(X0,X1),X2) [forward demodulation 79,26] {a:0,w:11,inD:0}
[SA] passive: 118. $sum(X0,$sum(X2,X1)) = $sum($sum(X0,X1),X2) [forward demodulation 79,26] {a:0,w:11,inD:0}
[SA] passive: 80. test(l9(zero)) = test(l9(X0)) [cnf transformation 13] {a:0,w:7,inD:0}
[SA] new: 119. test(l9(s(X0))) != $sum(1,test(l9(X0))) | X1 = X2 | test(l9(X1)) != test(l9(X2)) | ~'Sub'(s(X0),nl9) [forward demodulation 81,25] {a:0,w:24,inD:0}
[SA] new: 120. test(l9(zero)) != $sum(1,test(l9(X0))) | X1 = X2 | test(l9(X1)) != test(l9(X2)) | ~'Sub'(s(X0),nl9) [forward demodulation 119,80] {a:0,w:23,inD:0}
[SA] passive: 120. test(l9(zero)) != $sum(1,test(l9(X0))) | X1 = X2 | test(l9(X1)) != test(l9(X2)) | ~'Sub'(s(X0),nl9) [forward demodulation 119,80] {a:0,w:23,inD:0}
[SA] new: 121. i(l9(s(X0))) != $sum(1,i(l9(X0))) | X1 = X2 | i(l9(X1)) != i(l9(X2)) | ~'Sub'(s(X0),nl9) [forward demodulation 82,25] {a:0,w:24,inD:0}
[SA] passive: 121. i(l9(s(X0))) != $sum(1,i(l9(X0))) | X1 = X2 | i(l9(X1)) != i(l9(X2)) | ~'Sub'(s(X0),nl9) [forward demodulation 82,25] {a:0,w:24,inD:0}
[SA] new: 122. a(l9(s(X1)),X2) != $sum(1,a(l9(X1),X2)) | a(l9(sK0(X0,X2)),X2) = X0 | $less(X0,a(l9(zero),X2)) | ~$less(X0,a(l9(nl9),X2)) [forward demodulation 84,25] {a:0,w:32,inD:0}
[SA] passive: 122. a(l9(s(X1)),X2) != $sum(1,a(l9(X1),X2)) | a(l9(sK0(X0,X2)),X2) = X0 | $less(X0,a(l9(zero),X2)) | ~$less(X0,a(l9(nl9),X2)) [forward demodulation 84,25] {a:0,w:32,inD:0}
[SA] new: 123. a(l9(s(X1)),X2) != $sum(1,a(l9(X1),X2)) | 'Sub'(sK0(X0,X2),nl9) | $less(X0,a(l9(zero),X2)) | ~$less(X0,a(l9(nl9),X2)) [forward demodulation 83,25] {a:0,w:29,inD:0}
[SA] passive: 123. a(l9(s(X1)),X2) != $sum(1,a(l9(X1),X2)) | 'Sub'(sK0(X0,X2),nl9) | $less(X0,a(l9(zero),X2)) | ~$less(X0,a(l9(nl9),X2)) [forward demodulation 83,25] {a:0,w:29,inD:0}
[SA] new: 124. test(l9(zero)) = X0 | $less(X0,test(l9(zero))) | ~$less(X0,test(l9(nl9))) | test(l9(s(X1))) != $sum(test(l9(X1)),1) [forward demodulation 86,80] {a:0,w:25,inD:0}
[SA] new: 125. ~$less(X0,test(l9(zero))) | test(l9(zero)) = X0 | $less(X0,test(l9(zero))) | test(l9(s(X1))) != $sum(test(l9(X1)),1) [forward demodulation 124,80] {a:0,w:25,inD:0}
[SA] new: 126. ~$less(X0,test(l9(zero))) | 'Sub'(sK1(X0),nl9) | $less(X0,test(l9(zero))) | test(l9(s(X1))) != $sum(test(l9(X1)),1) [forward demodulation 85,80] {a:0,w:24,inD:0}
[SA] new: 127. i(l9(s(X1))) != $sum(1,i(l9(X1))) | i(l9(sK2(X0))) = X0 | $less(X0,i(l9(zero))) | ~$less(X0,i(l9(nl9))) [forward demodulation 88,25] {a:0,w:26,inD:0}
[SA] passive: 127. i(l9(s(X1))) != $sum(1,i(l9(X1))) | i(l9(sK2(X0))) = X0 | $less(X0,i(l9(zero))) | ~$less(X0,i(l9(nl9))) [forward demodulation 88,25] {a:0,w:26,inD:0}
[SA] new: 128. i(l9(s(X1))) != $sum(1,i(l9(X1))) | 'Sub'(sK2(X0),nl9) | $less(X0,i(l9(zero))) | ~$less(X0,i(l9(nl9))) [forward demodulation 87,25] {a:0,w:24,inD:0}
[SA] passive: 128. i(l9(s(X1))) != $sum(1,i(l9(X1))) | 'Sub'(sK2(X0),nl9) | $less(X0,i(l9(zero))) | ~$less(X0,i(l9(nl9))) [forward demodulation 87,25] {a:0,w:24,inD:0}
[SA] passive: 89. s(sK3) = nl9 | ~$less(i(l9(zero)),alength) [cnf transformation 67] {a:0,w:9,inD:0}
[SA] passive: 91. a(l9(zero),X0) = a(l9(nl9),X0) | a(l9(sK4(X0)),X0) != a(l9(s(sK4(X0))),X0) [cnf transformation 69] {a:0,w:21,inD:0}
[SA] passive: 90. a(l9(zero),X0) = a(l9(nl9),X0) | 'Sub'(sK4(X0),nl9) [cnf transformation 69] {a:0,w:13,inD:0}
[SA] passive: 95. i(l9(zero)) = i(l9(nl9)) | i(l9(sK6)) != i(l9(s(sK6))) [cnf transformation 73] {a:0,w:15,inD:0}
[SA] passive: 94. i(l9(zero)) = i(l9(nl9)) | 'Sub'(sK6,nl9) [cnf transformation 73] {a:0,w:10,inD:0}
[SA] passive: 115. test(l8) = 1 [cnf transformation 58] {a:0,w:4,inD:0}
[SA] passive: 114. a(l7,X0) = a(l8,X0) [cnf transformation 58] {a:0,w:7,inD:0}
[SA] new: 129. i(l9(nl9)) = 0 | 'Sub'(sK6,nl9) [backward demodulation 113,94] {a:0,w:8,inD:0}
[SA] new: 130. i(l9(nl9)) = 0 | i(l9(sK6)) != i(l9(s(sK6))) [backward demodulation 113,95] {a:0,w:13,inD:0}
[SA] new: 131. ~$less(0,alength) | s(sK3) = nl9 [backward demodulation 113,89] {a:0,w:7,inD:0}
[SA] new: 132. $less(X0,0) | i(l9(s(X1))) != $sum(1,i(l9(X1))) | 'Sub'(sK2(X0),nl9) | ~$less(X0,i(l9(nl9))) [backward demodulation 113,128] {a:0,w:22,inD:0}
[SA] new: 133. $less(X0,0) | i(l9(s(X1))) != $sum(1,i(l9(X1))) | i(l9(sK2(X0))) = X0 | ~$less(X0,i(l9(nl9))) [backward demodulation 113,127] {a:0,w:24,inD:0}
[SA] passive: 113. i(l9(zero)) = 0 [cnf transformation 58] {a:0,w:5,inD:0}
[SA] passive: 129. i(l9(nl9)) = 0 | 'Sub'(sK6,nl9) [backward demodulation 113,94] {a:0,w:8,inD:0}
[SA] passive: 130. i(l9(nl9)) = 0 | i(l9(sK6)) != i(l9(s(sK6))) [backward demodulation 113,95] {a:0,w:13,inD:0}
[SA] passive: 131. ~$less(0,alength) | s(sK3) = nl9 [backward demodulation 113,89] {a:0,w:7,inD:0}
[SA] passive: 132. $less(X0,0) | i(l9(s(X1))) != $sum(1,i(l9(X1))) | 'Sub'(sK2(X0),nl9) | ~$less(X0,i(l9(nl9))) [backward demodulation 113,128] {a:0,w:22,inD:0}
[SA] passive: 133. $less(X0,0) | i(l9(s(X1))) != $sum(1,i(l9(X1))) | i(l9(sK2(X0))) = X0 | ~$less(X0,i(l9(nl9))) [backward demodulation 113,127] {a:0,w:24,inD:0}
[SA] new: 134. a(l7,X1) = a(l9(zero),X1) [forward demodulation 112,114] {a:0,w:8,inD:0}
[SA] new: 135. a(l7,X0) = a(l9(nl9),X0) | 'Sub'(sK4(X0),nl9) [backward demodulation 134,90] {a:0,w:12,inD:0}
[SA] new: 136. a(l7,X0) = a(l9(nl9),X0) | a(l9(sK4(X0)),X0) != a(l9(s(sK4(X0))),X0) [backward demodulation 134,91] {a:0,w:20,inD:0}
[SA] new: 137. $less(X0,a(l7,X2)) | a(l9(s(X1)),X2) != $sum(1,a(l9(X1),X2)) | 'Sub'(sK0(X0,X2),nl9) | ~$less(X0,a(l9(nl9),X2)) [backward demodulation 134,123] {a:0,w:28,inD:0}
[SA] new: 138. $less(X0,a(l7,X2)) | a(l9(s(X1)),X2) != $sum(1,a(l9(X1),X2)) | a(l9(sK0(X0,X2)),X2) = X0 | ~$less(X0,a(l9(nl9),X2)) [backward demodulation 134,122] {a:0,w:31,inD:0}
[SA] passive: 134. a(l7,X1) = a(l9(zero),X1) [forward demodulation 112,114] {a:0,w:8,inD:0}
[SA] passive: 135. a(l7,X0) = a(l9(nl9),X0) | 'Sub'(sK4(X0),nl9) [backward demodulation 134,90] {a:0,w:12,inD:0}
[SA] passive: 136. a(l7,X0) = a(l9(nl9),X0) | a(l9(sK4(X0)),X0) != a(l9(s(sK4(X0))),X0) [backward demodulation 134,91] {a:0,w:20,inD:0}
[SA] passive: 137. $less(X0,a(l7,X2)) | a(l9(s(X1)),X2) != $sum(1,a(l9(X1),X2)) | 'Sub'(sK0(X0,X2),nl9) | ~$less(X0,a(l9(nl9),X2)) [backward demodulation 134,123] {a:0,w:28,inD:0}
[SA] passive: 138. $less(X0,a(l7,X2)) | a(l9(s(X1)),X2) != $sum(1,a(l9(X1),X2)) | a(l9(sK0(X0,X2)),X2) = X0 | ~$less(X0,a(l9(nl9),X2)) [backward demodulation 134,122] {a:0,w:31,inD:0}
[SA] new: 139. test(l9(zero)) = 1 [forward demodulation 111,115] {a:0,w:5,inD:0}
[SA] new: 140. 1 != $sum(1,test(l9(X0))) | X1 = X2 | test(l9(X1)) != test(l9(X2)) | ~'Sub'(s(X0),nl9) [backward demodulation 139,120] {a:0,w:21,inD:0}
[SA] new: 141. test(l9(X0)) = 1 [backward demodulation 139,80] {a:0,w:5,inD:0}
[SA] passive: 139. test(l9(zero)) = 1 [forward demodulation 111,115] {a:0,w:5,inD:0}
[SA] new: 142. test(l9(X0)) != $sum(1,-1) | X1 = X2 | test(l9(X1)) != test(l9(X2)) | ~'Sub'(s(X0),nl9) [evaluation 140] {a:0,w:21,inD:0}
[SA] new: 143. test(l9(X0)) != 0 | X1 = X2 | test(l9(X1)) != test(l9(X2)) | ~'Sub'(s(X0),nl9) [evaluation 142] {a:0,w:19,inD:0}
[SA] passive: 143. test(l9(X0)) != 0 | X1 = X2 | test(l9(X1)) != test(l9(X2)) | ~'Sub'(s(X0),nl9) [evaluation 142] {a:0,w:19,inD:0}
[SA] new: 144. test(l9(X2)) != 1 | test(l9(X0)) != 0 | X1 = X2 | ~'Sub'(s(X0),nl9) [backward demodulation 141,143] {a:0,w:17,inD:0}
[SA] passive: 141. test(l9(X0)) = 1 [backward demodulation 139,80] {a:0,w:5,inD:0}
[SA] new: 145. test(l9(X0)) != 0 | X1 = X2 | ~'Sub'(s(X0),nl9) [subsumption resolution 144,141] {a:0,w:12,inD:0}
[SA] new: 146. 1 != 0 | X1 = X2 | ~'Sub'(s(X0),nl9) [forward demodulation 145,141] {a:0,w:10,inD:0}
[SA] passive: 146. 1 != 0 | X1 = X2 | ~'Sub'(s(X0),nl9) [forward demodulation 145,141] {a:0,w:10,inD:0}
[SA] passive: 110. i(l9(X2)) = i(l11(X2)) [cnf transformation 58] {a:0,w:7,inD:0}
[SA] new: 147. test(l11(X2)) = 1 [forward demodulation 109,141] {a:0,w:5,inD:0}
[SA] passive: 147. test(l11(X2)) = 1 [forward demodulation 109,141] {a:0,w:5,inD:0}
[SA] passive: 108. a(l9(X2),X3) = a(l11(X2),X3) [cnf transformation 58] {a:0,w:9,inD:0}
[SA] new: 148. v = a(l12(X4),i(l9(X4))) [forward demodulation 107,110] {a:0,w:8,inD:0}
[SA] passive: 148. v = a(l12(X4),i(l9(X4))) [forward demodulation 107,110] {a:0,w:8,inD:0}
[SA] new: 149. a(l9(X4),X5) = a(l12(X4),X5) | i(l11(X4)) = X5 [forward demodulation 106,108] {a:0,w:14,inD:0}
[SA] new: 150. i(l9(X4)) = X5 | a(l9(X4),X5) = a(l12(X4),X5) [forward demodulation 149,110] {a:0,w:14,inD:0}
[SA] passive: 150. i(l9(X4)) = X5 | a(l9(X4),X5) = a(l12(X4),X5) [forward demodulation 149,110] {a:0,w:14,inD:0}
[SA] new: 151. test(l12(X4)) = 1 [forward demodulation 105,147] {a:0,w:5,inD:0}
[SA] passive: 151. test(l12(X4)) = 1 [forward demodulation 105,147] {a:0,w:5,inD:0}
[SA] new: 152. i(l9(X4)) = i(l12(X4)) [forward demodulation 104,110] {a:0,w:7,inD:0}
[SA] passive: 152. i(l9(X4)) = i(l12(X4)) [forward demodulation 104,110] {a:0,w:7,inD:0}
[SA] new: 153. i(l9(s(X4))) = $sum(i(l9(X4)),1) [forward demodulation 103,152] {a:0,w:10,inD:0}
[SA] passive: 153. i(l9(s(X4))) = $sum(i(l9(X4)),1) [forward demodulation 103,152] {a:0,w:10,inD:0}
[SA] new: 154. a(l12(X1),X2) != $sum(1,a(l9(X1),X2)) | $less(X0,a(l7,X2)) | a(l9(sK0(X0,X2)),X2) = X0 | ~$less(X0,a(l9(nl9),X2)) [backward demodulation 102,138] {a:0,w:30,inD:0}
[SA] new: 155. a(l12(X1),X2) != $sum(1,a(l9(X1),X2)) | $less(X0,a(l7,X2)) | 'Sub'(sK0(X0,X2),nl9) | ~$less(X0,a(l9(nl9),X2)) [backward demodulation 102,137] {a:0,w:27,inD:0}
[SA] passive: 102. a(l12(X4),X6) = a(l9(s(X4)),X6) [cnf transformation 58] {a:0,w:10,inD:0}
[SA] passive: 154. a(l12(X1),X2) != $sum(1,a(l9(X1),X2)) | $less(X0,a(l7,X2)) | a(l9(sK0(X0,X2)),X2) = X0 | ~$less(X0,a(l9(nl9),X2)) [backward demodulation 102,138] {a:0,w:30,inD:0}
[SA] passive: 155. a(l12(X1),X2) != $sum(1,a(l9(X1),X2)) | $less(X0,a(l7,X2)) | 'Sub'(sK0(X0,X2),nl9) | ~$less(X0,a(l9(nl9),X2)) [backward demodulation 102,137] {a:0,w:27,inD:0}
[SA] new: 156. test(l12(X4)) = 1 [forward demodulation 101,141] {a:0,w:5,inD:0}
[SA] passive: 100. $less(i(l9(X7)),alength) | ~'Sub'(X7,nl9) [cnf transformation 58] {a:0,w:8,inD:0}
[SA] passive: 99. ~$less(i(l9(nl9)),alength) [cnf transformation 58] {a:0,w:5,inD:0}
[SA] new: 157. ~$less(i(main_end),alength) [backward demodulation 98,99] {a:0,w:4,inD:0}
[SA] new: 158. ~$less(X0,i(main_end)) | $less(X0,0) | i(l9(s(X1))) != $sum(1,i(l9(X1))) | i(l9(sK2(X0))) = X0 [backward demodulation 98,133] {a:0,w:23,inD:0}
[SA] new: 159. ~$less(X0,i(main_end)) | $less(X0,0) | i(l9(s(X1))) != $sum(1,i(l9(X1))) | 'Sub'(sK2(X0),nl9) [backward demodulation 98,132] {a:0,w:21,inD:0}
[SA] new: 160. i(main_end) = 0 | i(l9(sK6)) != i(l9(s(sK6))) [backward demodulation 98,130] {a:0,w:12,inD:0}
[SA] new: 161. i(main_end) = 0 | 'Sub'(sK6,nl9) [backward demodulation 98,129] {a:0,w:7,inD:0}
[SA] passive: 98. i(main_end) = i(l9(nl9)) [cnf transformation 58] {a:0,w:6,inD:0}
[SA] passive: 157. ~$less(i(main_end),alength) [backward demodulation 98,99] {a:0,w:4,inD:0}
[SA] passive: 158. ~$less(X0,i(main_end)) | $less(X0,0) | i(l9(s(X1))) != $sum(1,i(l9(X1))) | i(l9(sK2(X0))) = X0 [backward demodulation 98,133] {a:0,w:23,inD:0}
[SA] passive: 159. ~$less(X0,i(main_end)) | $less(X0,0) | i(l9(s(X1))) != $sum(1,i(l9(X1))) | 'Sub'(sK2(X0),nl9) [backward demodulation 98,132] {a:0,w:21,inD:0}
[SA] passive: 160. i(main_end) = 0 | i(l9(sK6)) != i(l9(s(sK6))) [backward demodulation 98,130] {a:0,w:12,inD:0}
[SA] passive: 161. i(main_end) = 0 | 'Sub'(sK6,nl9) [backward demodulation 98,129] {a:0,w:7,inD:0}
[SA] new: 162. test(main_end) = 1 [forward demodulation 97,141] {a:0,w:4,inD:0}
[SA] passive: 162. test(main_end) = 1 [forward demodulation 97,141] {a:0,w:4,inD:0}
[SA] new: 163. ~$less(X0,a(main_end,X2)) | a(l12(X1),X2) != $sum(1,a(l9(X1),X2)) | $less(X0,a(l7,X2)) | 'Sub'(sK0(X0,X2),nl9) [backward demodulation 96,155] {a:0,w:26,inD:0}
[SA] new: 164. ~$less(X0,a(main_end,X2)) | a(l12(X1),X2) != $sum(1,a(l9(X1),X2)) | $less(X0,a(l7,X2)) | a(l9(sK0(X0,X2)),X2) = X0 [backward demodulation 96,154] {a:0,w:29,inD:0}
[SA] new: 165. a(main_end,X0) = a(l7,X0) | a(l9(sK4(X0)),X0) != a(l9(s(sK4(X0))),X0) [backward demodulation 96,136] {a:0,w:19,inD:0}
[SA] new: 166. a(main_end,X0) = a(l7,X0) | 'Sub'(sK4(X0),nl9) [backward demodulation 96,135] {a:0,w:11,inD:0}
[SA] passive: 96. a(main_end,X8) = a(l9(nl9),X8) [cnf transformation 58] {a:0,w:8,inD:0}
[SA] passive: 163. ~$less(X0,a(main_end,X2)) | a(l12(X1),X2) != $sum(1,a(l9(X1),X2)) | $less(X0,a(l7,X2)) | 'Sub'(sK0(X0,X2),nl9) [backward demodulation 96,155] {a:0,w:26,inD:0}
[SA] passive: 164. ~$less(X0,a(main_end,X2)) | a(l12(X1),X2) != $sum(1,a(l9(X1),X2)) | $less(X0,a(l7,X2)) | a(l9(sK0(X0,X2)),X2) = X0 [backward demodulation 96,154] {a:0,w:29,inD:0}
[SA] passive: 165. a(main_end,X0) = a(l7,X0) | a(l9(sK4(X0)),X0) != a(l9(s(sK4(X0))),X0) [backward demodulation 96,136] {a:0,w:19,inD:0}
[SA] passive: 166. a(main_end,X0) = a(l7,X0) | 'Sub'(sK4(X0),nl9) [backward demodulation 96,135] {a:0,w:11,inD:0}
[SA] passive: 116. 'Sub'(X0,s(X1)) | ~'Sub'(X0,X1) [cnf transformation 59] {a:0,w:7,inD:0}
[SA] passive: 117. 'Sub'(X0,s(X0)) [cnf transformation 1] {a:0,w:4,inD:0}
[SA] active: 75. $less(2,1) [cnf transformation 40] {a:0,w:3,nSel:1,goal:1,inD:0}
[SA] active: 77. ~$less(1,0) [cnf transformation 21] {a:0,w:3,nSel:1,inD:0}
[SA] active: 30. ~$less(X0,X0) [theory axiom] {a:0,w:3,nSel:1,tD:1,inD:0}
[SA] active: 115. test(l8) = 1 [cnf transformation 58] {a:0,w:4,nSel:1,inD:0}
[SA] active: 117. 'Sub'(X0,s(X0)) [cnf transformation 1] {a:0,w:4,nSel:1,inD:0}
[SA] active: 157. ~$less(i(main_end),alength) [backward demodulation 98,99] {a:0,w:4,nSel:1,inD:0}
[SA] active: 162. test(main_end) = 1 [forward demodulation 97,141] {a:0,w:4,nSel:1,inD:0}
[SA] active: 38. zero != s(X0) [term algebras distinctness] {a:0,w:4,nSel:1,tD:1,inD:0}
[SA] active: 113. i(l9(zero)) = 0 [cnf transformation 58] {a:0,w:5,nSel:1,inD:0}
[SA] active: 139. test(l9(zero)) = 1 [forward demodulation 111,115] {a:0,w:5,nSel:1,inD:0}
[SA] active: 141. test(l9(X0)) = 1 [backward demodulation 139,80] {a:0,w:5,nSel:1,inD:0}
[SA] active: 147. test(l11(X2)) = 1 [forward demodulation 109,141] {a:0,w:5,nSel:1,inD:0}
[SA] active: 151. test(l12(X4)) = 1 [forward demodulation 105,147] {a:0,w:5,nSel:1,inD:0}
[SA] active: 27. $sum(X0,0) = X0 [theory axiom] {a:0,w:5,nSel:1,tD:1,inD:0}
[SA] active: 35. $uminus($uminus(X0)) = X0 [theory axiom] {a:0,w:5,nSel:1,tD:1,inD:0}
[SA] active: 98. i(main_end) = i(l9(nl9)) [cnf transformation 58] {a:0,w:6,nSel:1,inD:0}
[SA] active: 29. $sum(X0,$uminus(X0)) = 0 [theory axiom] {a:0,w:6,nSel:1,tD:1,inD:0}
[SA] new: 167. $sum($uminus(X0),X0) = 0 [superposition 29,35] {a:1,w:6,tD:1,inD:0}
[SA] passive: 167. $sum($uminus(X0),X0) = 0 [superposition 29,35] {a:1,w:6,tD:1,inD:0}
[SA] active: 110. i(l9(X2)) = i(l11(X2)) [cnf transformation 58] {a:0,w:7,nSel:1,inD:0}
[SA] active: 167. $sum($uminus(X0),X0) = 0 [superposition 29,35] {a:1,w:6,nSel:1,tD:1,inD:0}
[SA] new: 168. $sum(X0,$uminus(X0)) = 0 [superposition 167,35] {a:2,w:6,tD:1,inD:0}
[SA] new: 169. 0 = $uminus(0) [superposition 167,27] {a:2,w:4,tD:1,inD:0}
[SA] new: 170. 0 = $uminus(0) [superposition 27,167] {a:2,w:4,tD:1,inD:0}
[SA] new: 171. 0 = 0 [evaluation 170] {a:2,w:3,tD:1,inD:0}
[SA] new: 172. 0 = 0 [evaluation 169] {a:2,w:3,tD:1,inD:0}
[SA] active: 114. a(l7,X0) = a(l8,X0) [cnf transformation 58] {a:0,w:7,nSel:1,inD:0}
[SA] active: 116. 'Sub'(X0,s(X1)) | ~'Sub'(X0,X1) [cnf transformation 59] {a:0,w:7,nSel:1,inD:0}
[SA] active: 131. ~$less(0,alength) | s(sK3) = nl9 [backward demodulation 113,89] {a:0,w:7,nSel:1,inD:0}
[SA] active: 152. i(l9(X4)) = i(l12(X4)) [forward demodulation 104,110] {a:0,w:7,nSel:1,inD:0}
[SA] active: 161. 'Sub'(sK6,nl9) | i(main_end) = 0 [backward demodulation 98,129] {a:0,w:7,nSel:1,inD:0}
[SA] active: 25. $sum(X0,X1) = $sum(X1,X0) [theory axiom] {a:0,w:7,nSel:1,tD:1,inD:0}
[SA] new: 173. $sum(0,X0) = X0 [superposition 25,27] {a:1,w:5,tD:1,inD:0}
[SA] new: 174. $sum(X6,$uminus(X6)) = 0 [superposition 25,167] {a:2,w:6,tD:1,inD:0}
[SA] new: 175. $sum(0,X0) = X0 [superposition 25,27] {a:1,w:5,tD:1,inD:0}
[SA] new: 176. $sum(X6,$uminus(X6)) = 0 [superposition 25,167] {a:2,w:6,tD:1,inD:0}
[SA] new: 177. $sum(0,X0) = X0 [superposition 27,25] {a:1,w:5,tD:1,inD:0}
[SA] new: 178. $sum(X2,$uminus(X2)) = 0 [superposition 167,25] {a:2,w:6,tD:1,inD:0}
[SA] new: 179. $sum(0,X0) = X0 [superposition 27,25] {a:1,w:5,tD:1,inD:0}
[SA] new: 180. $sum(X2,$uminus(X2)) = 0 [superposition 167,25] {a:2,w:6,tD:1,inD:0}
[SA] passive: 173. $sum(0,X0) = X0 [superposition 25,27] {a:1,w:5,tD:1,inD:0}
[SA] active: 96. a(main_end,X8) = a(l9(nl9),X8) [cnf transformation 58] {a:0,w:8,nSel:1,inD:0}
[SA] active: 173. $sum(0,X0) = X0 [superposition 25,27] {a:1,w:5,nSel:1,tD:1,inD:0}
[SA] new: 181. 0 = $uminus(0) [superposition 173,29] {a:2,w:4,tD:1,inD:0}
[SA] new: 182. $sum(X1,0) = X1 [superposition 173,25] {a:2,w:5,tD:1,inD:0}
[SA] new: 183. $sum(X2,0) = X2 [superposition 173,25] {a:2,w:5,tD:1,inD:0}
[SA] new: 184. 0 = $uminus(0) [superposition 29,173] {a:2,w:4,tD:1,inD:0}
[SA] new: 185. $sum(X0,0) = X0 [superposition 25,173] {a:2,w:5,tD:1,inD:0}
[SA] new: 186. $sum(X1,0) = X1 [superposition 25,173] {a:2,w:5,tD:1,inD:0}
[SA] new: 187. 0 = 0 [evaluation 184] {a:2,w:3,tD:1,inD:0}
[SA] new: 188. 0 = 0 [evaluation 181] {a:2,w:3,tD:1,inD:0}
[SA] active: 100. ~'Sub'(X7,nl9) | $less(i(l9(X7)),alength) [cnf transformation 58] {a:0,w:8,nSel:1,inD:0}
[SA] new: 189. $less(i(l9(sK6)),alength) | i(main_end) = 0 [resolution 100,161] {a:1,w:9,inD:0}
[SA] passive: 189. $less(i(l9(sK6)),alength) | i(main_end) = 0 [resolution 100,161] {a:1,w:9,inD:0}
[SA] active: 134. a(l7,X1) = a(l9(zero),X1) [forward demodulation 112,114] {a:0,w:8,nSel:1,inD:0}
[SA] active: 148. v = a(l12(X4),i(l9(X4))) [forward demodulation 107,110] {a:0,w:8,nSel:1,inD:0}
[SA] new: 190. v = a(l12(zero),0) [superposition 148,113] {a:1,w:6,inD:0}
[SA] new: 191. v = a(l12(nl9),i(main_end)) [superposition 148,98] {a:1,w:7,inD:0}
[SA] passive: 190. v = a(l12(zero),0) [superposition 148,113] {a:1,w:6,inD:0}
[SA] passive: 191. v = a(l12(nl9),i(main_end)) [superposition 148,98] {a:1,w:7,inD:0}
[SA] active: 190. v = a(l12(zero),0) [superposition 148,113] {a:1,w:6,nSel:1,inD:0}
[SA] active: 34. $less(X1,$sum(X0,1)) | $less(X0,X1) [theory axiom] {a:0,w:8,nSel:1,tD:1,inD:0}
[SA] new: 192. $less(X0,$sum(X0,1)) [resolution 34,30] {a:1,w:5,tD:1,inD:0}
[SA] new: 193. $less(X1,$sum(1,X0)) | $less(X0,X1) [superposition 34,25] {a:1,w:8,tD:1,inD:0}
[SA] new: 194. $less(X3,$sum(1,X2)) | $less(X2,X3) [superposition 34,25] {a:1,w:8,tD:1,inD:0}
[SA] new: 195. $less(X4,0) | $less($uminus(1),X4) [superposition 34,167] {a:2,w:7,tD:1,inD:0}
[SA] new: 196. $less(X5,1) | $less(0,X5) [superposition 34,173] {a:2,w:6,tD:1,inD:0}
[SA] new: 197. $less(X4,0) | $less(-1,X4) [evaluation 195] {a:2,w:6,tD:1,inD:0}
[SA] passive: 192. $less(X0,$sum(X0,1)) [resolution 34,30] {a:1,w:5,tD:1,inD:0}
[SA] passive: 193. $less(X1,$sum(1,X0)) | $less(X0,X1) [superposition 34,25] {a:1,w:8,tD:1,inD:0}
[SA] passive: 197. $less(X4,0) | $less(-1,X4) [evaluation 195] {a:2,w:6,tD:1,inD:0}
[SA] passive: 196. $less(X5,1) | $less(0,X5) [superposition 34,173] {a:2,w:6,tD:1,inD:0}
[SA] active: 192. $less(X0,$sum(X0,1)) [resolution 34,30] {a:1,w:5,nSel:1,tD:1,inD:0}
[SA] new: 198. $less(X0,$sum(1,X0)) [superposition 192,25] {a:2,w:5,tD:1,inD:0}
[SA] new: 199. $less(X1,$sum(1,X1)) [superposition 192,25] {a:2,w:5,tD:1,inD:0}
[SA] new: 200. $less($uminus(1),0) [superposition 192,167] {a:2,w:4,tD:1,inD:0}
[SA] new: 201. $less(0,1) [superposition 192,173] {a:2,w:3,tD:1,inD:0}
[SA] new: 202. $less(-1,0) [evaluation 200] {a:2,w:3,tD:1,inD:0}
[SA] passive: 198. $less(X0,$sum(1,X0)) [superposition 192,25] {a:2,w:5,tD:1,inD:0}
[SA] passive: 202. $less(-1,0) [evaluation 200] {a:2,w:3,tD:1,inD:0}
[SA] passive: 201. $less(0,1) [superposition 192,173] {a:2,w:3,tD:1,inD:0}
[SA] active: 36. ~$less(X1,$sum(X0,1)) | ~$less(X0,X1) [theory axiom] {a:0,w:8,nSel:1,tD:1,inD:0}
[SA] new: 203. ~$less(X0,X1) | $less(X0,X1) [resolution 36,34] {a:1,w:6,tD:1,inD:0}
[SA] new: 204. ~$less(X2,X2) [resolution 36,192] {a:2,w:3,tD:1,inD:0}
[SA] new: 205. ~$less(X1,$sum(1,X0)) | ~$less(X0,X1) [superposition 36,25] {a:1,w:8,tD:1,inD:0}
[SA] new: 206. ~$less(X3,$sum(1,X2)) | ~$less(X2,X3) [superposition 36,25] {a:1,w:8,tD:1,inD:0}
[SA] new: 207. ~$less(X4,0) | ~$less($uminus(1),X4) [superposition 36,167] {a:2,w:7,tD:1,inD:0}
[SA] new: 208. ~$less(X5,1) | ~$less(0,X5) [superposition 36,173] {a:2,w:6,tD:1,inD:0}
[SA] new: 209. ~$less(X4,0) | ~$less(-1,X4) [evaluation 207] {a:2,w:6,tD:1,inD:0}
[SA] passive: 205. ~$less(X1,$sum(1,X0)) | ~$less(X0,X1) [superposition 36,25] {a:1,w:8,tD:1,inD:0}
[SA] passive: 209. ~$less(X4,0) | ~$less(-1,X4) [evaluation 207] {a:2,w:6,tD:1,inD:0}
[SA] passive: 208. ~$less(X5,1) | ~$less(0,X5) [superposition 36,173] {a:2,w:6,tD:1,inD:0}
[SA] active: 201. $less(0,1) [superposition 192,173] {a:2,w:3,nSel:1,tD:1,inD:0}
[SA] active: 39. s(X0) != s(X1) | X0 = X1 [term algebras injectivity] {a:0,w:8,nSel:1,tD:1,inD:0}
[SA] new: 210. X0 = X0 [equality resolution 39] {a:1,w:3,tD:1,inD:0}
[SA] active: 202. $less(-1,0) [evaluation 200] {a:2,w:3,nSel:1,tD:1,inD:0}
[SA] active: 74. s(p(X0)) = X0 | zero = X0 [cnf transformation 37] {a:0,w:8,nSel:1,inD:0}
[SA] new: 211. s(X1) != X0 | p(X0) = X1 | zero = X0 [superposition 39,74] {a:1,w:11,inD:0}
[SA] new: 212. s(X3) != X2 | p(X2) = X3 | zero = X2 [superposition 39,74] {a:1,w:11,inD:0}
[SA] new: 213. 'Sub'(X5,X4) | ~'Sub'(X5,p(X4)) | zero = X4 [superposition 116,74] {a:1,w:10,inD:0}
[SA] new: 214. zero != X6 | zero = X6 [superposition 38,74] {a:1,w:6,inD:0}
[SA] new: 215. 'Sub'(p(X7),X7) | zero = X7 [superposition 117,74] {a:1,w:7,inD:0}
[SA] passive: 211. s(X1) != X0 | p(X0) = X1 | zero = X0 [superposition 39,74] {a:1,w:11,inD:0}
[SA] passive: 213. 'Sub'(X5,X4) | ~'Sub'(X5,p(X4)) | zero = X4 [superposition 116,74] {a:1,w:10,inD:0}
[SA] passive: 215. 'Sub'(p(X7),X7) | zero = X7 [superposition 117,74] {a:1,w:7,inD:0}
[SA] active: 198. $less(X0,$sum(1,X0)) [superposition 192,25] {a:2,w:5,nSel:1,tD:1,inD:0}
[SA] new: 216. ~$less(1,1) [resolution 198,36] {a:3,w:3,tD:1,inD:0}
[SA] new: 217. $less(0,1) [superposition 198,27] {a:3,w:3,tD:1,inD:0}
[SA] new: 218. $less($uminus(1),0) [superposition 198,29] {a:3,w:4,tD:1,inD:0}
[SA] new: 219. $less(X0,$sum(X0,1)) [superposition 198,25] {a:3,w:5,tD:1,inD:0}
[SA] new: 220. $less(X1,$sum(X1,1)) [superposition 198,25] {a:3,w:5,tD:1,inD:0}
[SA] new: 221. $less(-1,0) [evaluation 218] {a:3,w:3,tD:1,inD:0}
[SA] active: 108. a(l9(X2),X3) = a(l11(X2),X3) [cnf transformation 58] {a:0,w:9,nSel:1,inD:0}
[SA] active: 196. $less(X5,1) | $less(0,X5) [superposition 34,173] {a:2,w:6,nSel:2,tD:1,inD:0}
[SA] new: 222. $less(0,1) [resolution 196,30] {a:3,w:3,tD:1,inD:0}
[SA] new: 223. $less(alength,1) | s(sK3) = nl9 [resolution 196,131] {a:3,w:7,inD:0}
[SA] new: 224. $less(0,1) [resolution 196,30] {a:3,w:3,tD:1,inD:0}
[SA] new: 225. $less($sum(X0,1),1) | ~$less(X0,0) [resolution 196,36] {a:3,w:8,tD:1,inD:0}
[SA] new: 226. $less(X0,$sum(1,-1)) | ~$less(X0,0) [evaluation 225] {a:3,w:8,tD:1,inD:0}
[SA] new: 227. $less(X0,0) | ~$less(X0,0) [evaluation 226] {a:3,w:6,tD:1,inD:0}
[SA] passive: 223. $less(alength,1) | s(sK3) = nl9 [resolution 196,131] {a:3,w:7,inD:0}
[SA] active: 31. ~$less(X1,X2) | ~$less(X0,X1) | $less(X0,X2) [theory axiom] {a:0,w:9,nSel:1,tD:1,inD:0}
[SA] new: 228. ~$less(X0,X1) | $less(X0,$sum(X2,1)) | $less(X2,X1) [resolution 31,34] {a:1,w:11,tD:1,inD:0}
[SA] new: 229. ~$less(X3,X4) | $less(X3,$sum(X4,1)) [resolution 31,192] {a:2,w:8,tD:1,inD:0}
[SA] new: 230. ~$less(X5,X6) | $less(X5,$sum(1,X6)) [resolution 31,198] {a:3,w:8,tD:1,inD:0}
[SA] new: 231. ~$less(X7,X8) | $less(X7,1) | $less(0,X8) [resolution 31,196] {a:3,w:9,tD:1,inD:0}
[SA] new: 232. ~$less(X9,0) | $less(X9,1) [resolution 31,201] {a:3,w:6,tD:1,inD:0}
[SA] new: 233. ~$less(X10,0) | $less(X10,X11) | $less(X11,1) [resolution 31,196] {a:3,w:9,tD:1,inD:0}
[SA] new: 234. ~$less(X12,2) | $less(X12,1) [resolution 31,75] {a:1,w:6,goal:1,inD:0}
[SA] new: 235. ~$less(X13,-1) | $less(X13,0) [resolution 31,202] {a:3,w:6,tD:1,inD:0}
[SA] passive: 228. ~$less(X0,X1) | $less(X0,$sum(X2,1)) | $less(X2,X1) [resolution 31,34] {a:1,w:11,tD:1,inD:0}
[SA] passive: 229. ~$less(X3,X4) | $less(X3,$sum(X4,1)) [resolution 31,192] {a:2,w:8,tD:1,inD:0}
[SA] passive: 230. ~$less(X5,X6) | $less(X5,$sum(1,X6)) [resolution 31,198] {a:3,w:8,tD:1,inD:0}
[SA] passive: 231. ~$less(X7,X8) | $less(X7,1) | $less(0,X8) [resolution 31,196] {a:3,w:9,tD:1,inD:0}
[SA] passive: 232. ~$less(X9,0) | $less(X9,1) [resolution 31,201] {a:3,w:6,tD:1,inD:0}
[SA] passive: 233. ~$less(X10,0) | $less(X10,X11) | $less(X11,1) [resolution 31,196] {a:3,w:9,tD:1,inD:0}
[SA] passive: 234. ~$less(X12,2) | $less(X12,1) [resolution 31,75] {a:1,w:6,goal:1,inD:0}
[SA] passive: 235. ~$less(X13,-1) | $less(X13,0) [resolution 31,202] {a:3,w:6,tD:1,inD:0}
[SA] active: 234. ~$less(X12,2) | $less(X12,1) [resolution 31,75] {a:1,w:6,nSel:1,goal:1,inD:0}
[SA] new: 236. $less(0,1) | $less(2,1) [resolution 234,196] {a:3,w:6,goal:1,inD:0}
[SA] active: 32. $less(X1,X0) | $less(X0,X1) | X0 = X1 [theory axiom] {a:0,w:9,nSel:2,tD:1,inD:0}
[SA] new: 237. $less(X0,X0) | X0 = X0 [resolution 32,30] {a:1,w:6,tD:1,inD:0}
[SA] new: 238. $less(X1,X2) | X1 = X2 | ~$less(X3,X2) | $less(X3,X1) [resolution 32,31] {a:1,w:12,tD:1,inD:0}
[SA] new: 239. $less($sum(X4,1),X5) | $sum(X4,1) = X5 | ~$less(X4,X5) [resolution 32,36] {a:1,w:13,tD:1,inD:0}
[SA] new: 240. $less(2,X6) | 2 = X6 | $less(X6,1) [resolution 32,234] {a:2,w:9,goal:1,inD:0}
[SA] new: 241. $less(alength,i(main_end)) | alength = i(main_end) [resolution 32,157] {a:1,w:8,inD:0}
[SA] new: 242. $less(0,1) | 1 = 0 [resolution 32,77] {a:1,w:6,inD:0}
[SA] new: 243. $less(alength,0) | alength = 0 | s(sK3) = nl9 [resolution 32,131] {a:1,w:10,inD:0}
[SA] new: 244. $less(X0,X0) | X0 = X0 [resolution 32,30] {a:1,w:6,tD:1,inD:0}
[SA] new: 245. $less(X1,X2) | X1 = X2 | ~$less(X3,X2) | $less(X3,X1) [resolution 32,31] {a:1,w:12,tD:1,inD:0}
[SA] new: 246. $less($sum(X4,1),X5) | $sum(X4,1) = X5 | ~$less(X4,X5) [resolution 32,36] {a:1,w:13,tD:1,inD:0}
[SA] new: 247. $less(2,X6) | 2 = X6 | $less(X6,1) [resolution 32,234] {a:2,w:9,goal:1,inD:0}
[SA] new: 248. $less(alength,i(main_end)) | alength = i(main_end) [resolution 32,157] {a:1,w:8,inD:0}
[SA] new: 249. $less(0,1) | 1 = 0 [resolution 32,77] {a:1,w:6,inD:0}
[SA] new: 250. $less(alength,0) | alength = 0 | s(sK3) = nl9 [resolution 32,131] {a:1,w:10,inD:0}
[SA] new: 251. $less(X0,X0) | X0 = X0 [factoring 32] {a:1,w:6,tD:1,inD:0}
[SA] passive: 238. $less(X1,X2) | X1 = X2 | ~$less(X3,X2) | $less(X3,X1) [resolution 32,31] {a:1,w:12,tD:1,inD:0}
[SA] passive: 239. $less($sum(X4,1),X5) | $sum(X4,1) = X5 | ~$less(X4,X5) [resolution 32,36] {a:1,w:13,tD:1,inD:0}
[SA] passive: 240. $less(2,X6) | 2 = X6 | $less(X6,1) [resolution 32,234] {a:2,w:9,goal:1,inD:0}
[SA] passive: 241. $less(alength,i(main_end)) | alength = i(main_end) [resolution 32,157] {a:1,w:8,inD:0}
[SA] passive: 243. $less(alength,0) | alength = 0 | s(sK3) = nl9 [resolution 32,131] {a:1,w:10,inD:0}
[SA] active: 197. $less(X4,0) | $less(-1,X4) [evaluation 195] {a:2,w:6,nSel:2,tD:1,inD:0}
[SA] new: 252. $less(-1,0) [resolution 197,30] {a:3,w:3,tD:1,inD:0}
[SA] new: 253. $less(-1,X0) | ~$less(X1,X0) | $less(X1,0) [resolution 197,31] {a:3,w:9,tD:1,inD:0}
[SA] new: 254. $less(-1,1) [resolution 197,77] {a:3,w:3,inD:0}
[SA] new: 255. $less(-1,0) [resolution 197,30] {a:3,w:3,tD:1,inD:0}
[SA] new: 256. $less(X0,0) | ~$less(X1,-1) | $less(X1,X0) [resolution 197,31] {a:3,w:9,tD:1,inD:0}
[SA] new: 257. $less($sum(X2,1),0) | ~$less(X2,-1) [resolution 197,36] {a:3,w:8,tD:1,inD:0}
[SA] new: 258. $less(2,0) | $less(-1,1) [resolution 197,234] {a:3,w:6,goal:1,inD:0}
[SA] new: 259. $less(X2,$sum(0,-1)) | ~$less(X2,-1) [evaluation 257] {a:3,w:8,tD:1,inD:0}
[SA] new: 260. $less(X2,-1) | ~$less(X2,-1) [evaluation 259] {a:3,w:6,tD:1,inD:0}
[SA] passive: 253. $less(-1,X0) | ~$less(X1,X0) | $less(X1,0) [resolution 197,31] {a:3,w:9,tD:1,inD:0}
[SA] passive: 254. $less(-1,1) [resolution 197,77] {a:3,w:3,inD:0}
[SA] passive: 256. $less(X0,0) | ~$less(X1,-1) | $less(X1,X0) [resolution 197,31] {a:3,w:9,tD:1,inD:0}
[SA] active: 102. a(l12(X4),X6) = a(l9(s(X4)),X6) [cnf transformation 58] {a:0,w:10,nSel:1,inD:0}
[SA] new: 261. a(l9(X0),X1) = a(l12(p(X0)),X1) | zero = X0 [superposition 102,74] {a:1,w:13,inD:0}
[SA] passive: 261. a(l9(X0),X1) = a(l12(p(X0)),X1) | zero = X0 [superposition 102,74] {a:1,w:13,inD:0}
[SA] active: 254. $less(-1,1) [resolution 197,77] {a:3,w:3,nSel:1,inD:0}
[SA] new: 262. ~$less(X0,-1) | $less(X0,1) [resolution 254,31] {a:4,w:6,inD:0}
[SA] passive: 262. ~$less(X0,-1) | $less(X0,1) [resolution 254,31] {a:4,w:6,inD:0}
[SA] active: 146. 1 != 0 | X1 = X2 | ~'Sub'(s(X0),nl9) [forward demodulation 145,141] {a:0,w:10,nSel:1,inD:0}
[SA] active: 208. ~$less(0,X5) | ~$less(X5,1) [superposition 36,173] {a:2,w:6,nSel:1,tD:1,inD:0}
[SA] new: 263. ~$less(1,1) [resolution 208,201] {a:3,w:3,tD:1,inD:0}
[SA] new: 264. ~$less(X0,1) | $less(X0,1) [resolution 208,196] {a:3,w:6,tD:1,inD:0}
[SA] new: 265. ~$less(X1,1) | $less(X1,0) | 0 = X1 [resolution 208,32] {a:3,w:9,tD:1,inD:0}
[SA] new: 266. ~$less(X2,1) | $less(X2,0) | 0 = X2 [resolution 208,32] {a:3,w:9,tD:1,inD:0}
[SA] new: 267. ~$less(1,1) | $less(0,0) [resolution 208,196] {a:3,w:6,tD:1,inD:0}
[SA] new: 268. ~$less($sum(X3,1),1) | $less(X3,0) [resolution 208,34] {a:3,w:8,tD:1,inD:0}
[SA] new: 269. ~$less($sum(0,1),1) [resolution 208,192] {a:3,w:5,tD:1,inD:0}
[SA] new: 270. ~$less($sum(1,0),1) [resolution 208,198] {a:3,w:5,tD:1,inD:0}
[SA] new: 271. ~$less(0,1) | $less(-1,0) [resolution 208,197] {a:3,w:6,tD:1,inD:0}
[SA] new: 272. ~$less(1,1) [evaluation 270] {a:3,w:3,tD:1,inD:0}
[SA] new: 273. ~$less(1,1) [evaluation 269] {a:3,w:3,tD:1,inD:0}
[SA] new: 274. ~$less(X3,$sum(1,-1)) | $less(X3,0) [evaluation 268] {a:3,w:8,tD:1,inD:0}
[SA] new: 275. ~$less(X3,0) | $less(X3,0) [evaluation 274] {a:3,w:6,tD:1,inD:0}
[SA] passive: 265. ~$less(X1,1) | $less(X1,0) | 0 = X1 [resolution 208,32] {a:3,w:9,tD:1,inD:0}
[SA] active: 153. i(l9(s(X4))) = $sum(i(l9(X4)),1) [forward demodulation 103,152] {a:0,w:10,nSel:1,inD:0}
[SA] new: 276. i(l9(s(zero))) = $sum(0,1) [superposition 153,113] {a:1,w:8,inD:0}
[SA] new: 277. i(l9(s(nl9))) = $sum(i(main_end),1) [superposition 153,98] {a:1,w:9,inD:0}
[SA] new: 278. i(l9(s(X1))) = $sum(1,i(l9(X1))) [superposition 153,25] {a:1,w:10,inD:0}
[SA] new: 279. i(l9(s(X2))) = $sum(1,i(l9(X2))) [superposition 153,25] {a:1,w:10,inD:0}
[SA] new: 280. ~$less(X1,i(l9(s(X0)))) | ~$less(i(l9(X0)),X1) [superposition 36,153] {a:1,w:11,inD:0}
[SA] new: 281. $less(i(l9(X2)),i(l9(s(X2)))) [superposition 192,153] {a:2,w:8,inD:0}
[SA] new: 282. $less(X4,i(l9(s(X3)))) | $less(i(l9(X3)),X4) [superposition 34,153] {a:1,w:11,inD:0}
[SA] new: 283. i(l9(s(X5))) = $sum(1,i(l9(X5))) [superposition 25,153] {a:1,w:10,inD:0}
[SA] new: 284. i(l9(s(X6))) = $sum(1,i(l9(X6))) [superposition 25,153] {a:1,w:10,inD:0}
[SA] new: 285. i(l9(s(zero))) = 1 [evaluation 276] {a:1,w:6,inD:0}
[SA] passive: 285. i(l9(s(zero))) = 1 [evaluation 276] {a:1,w:6,inD:0}
[SA] new: 286. i(l9(s(nl9))) = $sum(1,i(main_end)) [forward demodulation 277,25] {a:1,w:9,inD:0}
[SA] passive: 286. i(l9(s(nl9))) = $sum(1,i(main_end)) [forward demodulation 277,25] {a:1,w:9,inD:0}
[SA] new: 287. i(l9(s(X1))) != i(l9(s(X1))) | ~$less(X0,i(main_end)) | $less(X0,0) | 'Sub'(sK2(X0),nl9) [backward demodulation 278,159] {a:0,w:20,inD:0}
[SA] new: 288. i(l9(s(X1))) != i(l9(s(X1))) | ~$less(X0,i(main_end)) | $less(X0,0) | i(l9(sK2(X0))) = X0 [backward demodulation 278,158] {a:0,w:22,inD:0}
[SA] new: 289. i(l9(s(X0))) != i(l9(s(X0))) | X1 = X2 | i(l9(X1)) != i(l9(X2)) | ~'Sub'(s(X0),nl9) [backward demodulation 278,121] {a:0,w:23,inD:0}
[SA] passive: 278. i(l9(s(X1))) = $sum(1,i(l9(X1))) [superposition 153,25] {a:1,w:10,inD:0}
[SA] new: 290. X1 = X2 | i(l9(X1)) != i(l9(X2)) | ~'Sub'(s(X0),nl9) [trivial inequality removal 289] {a:0,w:14,inD:0}
[SA] new: 291. ~$less(X0,i(main_end)) | $less(X0,0) | i(l9(sK2(X0))) = X0 [trivial inequality removal 288] {a:0,w:13,inD:0}
[SA] new: 292. ~$less(X0,i(main_end)) | $less(X0,0) | 'Sub'(sK2(X0),nl9) [trivial inequality removal 287] {a:0,w:11,inD:0}
[SA] passive: 292. ~$less(X0,i(main_end)) | $less(X0,0) | 'Sub'(sK2(X0),nl9) [trivial inequality removal 287] {a:0,w:11,inD:0}
[SA] passive: 291. ~$less(X0,i(main_end)) | $less(X0,0) | i(l9(sK2(X0))) = X0 [trivial inequality removal 288] {a:0,w:13,inD:0}
[SA] passive: 290. X1 = X2 | i(l9(X1)) != i(l9(X2)) | ~'Sub'(s(X0),nl9) [trivial inequality removal 289] {a:0,w:14,inD:0}
[SA] passive: 280. ~$less(X1,i(l9(s(X0)))) | ~$less(i(l9(X0)),X1) [superposition 36,153] {a:1,w:11,inD:0}
[SA] passive: 281. $less(i(l9(X2)),i(l9(s(X2)))) [superposition 192,153] {a:2,w:8,inD:0}
[SA] passive: 282. $less(X4,i(l9(s(X3)))) | $less(i(l9(X3)),X4) [superposition 34,153] {a:1,w:11,inD:0}
[SA] active: 285. i(l9(s(zero))) = 1 [evaluation 276] {a:1,w:6,nSel:1,inD:0}
[SA] new: 293. i(l9(s(s(zero)))) = $sum(1,1) [superposition 153,285] {a:2,w:9,inD:0}
[SA] new: 294. v = a(l12(s(zero)),1) [superposition 148,285] {a:2,w:7,inD:0}
[SA] new: 295. i(l9(s(s(zero)))) = 2 [evaluation 293] {a:2,w:7,inD:0}
[SA] passive: 295. i(l9(s(s(zero)))) = 2 [evaluation 293] {a:2,w:7,inD:0}
[SA] passive: 294. v = a(l12(s(zero)),1) [superposition 148,285] {a:2,w:7,inD:0}
[SA] active: 28. $sum($uminus(X1),$uminus(X0)) = $uminus($sum(X0,X1)) [theory axiom] {a:0,w:10,nSel:1,tD:1,inD:0}
[SA] new: 296. $sum(X0,$uminus(X1)) = $uminus($sum(X1,$uminus(X0))) [superposition 28,35] {a:1,w:10,tD:1,inD:0}
[SA] new: 297. $sum($uminus(X1),X0) = $uminus($sum($uminus(X0),X1)) [superposition 28,35] {a:1,w:10,tD:1,inD:0}
[SA] new: 298. 0 = $uminus($sum(X2,$uminus(X2))) [superposition 28,167] {a:2,w:7,tD:1,inD:0}
[SA] new: 299. 0 = $uminus($sum($uminus(X3),X3)) [superposition 28,29] {a:1,w:7,tD:1,inD:0}
[SA] new: 300. $sum($uminus(X5),$uminus(X4)) = $uminus($sum(X5,X4)) [superposition 28,25] {a:1,w:10,tD:1,inD:0}
[SA] new: 301. $sum($uminus(X7),$uminus(X6)) = $uminus($sum(X7,X6)) [superposition 28,25] {a:1,w:10,tD:1,inD:0}
[SA] new: 302. 0 = $uminus($sum(X0,$uminus(X0))) [superposition 167,28] {a:2,w:7,tD:1,inD:0}
[SA] new: 303. 0 = $uminus($sum($uminus(X1),X1)) [superposition 29,28] {a:1,w:7,tD:1,inD:0}
[SA] new: 304. $sum($uminus(X3),$uminus(X2)) = $uminus($sum(X3,X2)) [superposition 25,28] {a:1,w:10,tD:1,inD:0}
[SA] new: 305. $sum($uminus(X5),$uminus(X4)) = $uminus($sum(X5,X4)) [superposition 25,28] {a:1,w:10,tD:1,inD:0}
[SA] passive: 296. $sum(X0,$uminus(X1)) = $uminus($sum(X1,$uminus(X0))) [superposition 28,35] {a:1,w:10,tD:1,inD:0}
[SA] passive: 297. $sum($uminus(X1),X0) = $uminus($sum($uminus(X0),X1)) [superposition 28,35] {a:1,w:10,tD:1,inD:0}
[SA] new: 306. 0 = $uminus(0) [forward demodulation 298,29] {a:2,w:4,tD:1,inD:0}
[SA] new: 307. 0 = 0 [evaluation 306] {a:2,w:3,tD:1,inD:0}
[SA] new: 308. 0 = $uminus($sum(X3,$uminus(X3))) [forward demodulation 299,25] {a:1,w:7,tD:1,inD:0}
[SA] new: 309. 0 = $uminus(0) [forward demodulation 308,29] {a:1,w:4,tD:1,inD:0}
[SA] new: 310. 0 = 0 [evaluation 309] {a:1,w:3,tD:1,inD:0}
[SA] passive: 300. $sum($uminus(X5),$uminus(X4)) = $uminus($sum(X5,X4)) [superposition 28,25] {a:1,w:10,tD:1,inD:0}
[SA] new: 311. 0 = $uminus(0) [forward demodulation 302,29] {a:2,w:4,tD:1,inD:0}
[SA] new: 312. 0 = 0 [evaluation 311] {a:2,w:3,tD:1,inD:0}
[SA] new: 313. 0 = $uminus($sum(X1,$uminus(X1))) [forward demodulation 303,25] {a:1,w:7,tD:1,inD:0}
[SA] new: 314. 0 = $uminus(0) [forward demodulation 313,29] {a:1,w:4,tD:1,inD:0}
[SA] new: 315. 0 = 0 [evaluation 314] {a:1,w:3,tD:1,inD:0}
[SA] active: 209. ~$less(-1,X4) | ~$less(X4,0) [evaluation 207] {a:2,w:6,nSel:1,tD:1,inD:0}
[SA] new: 316. ~$less(0,0) [resolution 209,202] {a:3,w:3,tD:1,inD:0}
[SA] new: 317. ~$less(X0,0) | $less(X0,0) [resolution 209,197] {a:3,w:6,tD:1,inD:0}
[SA] new: 318. ~$less(1,0) [resolution 209,254] {a:4,w:3,inD:0}
[SA] new: 319. ~$less(X1,0) | $less(X1,-1) | -1 = X1 [resolution 209,32] {a:3,w:9,tD:1,inD:0}
[SA] new: 320. ~$less(X2,0) | $less(X2,-1) | -1 = X2 [resolution 209,32] {a:3,w:9,tD:1,inD:0}
[SA] new: 321. ~$less(1,0) | $less(0,-1) [resolution 209,196] {a:3,w:6,tD:1,inD:0}
[SA] new: 322. ~$less($sum(X3,1),0) | $less(X3,-1) [resolution 209,34] {a:3,w:8,tD:1,inD:0}
[SA] new: 323. ~$less($sum(-1,1),0) [resolution 209,192] {a:3,w:5,tD:1,inD:0}
[SA] new: 324. ~$less($sum(1,-1),0) [resolution 209,198] {a:3,w:5,tD:1,inD:0}
[SA] new: 325. ~$less(0,0) | $less(-1,-1) [resolution 209,197] {a:3,w:6,tD:1,inD:0}
[SA] new: 326. ~$less(0,0) [evaluation 324] {a:3,w:3,tD:1,inD:0}
[SA] new: 327. ~$less(0,0) [evaluation 323] {a:3,w:3,tD:1,inD:0}
[SA] new: 328. ~$less(X3,$sum(0,-1)) | $less(X3,-1) [evaluation 322] {a:3,w:8,tD:1,inD:0}
[SA] new: 329. ~$less(X3,-1) | $less(X3,-1) [evaluation 328] {a:3,w:6,tD:1,inD:0}
[SA] passive: 319. ~$less(X1,0) | $less(X1,-1) | -1 = X1 [resolution 209,32] {a:3,w:9,tD:1,inD:0}
[SA] active: 33. $less($sum(X0,X2),$sum(X1,X2)) | ~$less(X0,X1) [theory axiom] {a:0,w:10,nSel:1,tD:1,inD:0}
[SA] new: 330. ~$less(X0,X1) | ~$less(X1,$sum(X0,1)) [resolution 33,36] {a:1,w:8,tD:1,inD:0}
[SA] new: 331. ~$less(X2,X2) [resolution 33,30] {a:1,w:3,tD:1,inD:0}
[SA] new: 332. ~$less(X3,X4) | ~$less(X5,$sum(X3,X6)) | $less(X5,$sum(X4,X6)) [resolution 33,31] {a:1,w:13,tD:1,inD:0}
[SA] new: 333. $less(X0,$sum(X1,0)) | ~$less(X0,X1) [superposition 33,27] {a:1,w:8,tD:1,inD:0}
[SA] new: 334. $less(0,$sum(X3,$uminus(X2))) | ~$less(X2,X3) [superposition 33,29] {a:1,w:9,tD:1,inD:0}
[SA] new: 335. $less($sum(X5,X4),$sum(X6,X5)) | ~$less(X4,X6) [superposition 33,25] {a:1,w:10,tD:1,inD:0}
[SA] new: 336. $less($sum(X8,X7),$sum(X9,X8)) | ~$less(X7,X9) [superposition 33,25] {a:1,w:10,tD:1,inD:0}
[SA] new: 337. $less(i(l9(s(X10))),$sum(X11,1)) | ~$less(i(l9(X10)),X11) [superposition 33,153] {a:1,w:13,inD:0}
[SA] new: 338. $less(X12,$sum(X13,X12)) | ~$less(0,X13) [superposition 33,173] {a:2,w:8,tD:1,inD:0}
[SA] new: 339. $less(0,$sum(X15,X14)) | ~$less($uminus(X14),X15) [superposition 33,167] {a:2,w:9,tD:1,inD:0}
[SA] new: 340. $less($uminus($sum(X17,X16)),$sum(X18,$uminus(X17))) | ~$less($uminus(X16),X18) [superposition 33,28] {a:1,w:13,tD:1,inD:0}
[SA] new: 341. $less($sum(X1,0),X0) | ~$less(X1,X0) [superposition 33,27] {a:1,w:8,tD:1,inD:0}
[SA] new: 342. $less($sum(X3,$uminus(X2)),0) | ~$less(X3,X2) [superposition 33,29] {a:1,w:9,tD:1,inD:0}
[SA] new: 343. $less($sum(X6,X5),$sum(X5,X4)) | ~$less(X6,X4) [superposition 33,25] {a:1,w:10,tD:1,inD:0}
[SA] new: 344. $less($sum(X9,X8),$sum(X8,X7)) | ~$less(X9,X7) [superposition 33,25] {a:1,w:10,tD:1,inD:0}
[SA] new: 345. $less($sum(X11,1),i(l9(s(X10)))) | ~$less(X11,i(l9(X10))) [superposition 33,153] {a:1,w:13,inD:0}
[SA] new: 346. $less($sum(X13,X12),X12) | ~$less(X13,0) [superposition 33,173] {a:2,w:8,tD:1,inD:0}
[SA] new: 347. $less($sum(X15,X14),0) | ~$less(X15,$uminus(X14)) [superposition 33,167] {a:2,w:9,tD:1,inD:0}
[SA] new: 348. $less($sum(X18,$uminus(X17)),$uminus($sum(X17,X16))) | ~$less(X18,$uminus(X16)) [superposition 33,28] {a:1,w:13,tD:1,inD:0}
[SA] passive: 332. ~$less(X3,X4) | ~$less(X5,$sum(X3,X6)) | $less(X5,$sum(X4,X6)) [resolution 33,31] {a:1,w:13,tD:1,inD:0}
[SA] new: 349. $less(X0,X1) | ~$less(X0,X1) [forward demodulation 333,27] {a:1,w:6,tD:1,inD:0}
[SA] passive: 334. $less(0,$sum(X3,$uminus(X2))) | ~$less(X2,X3) [superposition 33,29] {a:1,w:9,tD:1,inD:0}
[SA] passive: 335. $less($sum(X5,X4),$sum(X6,X5)) | ~$less(X4,X6) [superposition 33,25] {a:1,w:10,tD:1,inD:0}
[SA] passive: 337. $less(i(l9(s(X10))),$sum(X11,1)) | ~$less(i(l9(X10)),X11) [superposition 33,153] {a:1,w:13,inD:0}
[SA] passive: 338. $less(X12,$sum(X13,X12)) | ~$less(0,X13) [superposition 33,173] {a:2,w:8,tD:1,inD:0}
[SA] passive: 339. $less(0,$sum(X15,X14)) | ~$less($uminus(X14),X15) [superposition 33,167] {a:2,w:9,tD:1,inD:0}
[SA] passive: 340. $less($uminus($sum(X17,X16)),$sum(X18,$uminus(X17))) | ~$less($uminus(X16),X18) [superposition 33,28] {a:1,w:13,tD:1,inD:0}
[SA] new: 350. $less(X1,X0) | ~$less(X1,X0) [forward demodulation 341,27] {a:1,w:6,tD:1,inD:0}
[SA] passive: 342. $less($sum(X3,$uminus(X2)),0) | ~$less(X3,X2) [superposition 33,29] {a:1,w:9,tD:1,inD:0}
[SA] passive: 343. $less($sum(X6,X5),$sum(X5,X4)) | ~$less(X6,X4) [superposition 33,25] {a:1,w:10,tD:1,inD:0}
[SA] passive: 345. $less($sum(X11,1),i(l9(s(X10)))) | ~$less(X11,i(l9(X10))) [superposition 33,153] {a:1,w:13,inD:0}
[SA] passive: 346. $less($sum(X13,X12),X12) | ~$less(X13,0) [superposition 33,173] {a:2,w:8,tD:1,inD:0}
[SA] passive: 347. $less($sum(X15,X14),0) | ~$less(X15,$uminus(X14)) [superposition 33,167] {a:2,w:9,tD:1,inD:0}
[SA] passive: 348. $less($sum(X18,$uminus(X17)),$uminus($sum(X17,X16))) | ~$less(X18,$uminus(X16)) [superposition 33,28] {a:1,w:13,tD:1,inD:0}
[SA] active: 232. ~$less(X9,0) | $less(X9,1) [resolution 31,201] {a:3,w:6,nSel:1,tD:1,inD:0}
[SA] new: 351. $less(X0,1) | $less(-1,X0) [resolution 232,197] {a:4,w:6,tD:1,inD:0}
[SA] new: 352. $less(X1,1) | $less(0,X1) | 0 = X1 [resolution 232,32] {a:4,w:9,tD:1,inD:0}
[SA] new: 353. $less(X2,1) | $less(0,X2) | 0 = X2 [resolution 232,32] {a:4,w:9,tD:1,inD:0}
[SA] new: 354. $less(0,1) | $less(0,1) [resolution 232,196] {a:4,w:6,tD:1,inD:0}
[SA] new: 355. $less(-1,1) [resolution 232,202] {a:4,w:3,tD:1,inD:0}
[SA] new: 356. $less(-1,1) | $less(0,0) [resolution 232,197] {a:4,w:6,tD:1,inD:0}
[SA] new: 357. $less(0,1) [duplicate literal removal 354] {a:4,w:3,tD:1,inD:0}
[SA] passive: 351. $less(X0,1) | $less(-1,X0) [resolution 232,197] {a:4,w:6,tD:1,inD:0}
[SA] active: 118. $sum(X0,$sum(X2,X1)) = $sum($sum(X0,X1),X2) [forward demodulation 79,26] {a:0,w:11,nSel:1,inD:0}
[SA] new: 358. $sum(X0,X1) = $sum(X0,$sum(X1,0)) [superposition 118,27] {a:1,w:9,inD:0}
[SA] new: 359. $sum(0,X3) = $sum(X2,$sum(X3,$uminus(X2))) [superposition 118,29] {a:1,w:10,inD:0}
[SA] new: 360. $sum(X4,$sum(X6,X5)) = $sum($sum(X5,X4),X6) [superposition 118,25] {a:1,w:11,inD:0}
[SA] new: 361. $sum(X7,$sum(X9,X8)) = $sum($sum(X8,X7),X9) [superposition 118,25] {a:1,w:11,inD:0}
[SA] new: 362. $sum(i(l9(s(X10))),X11) = $sum(i(l9(X10)),$sum(X11,1)) [superposition 118,153] {a:1,w:14,inD:0}
[SA] new: 363. $sum($sum(X12,X13),$sum(X15,X14)) = $sum($sum(X12,$sum(X14,X13)),X15) [superposition 118,118] {a:1,w:15,inD:0}
[SA] new: 364. $sum(X16,X17) = $sum(0,$sum(X17,X16)) [superposition 118,173] {a:2,w:9,inD:0}
[SA] new: 365. $sum(0,X19) = $sum($uminus(X18),$sum(X19,X18)) [superposition 118,167] {a:2,w:10,inD:0}
[SA] new: 366. $sum($uminus($sum(X21,X20)),X22) = $sum($uminus(X20),$sum(X22,$uminus(X21))) [superposition 118,28] {a:1,w:14,inD:0}
[SA] new: 367. $sum(X3,X4) = $sum(X3,$sum(0,X4)) [superposition 118,27] {a:1,w:9,inD:0}
[SA] new: 368. $sum(X5,$sum($uminus($sum(X5,X6)),X6)) = 0 [superposition 118,29] {a:1,w:10,inD:0}
[SA] new: 369. $sum(X7,$sum(X9,X8)) = $sum(X9,$sum(X7,X8)) [superposition 118,25] {a:1,w:11,inD:0}
[SA] new: 370. $sum(X10,$sum(X12,X11)) = $sum(X12,$sum(X10,X11)) [superposition 118,25] {a:1,w:11,inD:0}
[SA] new: 371. $less($sum(X3,X2),$sum(X0,$sum(X2,X1))) | ~$less(X3,$sum(X0,X1)) [superposition 33,118] {a:1,w:14,inD:0}
[SA] new: 372. $less($sum(X4,$sum(X6,X5)),$sum(X7,X6)) | ~$less($sum(X4,X5),X7) [superposition 33,118] {a:1,w:14,inD:0}
[SA] new: 373. $sum(X8,$sum(X10,X9)) = $sum(X10,$sum(X8,X9)) [superposition 25,118] {a:1,w:11,inD:0}
[SA] new: 374. $sum(X11,$sum(X13,X12)) = $sum(X13,$sum(X11,X12)) [superposition 25,118] {a:1,w:11,inD:0}
[SA] new: 375. ~$less(X16,$sum(X14,$sum(1,X15))) | ~$less($sum(X14,X15),X16) [superposition 36,118] {a:1,w:12,inD:0}
[SA] new: 376. $less($sum(X17,X18),$sum(X17,$sum(1,X18))) [superposition 192,118] {a:2,w:9,inD:0}
[SA] new: 377. $less(X21,$sum(X19,$sum(1,X20))) | $less($sum(X19,X20),X21) [superposition 34,118] {a:1,w:12,inD:0}
[SA] new: 378. $sum(X22,X23) = $sum(X22,$sum(0,X23)) [superposition 27,118] {a:1,w:9,inD:0}
[SA] new: 379. $sum(X24,$sum($uminus($sum(X24,X25)),X25)) = 0 [superposition 29,118] {a:1,w:10,inD:0}
[SA] new: 380. $sum(X2,$sum(X3,$uminus(X2))) = X3 [forward demodulation 359,173] {a:1,w:8,inD:0}
[SA] passive: 380. $sum(X2,$sum(X3,$uminus(X2))) = X3 [forward demodulation 359,173] {a:1,w:8,inD:0}
[SA] passive: 360. $sum(X4,$sum(X6,X5)) = $sum($sum(X5,X4),X6) [superposition 118,25] {a:1,w:11,inD:0}
[SA] passive: 362. $sum(i(l9(s(X10))),X11) = $sum(i(l9(X10)),$sum(X11,1)) [superposition 118,153] {a:1,w:14,inD:0}
[SA] new: 381. $sum(X12,$sum($sum(X14,X13),X15)) = $sum($sum(X12,X13),$sum(X15,X14)) [forward demodulation 363,26] {a:1,w:15,inD:0}
[SA] new: 382. $sum(X12,$sum(X13,$sum(X15,X14))) = $sum(X12,$sum($sum(X14,X13),X15)) [forward demodulation 381,26] {a:1,w:15,inD:0}
[SA] new: 383. $sum(X12,$sum(X13,$sum(X15,X14))) = $sum(X12,$sum(X14,$sum(X13,X15))) [forward demodulation 382,26] {a:1,w:15,inD:0}
[SA] passive: 383. $sum(X12,$sum(X13,$sum(X15,X14))) = $sum(X12,$sum(X14,$sum(X13,X15))) [forward demodulation 382,26] {a:1,w:15,inD:0}
[SA] passive: 364. $sum(X16,X17) = $sum(0,$sum(X17,X16)) [superposition 118,173] {a:2,w:9,inD:0}
[SA] new: 384. $sum($uminus(X18),$sum(X19,X18)) = X19 [forward demodulation 365,173] {a:2,w:8,inD:0}
[SA] passive: 384. $sum($uminus(X18),$sum(X19,X18)) = X19 [forward demodulation 365,173] {a:2,w:8,inD:0}
[SA] passive: 366. $sum($uminus($sum(X21,X20)),X22) = $sum($uminus(X20),$sum(X22,$uminus(X21))) [superposition 118,28] {a:1,w:14,inD:0}
[SA] new: 385. $sum(X5,$sum(X6,$uminus($sum(X5,X6)))) = 0 [forward demodulation 368,25] {a:1,w:10,inD:0}
[SA] passive: 385. $sum(X5,$sum(X6,$uminus($sum(X5,X6)))) = 0 [forward demodulation 368,25] {a:1,w:10,inD:0}
[SA] passive: 369. $sum(X7,$sum(X9,X8)) = $sum(X9,$sum(X7,X8)) [superposition 118,25] {a:1,w:11,inD:0}
[SA] passive: 371. $less($sum(X3,X2),$sum(X0,$sum(X2,X1))) | ~$less(X3,$sum(X0,X1)) [superposition 33,118] {a:1,w:14,inD:0}
[SA] passive: 372. $less($sum(X4,$sum(X6,X5)),$sum(X7,X6)) | ~$less($sum(X4,X5),X7) [superposition 33,118] {a:1,w:14,inD:0}
[SA] passive: 375. ~$less(X16,$sum(X14,$sum(1,X15))) | ~$less($sum(X14,X15),X16) [superposition 36,118] {a:1,w:12,inD:0}
[SA] passive: 376. $less($sum(X17,X18),$sum(X17,$sum(1,X18))) [superposition 192,118] {a:2,w:9,inD:0}
[SA] passive: 377. $less(X21,$sum(X19,$sum(1,X20))) | $less($sum(X19,X20),X21) [superposition 34,118] {a:1,w:12,inD:0}
[SA] new: 386. $sum(X24,$sum(X25,$uminus($sum(X24,X25)))) = 0 [forward demodulation 379,25] {a:1,w:10,inD:0}
[SA] active: 235. ~$less(X13,-1) | $less(X13,0) [resolution 31,202] {a:3,w:6,nSel:1,tD:1,inD:0}
[SA] new: 387. $less(X0,0) | $less(-1,X0) | -1 = X0 [resolution 235,32] {a:4,w:9,tD:1,inD:0}
[SA] new: 388. $less(X1,0) | $less(-1,X1) | -1 = X1 [resolution 235,32] {a:4,w:9,tD:1,inD:0}
[SA] new: 389. $less(0,0) | $less(-1,1) [resolution 235,196] {a:4,w:6,tD:1,inD:0}
[SA] new: 390. $less(-1,0) | $less(-1,0) [resolution 235,197] {a:4,w:6,tD:1,inD:0}
[SA] new: 391. $less(-1,0) [duplicate literal removal 390] {a:4,w:3,tD:1,inD:0}
[SA] active: 166. 'Sub'(sK4(X0),nl9) | a(main_end,X0) = a(l7,X0) [backward demodulation 96,135] {a:0,w:11,nSel:1,inD:0}
[SA] new: 392. a(main_end,X0) = a(l7,X0) | $less(i(l9(sK4(X0))),alength) [resolution 166,100] {a:1,w:13,inD:0}
[SA] passive: 392. a(main_end,X0) = a(l7,X0) | $less(i(l9(sK4(X0))),alength) [resolution 166,100] {a:1,w:13,inD:0}
[SA] active: 262. ~$less(X0,-1) | $less(X0,1) [resolution 254,31] {a:4,w:6,nSel:1,inD:0}
[SA] new: 393. $less(X0,1) | $less(-1,X0) | -1 = X0 [resolution 262,32] {a:5,w:9,inD:0}
[SA] new: 394. $less(X1,1) | $less(-1,X1) | -1 = X1 [resolution 262,32] {a:5,w:9,inD:0}
[SA] new: 395. $less(0,1) | $less(-1,1) [resolution 262,196] {a:5,w:6,inD:0}
[SA] new: 396. $less(-1,1) | $less(-1,0) [resolution 262,197] {a:5,w:6,inD:0}
[SA] active: 292. ~$less(X0,i(main_end)) | $less(X0,0) | 'Sub'(sK2(X0),nl9) [trivial inequality removal 287] {a:0,w:11,nSel:1,inD:0}
[SA] new: 397. $less(X0,0) | 'Sub'(sK2(X0),nl9) | $less(i(main_end),X0) | i(main_end) = X0 [resolution 292,32] {a:1,w:15,inD:0}
[SA] new: 398. $less(X1,0) | 'Sub'(sK2(X1),nl9) | $less(i(main_end),X1) | i(main_end) = X1 [resolution 292,32] {a:1,w:15,inD:0}
[SA] new: 399. $less(0,0) | 'Sub'(sK2(0),nl9) | $less(i(main_end),1) [resolution 292,196] {a:3,w:11,inD:0}
[SA] new: 400. $less(-1,0) | 'Sub'(sK2(-1),nl9) | $less(i(main_end),0) [resolution 292,197] {a:3,w:11,inD:0}
[SA] passive: 397. $less(X0,0) | 'Sub'(sK2(X0),nl9) | $less(i(main_end),X0) | i(main_end) = X0 [resolution 292,32] {a:1,w:15,inD:0}
[SA] new: 401. 'Sub'(sK2(0),nl9) | $less(i(main_end),1) [subsumption resolution 399,30] {a:3,w:8,inD:0}
[SA] passive: 401. 'Sub'(sK2(0),nl9) | $less(i(main_end),1) [subsumption resolution 399,30] {a:3,w:8,inD:0}
[SA] active: 351. $less(X0,1) | $less(-1,X0) [resolution 232,197] {a:4,w:6,nSel:2,tD:1,inD:0}
[SA] new: 402. $less(-1,1) [resolution 351,30] {a:5,w:3,tD:1,inD:0}
[SA] new: 403. $less(-1,X0) | ~$less(X1,X0) | $less(X1,1) [resolution 351,31] {a:5,w:9,tD:1,inD:0}
[SA] new: 404. $less(-1,0) | ~$less(1,1) [resolution 351,208] {a:5,w:6,tD:1,inD:0}
[SA] new: 405. $less(-1,-1) | ~$less(1,0) [resolution 351,209] {a:5,w:6,tD:1,inD:0}
[SA] new: 406. $less(X0,1) | ~$less(X0,0) [resolution 351,209] {a:5,w:6,tD:1,inD:0}
[SA] new: 407. $less(-1,1) [resolution 351,30] {a:5,w:3,tD:1,inD:0}
[SA] new: 408. $less(X1,1) | ~$less(X2,-1) | $less(X2,X1) [resolution 351,31] {a:5,w:9,tD:1,inD:0}
[SA] new: 409. $less(i(main_end),1) | $less(-1,0) | 'Sub'(sK2(-1),nl9) [resolution 351,292] {a:5,w:11,inD:0}
[SA] new: 410. $less($sum(X3,1),1) | ~$less(X3,-1) [resolution 351,36] {a:5,w:8,tD:1,inD:0}
[SA] new: 411. $less(0,1) | $less(-1,1) [resolution 351,232] {a:5,w:6,tD:1,inD:0}
[SA] new: 412. $less(2,1) | $less(-1,1) [resolution 351,234] {a:5,w:6,goal:1,inD:0}
[SA] new: 413. $less(-1,1) | $less(-1,1) [resolution 351,262] {a:5,w:6,inD:0}
[SA] new: 414. $less(-1,1) | $less(-1,0) [resolution 351,235] {a:5,w:6,tD:1,inD:0}
[SA] new: 415. $less(-1,1) [duplicate literal removal 413] {a:5,w:3,inD:0}
[SA] new: 416. $less(X3,$sum(1,-1)) | ~$less(X3,-1) [evaluation 410] {a:5,w:8,tD:1,inD:0}
[SA] new: 417. $less(X3,0) | ~$less(X3,-1) [evaluation 416] {a:5,w:6,tD:1,inD:0}
[SA] passive: 403. $less(-1,X0) | ~$less(X1,X0) | $less(X1,1) [resolution 351,31] {a:5,w:9,tD:1,inD:0}
[SA] passive: 408. $less(X1,1) | ~$less(X2,-1) | $less(X2,X1) [resolution 351,31] {a:5,w:9,tD:1,inD:0}
[SA] active: 26. $sum(X0,$sum(X1,X2)) = $sum($sum(X0,X1),X2) [theory axiom] {a:0,w:11,nSel:1,tD:1,inD:0}
[SA] new: 418. $sum(X0,X1) = $sum(X0,$sum(0,X1)) [superposition 26,27] {a:1,w:9,tD:1,inD:0}
[SA] new: 419. $sum(0,X3) = $sum(X2,$sum($uminus(X2),X3)) [superposition 26,29] {a:1,w:10,tD:1,inD:0}
[SA] new: 420. $sum(X4,$sum(X5,X6)) = $sum($sum(X5,X4),X6) [superposition 26,25] {a:1,w:11,tD:1,inD:0}
[SA] new: 421. $sum(X7,$sum(X8,X9)) = $sum($sum(X8,X7),X9) [superposition 26,25] {a:1,w:11,tD:1,inD:0}
[SA] new: 422. $sum(i(l9(s(X10))),X11) = $sum(i(l9(X10)),$sum(1,X11)) [superposition 26,153] {a:1,w:14,inD:0}
[SA] new: 423. $sum($sum(X12,X13),$sum(X14,X15)) = $sum($sum(X12,$sum(X13,X14)),X15) [superposition 26,26] {a:1,w:15,tD:1,inD:0}
[SA] new: 424. $sum($sum(X16,X17),$sum(X18,X19)) = $sum($sum(X16,$sum(X18,X17)),X19) [superposition 26,118] {a:1,w:15,inD:0}
[SA] new: 425. $sum(X20,X21) = $sum(0,$sum(X20,X21)) [superposition 26,173] {a:2,w:9,tD:1,inD:0}
[SA] new: 426. $sum(0,X23) = $sum($uminus(X22),$sum(X22,X23)) [superposition 26,167] {a:2,w:10,tD:1,inD:0}
[SA] new: 427. $sum($uminus($sum(X25,X24)),X26) = $sum($uminus(X24),$sum($uminus(X25),X26)) [superposition 26,28] {a:1,w:14,tD:1,inD:0}
[SA] new: 428. $sum(X3,$sum(X4,X5)) = $sum(X3,$sum(X5,X4)) [superposition 26,118] {a:1,w:11,inD:0}
[SA] new: 429. $sum(X6,X7) = $sum(X6,$sum(X7,0)) [superposition 26,27] {a:1,w:9,tD:1,inD:0}
[SA] new: 430. $sum(X8,$sum(X9,$uminus($sum(X8,X9)))) = 0 [superposition 26,29] {a:1,w:10,tD:1,inD:0}
[SA] new: 431. $sum(X10,$sum(X11,X12)) = $sum(X12,$sum(X10,X11)) [superposition 26,25] {a:1,w:11,tD:1,inD:0}
[SA] new: 432. $sum(X13,$sum(X14,X15)) = $sum(X15,$sum(X13,X14)) [superposition 26,25] {a:1,w:11,tD:1,inD:0}
[SA] new: 433. $sum(X0,$sum(X1,X2)) = $sum(X0,$sum(X2,X1)) [superposition 118,26] {a:1,w:11,inD:0}
[SA] new: 434. $sum($sum(X3,X4),$sum(X6,X5)) = $sum($sum(X3,$sum(X4,X5)),X6) [superposition 118,26] {a:1,w:15,inD:0}
[SA] new: 435. $less($sum(X10,X9),$sum(X7,$sum(X8,X9))) | ~$less(X10,$sum(X7,X8)) [superposition 33,26] {a:1,w:14,tD:1,inD:0}
[SA] new: 436. $less($sum(X11,$sum(X12,X13)),$sum(X14,X13)) | ~$less($sum(X11,X12),X14) [superposition 33,26] {a:1,w:14,tD:1,inD:0}
[SA] new: 437. $sum(X15,$sum(X16,X17)) = $sum(X17,$sum(X15,X16)) [superposition 25,26] {a:1,w:11,tD:1,inD:0}
[SA] new: 438. $sum(X18,$sum(X19,X20)) = $sum(X20,$sum(X18,X19)) [superposition 25,26] {a:1,w:11,tD:1,inD:0}
[SA] new: 439. ~$less(X23,$sum(X21,$sum(X22,1))) | ~$less($sum(X21,X22),X23) [superposition 36,26] {a:1,w:12,tD:1,inD:0}
[SA] new: 440. $less($sum(X24,X25),$sum(X24,$sum(X25,1))) [superposition 192,26] {a:2,w:9,tD:1,inD:0}
[SA] new: 441. $less(X28,$sum(X26,$sum(X27,1))) | $less($sum(X26,X27),X28) [superposition 34,26] {a:1,w:12,tD:1,inD:0}
[SA] new: 442. $sum(X29,X30) = $sum(X29,$sum(X30,0)) [superposition 27,26] {a:1,w:9,tD:1,inD:0}
[SA] new: 443. $sum(X31,$sum(X32,$uminus($sum(X31,X32)))) = 0 [superposition 29,26] {a:1,w:10,tD:1,inD:0}
[SA] new: 444. $sum(X2,$sum($uminus(X2),X3)) = X3 [forward demodulation 419,173] {a:1,w:8,tD:1,inD:0}
[SA] passive: 444. $sum(X2,$sum($uminus(X2),X3)) = X3 [forward demodulation 419,173] {a:1,w:8,tD:1,inD:0}
[SA] passive: 420. $sum(X4,$sum(X5,X6)) = $sum($sum(X5,X4),X6) [superposition 26,25] {a:1,w:11,tD:1,inD:0}
[SA] passive: 422. $sum(i(l9(s(X10))),X11) = $sum(i(l9(X10)),$sum(1,X11)) [superposition 26,153] {a:1,w:14,inD:0}
[SA] new: 445. $sum(X12,$sum($sum(X13,X14),X15)) = $sum($sum(X12,X13),$sum(X14,X15)) [forward demodulation 423,26] {a:1,w:15,tD:1,inD:0}
[SA] new: 446. $sum(X12,$sum(X13,$sum(X14,X15))) = $sum(X12,$sum($sum(X13,X14),X15)) [forward demodulation 445,26] {a:1,w:15,tD:1,inD:0}
[SA] new: 447. $sum(X16,$sum($sum(X18,X17),X19)) = $sum($sum(X16,X17),$sum(X18,X19)) [forward demodulation 424,26] {a:1,w:15,inD:0}
[SA] new: 448. $sum(X16,$sum(X17,$sum(X18,X19))) = $sum(X16,$sum($sum(X18,X17),X19)) [forward demodulation 447,26] {a:1,w:15,inD:0}
[SA] new: 449. $sum(X16,$sum(X17,$sum(X18,X19))) = $sum(X16,$sum(X18,$sum(X17,X19))) [forward demodulation 448,26] {a:1,w:15,inD:0}
[SA] passive: 449. $sum(X16,$sum(X17,$sum(X18,X19))) = $sum(X16,$sum(X18,$sum(X17,X19))) [forward demodulation 448,26] {a:1,w:15,inD:0}
[SA] new: 450. $sum($uminus(X22),$sum(X22,X23)) = X23 [forward demodulation 426,173] {a:2,w:8,tD:1,inD:0}
[SA] passive: 450. $sum($uminus(X22),$sum(X22,X23)) = X23 [forward demodulation 426,173] {a:2,w:8,tD:1,inD:0}
[SA] passive: 427. $sum($uminus($sum(X25,X24)),X26) = $sum($uminus(X24),$sum($uminus(X25),X26)) [superposition 26,28] {a:1,w:14,tD:1,inD:0}
[SA] passive: 428. $sum(X3,$sum(X4,X5)) = $sum(X3,$sum(X5,X4)) [superposition 26,118] {a:1,w:11,inD:0}
[SA] passive: 431. $sum(X10,$sum(X11,X12)) = $sum(X12,$sum(X10,X11)) [superposition 26,25] {a:1,w:11,tD:1,inD:0}
[SA] new: 451. $sum(X3,$sum($sum(X4,X5),X6)) = $sum($sum(X3,X4),$sum(X6,X5)) [forward demodulation 434,26] {a:1,w:15,inD:0}
[SA] new: 452. $sum(X3,$sum(X4,$sum(X6,X5))) = $sum(X3,$sum($sum(X4,X5),X6)) [forward demodulation 451,26] {a:1,w:15,inD:0}
[SA] new: 453. $sum(X3,$sum(X4,$sum(X5,X6))) = $sum(X3,$sum(X4,$sum(X6,X5))) [forward demodulation 452,26] {a:1,w:15,inD:0}
[SA] passive: 453. $sum(X3,$sum(X4,$sum(X5,X6))) = $sum(X3,$sum(X4,$sum(X6,X5))) [forward demodulation 452,26] {a:1,w:15,inD:0}
[SA] passive: 435. $less($sum(X10,X9),$sum(X7,$sum(X8,X9))) | ~$less(X10,$sum(X7,X8)) [superposition 33,26] {a:1,w:14,tD:1,inD:0}
[SA] passive: 436. $less($sum(X11,$sum(X12,X13)),$sum(X14,X13)) | ~$less($sum(X11,X12),X14) [superposition 33,26] {a:1,w:14,tD:1,inD:0}
[SA] passive: 439. ~$less(X23,$sum(X21,$sum(X22,1))) | ~$less($sum(X21,X22),X23) [superposition 36,26] {a:1,w:12,tD:1,inD:0}
[SA] passive: 440. $less($sum(X24,X25),$sum(X24,$sum(X25,1))) [superposition 192,26] {a:2,w:9,tD:1,inD:0}
[SA] passive: 441. $less(X28,$sum(X26,$sum(X27,1))) | $less($sum(X26,X27),X28) [superposition 34,26] {a:1,w:12,tD:1,inD:0}
[SA] active: 191. v = a(l12(nl9),i(main_end)) [superposition 148,98] {a:1,w:7,nSel:1,inD:0}
[SA] active: 160. i(l9(sK6)) != i(l9(s(sK6))) | i(main_end) = 0 [backward demodulation 98,130] {a:0,w:12,nSel:1,inD:0}
[SA] active: 215. 'Sub'(p(X7),X7) | zero = X7 [superposition 117,74] {a:1,w:7,nSel:1,inD:0}
[SA] new: 454. zero = nl9 | $less(i(l9(p(nl9))),alength) [resolution 215,100] {a:2,w:9,inD:0}
[SA] passive: 454. zero = nl9 | $less(i(l9(p(nl9))),alength) [resolution 215,100] {a:2,w:9,inD:0}
[SA] active: 291. ~$less(X0,i(main_end)) | $less(X0,0) | i(l9(sK2(X0))) = X0 [trivial inequality removal 288] {a:0,w:13,nSel:1,inD:0}
[SA] new: 455. $less(X0,0) | i(l9(sK2(X0))) = X0 | $less(i(main_end),X0) | i(main_end) = X0 [resolution 291,32] {a:1,w:17,inD:0}
[SA] new: 456. $less(X1,0) | i(l9(sK2(X1))) = X1 | $less(i(main_end),X1) | i(main_end) = X1 [resolution 291,32] {a:1,w:17,inD:0}
[SA] new: 457. $less(0,0) | i(l9(sK2(0))) = 0 | $less(i(main_end),1) [resolution 291,196] {a:3,w:13,inD:0}
[SA] new: 458. $less(-1,0) | i(l9(sK2(-1))) = -1 | $less(i(main_end),1) [resolution 291,351] {a:5,w:13,inD:0}
[SA] new: 459. $less(-1,0) | i(l9(sK2(-1))) = -1 | $less(i(main_end),0) [resolution 291,197] {a:3,w:13,inD:0}
[SA] passive: 455. $less(X0,0) | i(l9(sK2(X0))) = X0 | $less(i(main_end),X0) | i(main_end) = X0 [resolution 291,32] {a:1,w:17,inD:0}
[SA] new: 460. i(l9(sK2(0))) = 0 | $less(i(main_end),1) [subsumption resolution 457,30] {a:3,w:10,inD:0}
[SA] passive: 460. i(l9(sK2(0))) = 0 | $less(i(main_end),1) [subsumption resolution 457,30] {a:3,w:10,inD:0}
[SA] active: 294. v = a(l12(s(zero)),1) [superposition 148,285] {a:2,w:7,nSel:1,inD:0}
[SA] active: 150. a(l9(X4),X5) = a(l12(X4),X5) | i(l9(X4)) = X5 [forward demodulation 149,110] {a:0,w:14,nSel:1,inD:0}
[SA] new: 462. v = a(l9(nl9),i(main_end)) | i(main_end) = i(l9(nl9)) [superposition 150,191] {a:2,w:13,inD:0}
[SA] new: 463. v = a(l9(s(zero)),1) | i(l9(s(zero))) = 1 [superposition 150,294] {a:3,w:13,inD:0}
[SA] new: 464. v = a(l9(zero),0) | i(l9(zero)) = 0 [superposition 150,190] {a:2,w:11,inD:0}
[SA] new: 466. v = a(l9(nl9),i(main_end)) | i(main_end) = i(l9(nl9)) [superposition 191,150] {a:2,w:13,inD:0}
[SA] new: 467. v = a(l9(s(zero)),1) | i(l9(s(zero))) = 1 [superposition 294,150] {a:3,w:13,inD:0}
[SA] new: 468. v = a(l9(zero),0) | i(l9(zero)) = 0 [superposition 190,150] {a:2,w:11,inD:0}
[SA] active: 295. i(l9(s(s(zero)))) = 2 [evaluation 293] {a:2,w:7,nSel:1,inD:0}
[SA] new: 469. i(l9(s(s(s(zero))))) = $sum(2,1) [superposition 153,295] {a:3,w:10,inD:0}
[SA] new: 470. v = a(l12(s(s(zero))),2) [superposition 148,295] {a:3,w:8,inD:0}
[SA] new: 471. i(l9(s(s(s(zero))))) = 3 [evaluation 469] {a:3,w:8,inD:0}
[SA] passive: 471. i(l9(s(s(s(zero))))) = 3 [evaluation 469] {a:3,w:8,inD:0}
[SA] passive: 470. v = a(l12(s(s(zero))),2) [superposition 148,295] {a:3,w:8,inD:0}
[SA] active: 290. i(l9(X1)) != i(l9(X2)) | X1 = X2 | ~'Sub'(s(X0),nl9) [trivial inequality removal 289] {a:0,w:14,nSel:1,inD:0}
[SA] new: 472. i(l9(X0)) != 0 | zero = X0 | ~'Sub'(s(X1),nl9) [superposition 290,113] {a:1,w:12,inD:0}
[SA] new: 473. i(main_end) != i(l9(X2)) | nl9 = X2 | ~'Sub'(s(X3),nl9) [superposition 290,98] {a:1,w:13,inD:0}
[SA] new: 474. i(l9(X4)) != 1 | s(zero) = X4 | ~'Sub'(s(X5),nl9) [superposition 290,285] {a:2,w:13,inD:0}
[SA] new: 475. i(l9(X6)) != 2 | s(s(zero)) = X6 | ~'Sub'(s(X7),nl9) [superposition 290,295] {a:3,w:14,inD:0}
[SA] new: 476. i(l9(X0)) != 0 | zero = X0 | ~'Sub'(s(X1),nl9) [superposition 290,113] {a:1,w:12,inD:0}
[SA] new: 477. i(main_end) != i(l9(X2)) | nl9 = X2 | ~'Sub'(s(X3),nl9) [superposition 290,98] {a:1,w:13,inD:0}
[SA] new: 478. i(l9(X4)) != 1 | s(zero) = X4 | ~'Sub'(s(X5),nl9) [superposition 290,285] {a:2,w:13,inD:0}
[SA] new: 479. i(l9(X6)) != 2 | s(s(zero)) = X6 | ~'Sub'(s(X7),nl9) [superposition 290,295] {a:3,w:14,inD:0}
[SA] new: 480. X0 = X0 | ~'Sub'(s(X1),nl9) [equality resolution 290] {a:1,w:7,inD:0}
[SA] passive: 472. i(l9(X0)) != 0 | zero = X0 | ~'Sub'(s(X1),nl9) [superposition 290,113] {a:1,w:12,inD:0}
[SA] passive: 473. i(main_end) != i(l9(X2)) | nl9 = X2 | ~'Sub'(s(X3),nl9) [superposition 290,98] {a:1,w:13,inD:0}
[SA] passive: 474. i(l9(X4)) != 1 | s(zero) = X4 | ~'Sub'(s(X5),nl9) [superposition 290,285] {a:2,w:13,inD:0}
[SA] passive: 475. i(l9(X6)) != 2 | s(s(zero)) = X6 | ~'Sub'(s(X7),nl9) [superposition 290,295] {a:3,w:14,inD:0}
[SA] active: 223. $less(alength,1) | s(sK3) = nl9 [resolution 196,131] {a:3,w:7,nSel:1,inD:0}
[SA] new: 481. s(sK3) = nl9 | ~$less(X0,alength) | $less(X0,1) [resolution 223,31] {a:4,w:10,inD:0}
[SA] passive: 481. s(sK3) = nl9 | ~$less(X0,alength) | $less(X0,1) [resolution 223,31] {a:4,w:10,inD:0}
[SA] active: 165. a(l9(sK4(X0)),X0) != a(l9(s(sK4(X0))),X0) | a(main_end,X0) = a(l7,X0) [backward demodulation 96,136] {a:0,w:19,nSel:1,inD:0}
[SA] new: 482. a(l9(sK4(X0)),X0) != a(l12(sK4(X0)),X0) | a(main_end,X0) = a(l7,X0) [superposition 165,102] {a:1,w:18,inD:0}
[SA] passive: 482. a(l9(sK4(X0)),X0) != a(l12(sK4(X0)),X0) | a(main_end,X0) = a(l7,X0) [superposition 165,102] {a:1,w:18,inD:0}
[SA] active: 241. $less(alength,i(main_end)) | alength = i(main_end) [resolution 32,157] {a:1,w:8,nSel:1,inD:0}
[SA] new: 483. alength = i(main_end) | $less(alength,0) | alength = i(l9(sK2(alength))) [resolution 241,291] {a:2,w:13,inD:0}
[SA] new: 484. alength = i(main_end) | $less(alength,0) | 'Sub'(sK2(alength),nl9) [resolution 241,292] {a:2,w:11,inD:0}
[SA] new: 485. alength = i(main_end) | ~$less(X0,alength) | $less(X0,i(main_end)) [resolution 241,31] {a:2,w:11,inD:0}
[SA] passive: 483. alength = i(main_end) | $less(alength,0) | alength = i(l9(sK2(alength))) [resolution 241,291] {a:2,w:13,inD:0}
[SA] passive: 484. alength = i(main_end) | $less(alength,0) | 'Sub'(sK2(alength),nl9) [resolution 241,292] {a:2,w:11,inD:0}
[SA] passive: 485. alength = i(main_end) | ~$less(X0,alength) | $less(X0,i(main_end)) [resolution 241,31] {a:2,w:11,inD:0}
[SA] active: 163. a(l12(X1),X2) != $sum(1,a(l9(X1),X2)) | ~$less(X0,a(main_end,X2)) | $less(X0,a(l7,X2)) | 'Sub'(sK0(X0,X2),nl9) [backward demodulation 96,155] {a:0,w:26,nSel:1,inD:0}
[SA] new: 486. a(l12(nl9),X0) != $sum(1,a(main_end,X0)) | ~$less(X1,a(main_end,X0)) | $less(X1,a(l7,X0)) | 'Sub'(sK0(X1,X0),nl9) [superposition 163,96] {a:1,w:25,inD:0}
[SA] new: 487. a(l12(zero),X2) != $sum(1,a(l7,X2)) | ~$less(X3,a(main_end,X2)) | $less(X3,a(l7,X2)) | 'Sub'(sK0(X3,X2),nl9) [superposition 163,134] {a:1,w:25,inD:0}
[SA] new: 488. a(l12(s(X4)),X5) != $sum(1,a(l12(X4),X5)) | ~$less(X6,a(main_end,X5)) | $less(X6,a(l7,X5)) | 'Sub'(sK0(X6,X5),nl9) [superposition 163,102] {a:1,w:27,inD:0}
[SA] passive: 486. a(l12(nl9),X0) != $sum(1,a(main_end,X0)) | ~$less(X1,a(main_end,X0)) | $less(X1,a(l7,X0)) | 'Sub'(sK0(X1,X0),nl9) [superposition 163,96] {a:1,w:25,inD:0}
[SA] passive: 487. a(l12(zero),X2) != $sum(1,a(l7,X2)) | ~$less(X3,a(main_end,X2)) | $less(X3,a(l7,X2)) | 'Sub'(sK0(X3,X2),nl9) [superposition 163,134] {a:1,w:25,inD:0}
[SA] passive: 488. a(l12(s(X4)),X5) != $sum(1,a(l12(X4),X5)) | ~$less(X6,a(main_end,X5)) | $less(X6,a(l7,X5)) | 'Sub'(sK0(X6,X5),nl9) [superposition 163,102] {a:1,w:27,inD:0}
[SA] active: 380. $sum(X2,$sum(X3,$uminus(X2))) = X3 [forward demodulation 359,173] {a:1,w:8,nSel:1,inD:0}
[SA] new: 489. $sum($uminus(X0),$sum(X1,X0)) = X1 [superposition 380,35] {a:2,w:8,inD:0}
[SA] new: 490. $sum(X0,0) = X0 [superposition 380,29] {a:2,w:5,inD:0}
[SA] new: 491. $sum(X2,$sum($uminus(X2),X1)) = X1 [superposition 380,25] {a:2,w:8,inD:0}
[SA] new: 492. $sum(X4,$sum($uminus(X4),X3)) = X3 [superposition 380,25] {a:2,w:8,inD:0}
[SA] new: 493. $sum(X5,X6) = $sum(X7,$sum(X5,$sum(X6,$uminus(X7)))) [superposition 380,26] {a:2,w:12,inD:0}
[SA] new: 494. $sum(X8,X9) = $sum(X10,$sum(X8,$sum($uminus(X10),X9))) [superposition 380,118] {a:2,w:12,inD:0}
[SA] new: 495. $sum(X11,$uminus(X11)) = 0 [superposition 380,173] {a:2,w:6,inD:0}
[SA] new: 496. $sum(X13,$uminus($sum(X13,X12))) = $uminus(X12) [superposition 380,28] {a:2,w:9,inD:0}
[SA] new: 497. $sum(X14,0) = $uminus($uminus(X14)) [superposition 380,167] {a:2,w:7,inD:0}
[SA] new: 498. $sum(X6,$sum(X7,$sum(X8,$uminus($sum(X6,X7))))) = X8 [superposition 380,26] {a:2,w:12,inD:0}
[SA] new: 499. $sum(X9,$sum($sum(X11,$uminus($sum(X9,X10))),X10)) = X11 [superposition 380,118] {a:2,w:12,inD:0}
[SA] new: 500. $sum(X12,$uminus(0)) = X12 [superposition 380,173] {a:2,w:6,inD:0}
[SA] new: 501. $sum(X1,X2) = $sum(X0,$sum($sum(X1,$uminus(X0)),X2)) [superposition 26,380] {a:2,w:12,inD:0}
[SA] new: 502. $sum(X4,X5) = $sum(X3,$sum(X5,$sum(X4,$uminus(X3)))) [superposition 118,380] {a:2,w:12,inD:0}
[SA] new: 503. $less($sum(X8,$sum(X7,$uminus(X6))),X7) | ~$less(X8,X6) [superposition 33,380] {a:2,w:11,inD:0}
[SA] new: 504. $less(X10,$sum(X11,$sum(X10,$uminus(X9)))) | ~$less(X9,X11) [superposition 33,380] {a:2,w:11,inD:0}
[SA] new: 505. $less($sum(X16,$uminus(1)),X16) [superposition 198,380] {a:3,w:6,inD:0}
[SA] new: 506. $sum(X17,$sum(X18,$sum(X19,$uminus($sum(X17,X18))))) = X19 [superposition 26,380] {a:2,w:12,inD:0}
[SA] new: 507. $sum(X20,$sum($sum(X22,$uminus($sum(X20,X21))),X21)) = X22 [superposition 118,380] {a:2,w:12,inD:0}
[SA] new: 508. $sum(X23,$uminus(0)) = X23 [superposition 173,380] {a:2,w:6,inD:0}
[SA] new: 509. $sum(X23,0) = X23 [evaluation 508] {a:2,w:5,inD:0}
[SA] new: 510. $less($sum(X16,-1),X16) [evaluation 505] {a:3,w:5,inD:0}
[SA] new: 511. $sum(X12,0) = X12 [evaluation 500] {a:2,w:5,inD:0}
[SA] passive: 493. $sum(X5,X6) = $sum(X7,$sum(X5,$sum(X6,$uminus(X7)))) [superposition 380,26] {a:2,w:12,inD:0}
[SA] passive: 494. $sum(X8,X9) = $sum(X10,$sum(X8,$sum($uminus(X10),X9))) [superposition 380,118] {a:2,w:12,inD:0}
[SA] passive: 496. $sum(X13,$uminus($sum(X13,X12))) = $uminus(X12) [superposition 380,28] {a:2,w:9,inD:0}
[SA] new: 512. $sum(X14,0) = X14 [forward demodulation 497,35] {a:2,w:5,inD:0}
[SA] passive: 498. $sum(X6,$sum(X7,$sum(X8,$uminus($sum(X6,X7))))) = X8 [superposition 380,26] {a:2,w:12,inD:0}
[SA] new: 513. $sum(X9,$sum(X10,$sum(X11,$uminus($sum(X9,X10))))) = X11 [forward demodulation 499,25] {a:2,w:12,inD:0}
[SA] new: 514. $sum(X1,X2) = $sum(X0,$sum(X1,$sum($uminus(X0),X2))) [forward demodulation 501,26] {a:2,w:12,inD:0}
[SA] passive: 502. $sum(X4,X5) = $sum(X3,$sum(X5,$sum(X4,$uminus(X3)))) [superposition 118,380] {a:2,w:12,inD:0}
[SA] passive: 503. $less($sum(X8,$sum(X7,$uminus(X6))),X7) | ~$less(X8,X6) [superposition 33,380] {a:2,w:11,inD:0}
[SA] passive: 504. $less(X10,$sum(X11,$sum(X10,$uminus(X9)))) | ~$less(X9,X11) [superposition 33,380] {a:2,w:11,inD:0}
[SA] passive: 510. $less($sum(X16,-1),X16) [evaluation 505] {a:3,w:5,inD:0}
[SA] new: 515. $sum(X20,$sum(X21,$sum(X22,$uminus($sum(X20,X21))))) = X22 [forward demodulation 507,25] {a:2,w:12,inD:0}
[SA] active: 164. a(l12(X1),X2) != $sum(1,a(l9(X1),X2)) | ~$less(X0,a(main_end,X2)) | $less(X0,a(l7,X2)) | a(l9(sK0(X0,X2)),X2) = X0 [backward demodulation 96,154] {a:0,w:29,nSel:1,inD:0}
[SA] new: 516. a(l12(nl9),X0) != $sum(1,a(main_end,X0)) | ~$less(X1,a(main_end,X0)) | $less(X1,a(l7,X0)) | a(l9(sK0(X1,X0)),X0) = X1 [superposition 164,96] {a:1,w:28,inD:0}
[SA] new: 517. a(l12(zero),X2) != $sum(1,a(l7,X2)) | ~$less(X3,a(main_end,X2)) | $less(X3,a(l7,X2)) | a(l9(sK0(X3,X2)),X2) = X3 [superposition 164,134] {a:1,w:28,inD:0}
[SA] new: 518. a(l12(s(X4)),X5) != $sum(1,a(l12(X4),X5)) | ~$less(X6,a(main_end,X5)) | $less(X6,a(l7,X5)) | a(l9(sK0(X6,X5)),X5) = X6 [superposition 164,102] {a:1,w:30,inD:0}
[SA] passive: 516. a(l12(nl9),X0) != $sum(1,a(main_end,X0)) | ~$less(X1,a(main_end,X0)) | $less(X1,a(l7,X0)) | a(l9(sK0(X1,X0)),X0) = X1 [superposition 164,96] {a:1,w:28,inD:0}
[SA] passive: 517. a(l12(zero),X2) != $sum(1,a(l7,X2)) | ~$less(X3,a(main_end,X2)) | $less(X3,a(l7,X2)) | a(l9(sK0(X3,X2)),X2) = X3 [superposition 164,134] {a:1,w:28,inD:0}
[SA] passive: 518. a(l12(s(X4)),X5) != $sum(1,a(l12(X4),X5)) | ~$less(X6,a(main_end,X5)) | $less(X6,a(l7,X5)) | a(l9(sK0(X6,X5)),X5) = X6 [superposition 164,102] {a:1,w:30,inD:0}
[SA] active: 510. $less($sum(X16,-1),X16) [evaluation 505] {a:3,w:5,nSel:1,inD:0}
[SA] new: 519. ~$less(X0,$sum(X1,-1)) | $less(X0,X1) [resolution 510,31] {a:4,w:8,inD:0}
[SA] new: 520. $less($sum(i(main_end),-1),0) | i(l9(sK2($sum(i(main_end),-1)))) = $sum(i(main_end),-1) [resolution 510,291] {a:4,w:18,inD:0}
[SA] new: 521. $less($sum(i(main_end),-1),0) | 'Sub'(sK2($sum(i(main_end),-1)),nl9) [resolution 510,292] {a:4,w:13,inD:0}
[SA] new: 522. ~$less(X2,$sum($sum(X2,1),-1)) [resolution 510,36] {a:4,w:7,inD:0}
[SA] new: 523. $less($sum(0,-1),1) [resolution 510,232] {a:4,w:5,inD:0}
[SA] new: 524. $less($sum(2,-1),1) [resolution 510,234] {a:4,w:5,goal:1,inD:0}
[SA] new: 525. $less($sum(-1,-1),1) [resolution 510,262] {a:5,w:5,inD:0}
[SA] new: 526. $less($sum(-1,-1),0) [resolution 510,235] {a:4,w:5,inD:0}
[SA] new: 527. $less($sum(-1,X0),X0) [superposition 510,25] {a:4,w:5,inD:0}
[SA] new: 528. $less($sum(-1,X1),X1) [superposition 510,25] {a:4,w:5,inD:0}
[SA] new: 529. $less($sum(X2,$sum(X3,-1)),$sum(X2,X3)) [superposition 510,26] {a:4,w:9,inD:0}
[SA] new: 530. $less($sum(X4,$sum(-1,X5)),$sum(X4,X5)) [superposition 510,118] {a:4,w:9,inD:0}
[SA] new: 531. $less(-1,0) [superposition 510,173] {a:4,w:3,inD:0}
[SA] new: 532. $less(0,$uminus(-1)) [superposition 510,167] {a:4,w:4,inD:0}
[SA] new: 533. $less(0,1) [evaluation 532] {a:4,w:3,inD:0}
[SA] new: 534. $less(-2,0) [evaluation 526] {a:4,w:3,inD:0}
[SA] new: 535. $less(-2,1) [evaluation 525] {a:5,w:3,inD:0}
[SA] new: 536. $less(1,1) [evaluation 524] {a:4,w:3,goal:1,inD:0}
[SA] new: 537. $less(-1,1) [evaluation 523] {a:4,w:3,inD:0}
[SA] new: 538. ~$less(X2,X2) [evaluation 522] {a:4,w:3,inD:0}
[SA] new: 539. $less(i(main_end),$sum(0,1)) | 'Sub'(sK2($sum(i(main_end),-1)),nl9) [evaluation 521] {a:4,w:13,inD:0}
[SA] new: 540. $less(i(main_end),1) | 'Sub'(sK2($sum(i(main_end),-1)),nl9) [evaluation 539] {a:4,w:11,inD:0}
[SA] new: 541. $less(i(main_end),$sum(0,1)) | i(l9(sK2($sum(i(main_end),-1)))) = $sum(i(main_end),-1) [evaluation 520] {a:4,w:18,inD:0}
[SA] new: 542. $less(i(main_end),1) | i(l9(sK2($sum(i(main_end),-1)))) = $sum(i(main_end),-1) [evaluation 541] {a:4,w:16,inD:0}
[SA] passive: 519. ~$less(X0,$sum(X1,-1)) | $less(X0,X1) [resolution 510,31] {a:4,w:8,inD:0}
[SA] new: 543. i(l9(sK2($sum(-1,i(main_end))))) = $sum(-1,i(main_end)) | $less(i(main_end),1) [forward demodulation 542,25] {a:4,w:16,inD:0}
[SA] passive: 543. i(l9(sK2($sum(-1,i(main_end))))) = $sum(-1,i(main_end)) | $less(i(main_end),1) [forward demodulation 542,25] {a:4,w:16,inD:0}
[SA] new: 544. 'Sub'(sK2($sum(-1,i(main_end))),nl9) | $less(i(main_end),1) [forward demodulation 540,25] {a:4,w:11,inD:0}
[SA] passive: 544. 'Sub'(sK2($sum(-1,i(main_end))),nl9) | $less(i(main_end),1) [forward demodulation 540,25] {a:4,w:11,inD:0}
[SA] new: 545. $false [subsumption resolution 536,30] {a:4,w:0,goal:1,inD:0}
% Refutation found. Thanks to Tanya!
% SZS status Unsatisfiable for test-user-lemma-1
% SZS output start Proof for test-user-lemma-1
14. ! [X0 : $int,X2 : $int,X1 : $int] : $sum($sum(X0,X1),X2) = $sum($sum(X0,X2),X1) [input]
18. $lesseq(1,2) [input]
19. ~$lesseq(1,2) [negated conjecture 18]
20. ~~$less(2,1) [evaluation 19]
25. $sum(X0,X1) = $sum(X1,X0) [theory axiom] {a:0,w:7,nSel:1,tD:1,inD:0}
26. $sum(X0,$sum(X1,X2)) = $sum($sum(X0,X1),X2) [theory axiom] {a:0,w:11,nSel:1,tD:1,inD:0}
27. $sum(X0,0) = X0 [theory axiom] {a:0,w:5,nSel:1,tD:1,inD:0}
29. $sum(X0,$uminus(X0)) = 0 [theory axiom] {a:0,w:6,nSel:1,tD:1,inD:0}
30. ~$less(X0,X0) [theory axiom] {a:0,w:3,nSel:1,tD:1,inD:0}
31. ~$less(X1,X2) | ~$less(X0,X1) | $less(X0,X2) [theory axiom] {a:0,w:9,nSel:1,tD:1,inD:0}
34. $less(X1,$sum(X0,1)) | $less(X0,X1) [theory axiom] {a:0,w:8,nSel:1,tD:1,inD:0}
40. $less(2,1) [flattening 20]
41. ! [X0 : $int,X1 : $int,X2 : $int] : $sum($sum(X0,X1),X2) = $sum($sum(X0,X2),X1) [rectify 14]
75. $less(2,1) [cnf transformation 40] {a:0,w:3,nSel:1,goal:1,inD:0}
79. $sum($sum(X0,X1),X2) = $sum($sum(X0,X2),X1) [cnf transformation 41] {a:0,w:11,inD:0}
118. $sum(X0,$sum(X2,X1)) = $sum($sum(X0,X1),X2) [forward demodulation 79,26] {a:0,w:11,nSel:1,inD:0}
173. $sum(0,X0) = X0 [superposition 25,27] {a:1,w:5,nSel:1,tD:1,inD:0}
192. $less(X0,$sum(X0,1)) [resolution 34,30] {a:1,w:5,nSel:1,tD:1,inD:0}
198. $less(X0,$sum(1,X0)) [superposition 192,25] {a:2,w:5,nSel:1,tD:1,inD:0}
234. ~$less(X12,2) | $less(X12,1) [resolution 31,75] {a:1,w:6,nSel:1,goal:1,inD:0}
359. $sum(0,X3) = $sum(X2,$sum(X3,$uminus(X2))) [superposition 118,29] {a:1,w:10,inD:0}
380. $sum(X2,$sum(X3,$uminus(X2))) = X3 [forward demodulation 359,173] {a:1,w:8,nSel:1,inD:0}
505. $less($sum(X16,$uminus(1)),X16) [superposition 198,380] {a:3,w:6,inD:0}
510. $less($sum(X16,-1),X16) [evaluation 505] {a:3,w:5,nSel:1,inD:0}
524. $less($sum(2,-1),1) [resolution 510,234] {a:4,w:5,goal:1,inD:0}
536. $less(1,1) [evaluation 524] {a:4,w:3,goal:1,inD:0}
545. $false [subsumption resolution 536,30] {a:4,w:0,goal:1,inD:0}
% SZS output end Proof for test-user-lemma-1
% ------------------------------
% Version: Vampire 4.3.0 (commit 189fceff on 2019-04-16 16:54:47 +0100)
% Termination reason: Refutation

% Memory used [KB]: 895
% Time elapsed: 0.044 s
% ------------------------------
% ------------------------------
