1. ! [X0 : 'Nat()'] : 'Sub'(X0,s(X0)) [input]
2. ! [X0 : 'Nat()',X1 : 'Nat()'] : ('Sub'(X0,X1) => 'Sub'(X0,s(X1))) [input]
3. (i(l8(zero)) = 0 & ! [X7 : $int] : a(l7,X7) = a(l8(zero),X7)) & (! [X5 : 'Nat()'] : (i(l8(X5)) = i(l10(X5)) & ! [X6 : $int] : a(l8(X5),X6) = a(l10(X5),X6)) & ! [X2 : 'Nat()'] : ((v = a(l11(X2),i(l10(X2))) & ! [X4 : $int] : (~i(l10(X2)) = X4 => a(l10(X2),X4) = a(l11(X2),X4)) & i(l10(X2)) = i(l11(X2))) & (i(l8(s(X2))) = $sum(i(l11(X2)),1) & ! [X3 : $int] : a(l11(X2),X3) = a(l8(s(X2)),X3))) & ! [X1 : 'Nat()'] : ('Sub'(X1,nl8) => $less(i(l8(X1)),alength)) & ~$less(i(l8(nl8)),alength) & (i(main_end) = i(l8(nl8)) & ! [X0 : $int] : a(main_end,X0) = a(l8(nl8),X0))) [input]
4. ! [X1 : 'Nat()',X0 : $int] : (($lesseq(i(l8(zero)),X0) & $less(X0,i(l8(nl8))) & i(l8(s(X1))) = $sum(i(l8(X1)),1)) => ? [X2 : 'Nat()'] : (i(l8(X2)) = X0 & 'Sub'(X2,nl8))) [input]
5. ! [X0 : $int] : ! [X1 : 'Nat()'] : (('Sub'(X1,nl8) & ! [X2 : 'Nat()'] : (~X1 = X2 => a(l8(X2),X0) = a(l8(s(X2)),X0))) => a(l8(nl8),X0) = a(l8(s(X1)),X0)) [input]
6. ! [X0 : 'Nat()',X1 : 'Nat()'] : (i(l8(X0)) = i(l8(X1)) => X0 = X1) [input]
7. ~! [X0 : $int] : (($lesseq(0,X0) & $less(X0,alength)) => v = a(main_end,X0)) [input]
13. ~$less(X0,X0) (TD) [theory axiom]
14. ~$less(X0,X1) | ~$less(X1,X2) | $less(X0,X2) (TD) [theory axiom]
15. $less(X0,X1) | $less(X1,X0) | X0 = X1 (TD) [theory axiom]
20. ! [X0 : 'Nat()'] : (zero = X0 | s(p(X0)) = X0) [term algebras exhaustiveness]
21. zero != s(X0) (TD) [term algebras distinctness]
22. s(X0) != s(X1) | X0 = X1 (TD) [term algebras injectivity]
23. ~! [X0 : $int] : ((~$less(X0,0) & $less(X0,alength)) => v = a(main_end,X0)) [evaluation 7]
24. ! [X1 : 'Nat()',X0 : $int] : ((~$less(X0,i(l8(zero))) & $less(X0,i(l8(nl8))) & i(l8(s(X1))) = $sum(i(l8(X1)),1)) => ? [X2 : 'Nat()'] : (i(l8(X2)) = X0 & 'Sub'(X2,nl8))) [evaluation 4]
25. ! [X0 : $int,X1 : 'Nat()'] : (('Sub'(X1,nl8) & ! [X2 : 'Nat()'] : (X1 != X2 => a(l8(X2),X0) = a(l8(s(X2)),X0))) => a(l8(nl8),X0) = a(l8(s(X1)),X0)) [flattening 5]
26. ! [X0 : 'Nat()',X1 : $int] : ((~$less(X1,i(l8(zero))) & $less(X1,i(l8(nl8))) & i(l8(s(X0))) = $sum(i(l8(X0)),1)) => ? [X2 : 'Nat()'] : (i(l8(X2)) = X1 & 'Sub'(X2,nl8))) [rectify 24]
27. (i(l8(zero)) = 0 & ! [X0 : $int] : a(l7,X0) = a(l8(zero),X0)) & (! [X1 : 'Nat()'] : (i(l8(X1)) = i(l10(X1)) & ! [X2 : $int] : a(l8(X1),X2) = a(l10(X1),X2)) & ! [X3 : 'Nat()'] : ((v = a(l11(X3),i(l10(X3))) & ! [X4 : $int] : (~i(l10(X3)) = X4 => a(l10(X3),X4) = a(l11(X3),X4)) & i(l10(X3)) = i(l11(X3))) & (i(l8(s(X3))) = $sum(i(l11(X3)),1) & ! [X5 : $int] : a(l11(X3),X5) = a(l8(s(X3)),X5))) & ! [X6 : 'Nat()'] : ('Sub'(X6,nl8) => $less(i(l8(X6)),alength)) & ~$less(i(l8(nl8)),alength) & (i(main_end) = i(l8(nl8)) & ! [X7 : $int] : a(main_end,X7) = a(l8(nl8),X7))) [rectify 3]
28. i(l8(zero)) = 0 & ! [X0 : $int] : a(l7,X0) = a(l8(zero),X0) & ! [X1 : 'Nat()'] : (i(l8(X1)) = i(l10(X1)) & ! [X2 : $int] : a(l8(X1),X2) = a(l10(X1),X2)) & ! [X3 : 'Nat()'] : (v = a(l11(X3),i(l10(X3))) & ! [X4 : $int] : (i(l10(X3)) != X4 => a(l10(X3),X4) = a(l11(X3),X4)) & i(l10(X3)) = i(l11(X3)) & i(l8(s(X3))) = $sum(i(l11(X3)),1) & ! [X5 : $int] : a(l11(X3),X5) = a(l8(s(X3)),X5)) & ! [X6 : 'Nat()'] : ('Sub'(X6,nl8) => $less(i(l8(X6)),alength)) & ~$less(i(l8(nl8)),alength) & i(main_end) = i(l8(nl8)) & ! [X7 : $int] : a(main_end,X7) = a(l8(nl8),X7) [flattening 27]
29. ? [X0 : $int] : (v != a(main_end,X0) & (~$less(X0,0) & $less(X0,alength))) [ennf transformation 23]
30. ? [X0 : $int] : (v != a(main_end,X0) & ~$less(X0,0) & $less(X0,alength)) [flattening 29]
31. ! [X0 : 'Nat()',X1 : 'Nat()'] : (X0 = X1 | i(l8(X0)) != i(l8(X1))) [ennf transformation 6]
32. ! [X0 : $int,X1 : 'Nat()'] : (a(l8(nl8),X0) = a(l8(s(X1)),X0) | (~'Sub'(X1,nl8) | ? [X2 : 'Nat()'] : (a(l8(X2),X0) != a(l8(s(X2)),X0) & X1 != X2))) [ennf transformation 25]
33. ! [X0 : $int,X1 : 'Nat()'] : (a(l8(nl8),X0) = a(l8(s(X1)),X0) | ~'Sub'(X1,nl8) | ? [X2 : 'Nat()'] : (a(l8(X2),X0) != a(l8(s(X2)),X0) & X1 != X2)) [flattening 32]
34. ! [X0 : 'Nat()',X1 : $int] : (? [X2 : 'Nat()'] : (i(l8(X2)) = X1 & 'Sub'(X2,nl8)) | ($less(X1,i(l8(zero))) | ~$less(X1,i(l8(nl8))) | i(l8(s(X0))) != $sum(i(l8(X0)),1))) [ennf transformation 26]
35. ! [X0 : 'Nat()',X1 : $int] : (? [X2 : 'Nat()'] : (i(l8(X2)) = X1 & 'Sub'(X2,nl8)) | $less(X1,i(l8(zero))) | ~$less(X1,i(l8(nl8))) | i(l8(s(X0))) != $sum(i(l8(X0)),1)) [flattening 34]
36. i(l8(zero)) = 0 & ! [X0 : $int] : a(l7,X0) = a(l8(zero),X0) & ! [X1 : 'Nat()'] : (i(l8(X1)) = i(l10(X1)) & ! [X2 : $int] : a(l8(X1),X2) = a(l10(X1),X2)) & ! [X3 : 'Nat()'] : (v = a(l11(X3),i(l10(X3))) & ! [X4 : $int] : (a(l10(X3),X4) = a(l11(X3),X4) | i(l10(X3)) = X4) & i(l10(X3)) = i(l11(X3)) & i(l8(s(X3))) = $sum(i(l11(X3)),1) & ! [X5 : $int] : a(l11(X3),X5) = a(l8(s(X3)),X5)) & ! [X6 : 'Nat()'] : ($less(i(l8(X6)),alength) | ~'Sub'(X6,nl8)) & ~$less(i(l8(nl8)),alength) & i(main_end) = i(l8(nl8)) & ! [X7 : $int] : a(main_end,X7) = a(l8(nl8),X7) [ennf transformation 28]
37. ! [X0 : 'Nat()',X1 : 'Nat()'] : ('Sub'(X0,s(X1)) | ~'Sub'(X0,X1)) [ennf transformation 2]
38. ? [X0 : $int] : (v != a(main_end,X0) & ~$less(X0,0) & $less(X0,alength)) => (v != a(main_end,sK0) & ~$less(sK0,0) & $less(sK0,alength)) [choice axiom]
39. v != a(main_end,sK0) & ~$less(sK0,0) & $less(sK0,alength) [skolemisation 30,38]
40. ! [X1 : 'Nat()',X0 : $int] : (? [X2 : 'Nat()'] : (a(l8(X2),X0) != a(l8(s(X2)),X0) & X1 != X2) => (a(l8(sK1(X0,X1)),X0) != a(l8(s(sK1(X0,X1))),X0) & sK1(X0,X1) != X1)) [choice axiom]
41. ! [X0 : $int,X1 : 'Nat()'] : (a(l8(nl8),X0) = a(l8(s(X1)),X0) | ~'Sub'(X1,nl8) | (a(l8(sK1(X0,X1)),X0) != a(l8(s(sK1(X0,X1))),X0) & sK1(X0,X1) != X1)) [skolemisation 33,40]
42. ! [X1 : $int] : (? [X2 : 'Nat()'] : (i(l8(X2)) = X1 & 'Sub'(X2,nl8)) => (i(l8(sK2(X1))) = X1 & 'Sub'(sK2(X1),nl8))) [choice axiom]
43. ! [X0 : 'Nat()',X1 : $int] : ((i(l8(sK2(X1))) = X1 & 'Sub'(sK2(X1),nl8)) | $less(X1,i(l8(zero))) | ~$less(X1,i(l8(nl8))) | i(l8(s(X0))) != $sum(i(l8(X0)),1)) [skolemisation 35,42]
44. zero = X0 | s(p(X0)) = X0 [cnf transformation 20]
45. $less(sK0,alength) [cnf transformation 39]
46. ~$less(sK0,0) [cnf transformation 39]
47. v != a(main_end,sK0) [cnf transformation 39]
48. X0 = X1 | i(l8(X0)) != i(l8(X1)) [cnf transformation 31]
49. a(l8(nl8),X0) = a(l8(s(X1)),X0) | ~'Sub'(X1,nl8) | sK1(X0,X1) != X1 [cnf transformation 41]
50. a(l8(nl8),X0) = a(l8(s(X1)),X0) | ~'Sub'(X1,nl8) | a(l8(sK1(X0,X1)),X0) != a(l8(s(sK1(X0,X1))),X0) [cnf transformation 41]
51. 'Sub'(sK2(X1),nl8) | $less(X1,i(l8(zero))) | ~$less(X1,i(l8(nl8))) | i(l8(s(X0))) != $sum(i(l8(X0)),1) [cnf transformation 43]
52. i(l8(sK2(X1))) = X1 | $less(X1,i(l8(zero))) | ~$less(X1,i(l8(nl8))) | i(l8(s(X0))) != $sum(i(l8(X0)),1) [cnf transformation 43]
53. a(main_end,X7) = a(l8(nl8),X7) [cnf transformation 36]
54. i(main_end) = i(l8(nl8)) [cnf transformation 36]
55. ~$less(i(l8(nl8)),alength) [cnf transformation 36]
56. $less(i(l8(X6)),alength) | ~'Sub'(X6,nl8) [cnf transformation 36]
57. a(l11(X3),X5) = a(l8(s(X3)),X5) [cnf transformation 36]
58. i(l8(s(X3))) = $sum(i(l11(X3)),1) [cnf transformation 36]
59. i(l10(X3)) = i(l11(X3)) [cnf transformation 36]
60. a(l10(X3),X4) = a(l11(X3),X4) | i(l10(X3)) = X4 [cnf transformation 36]
61. v = a(l11(X3),i(l10(X3))) [cnf transformation 36]
62. a(l8(X1),X2) = a(l10(X1),X2) [cnf transformation 36]
63. i(l8(X1)) = i(l10(X1)) [cnf transformation 36]
64. a(l7,X0) = a(l8(zero),X0) [cnf transformation 36]
65. i(l8(zero)) = 0 [cnf transformation 36]
66. 'Sub'(X0,s(X1)) | ~'Sub'(X0,X1) [cnf transformation 37]
67. 'Sub'(X0,s(X0)) [cnf transformation 1]
[SA] new: 22. s(X0) != s(X1) | X0 = X1 (0:8) T [term algebras injectivity]
[SA] new: 21. zero != s(X0) (0:4) T [term algebras distinctness]
[SA] new: 44. zero = X0 | s(p(X0)) = X0 (0:8) [cnf transformation 20]
[SA] new: 15. $less(X0,X1) | $less(X1,X0) | X0 = X1 (0:9) T [theory axiom]
[SA] new: 14. ~$less(X0,X1) | ~$less(X1,X2) | $less(X0,X2) (0:9) T [theory axiom]
[SA] new: 13. ~$less(X0,X0) (0:3) T [theory axiom]
[SA] new: 47. v != a(main_end,sK0) (0:5) [cnf transformation 39]
[SA] new: 46. ~$less(sK0,0) (0:3) [cnf transformation 39]
[SA] new: 45. $less(sK0,alength) (0:3) [cnf transformation 39]
[SA] new: 48. X0 = X1 | i(l8(X0)) != i(l8(X1)) (0:10) [cnf transformation 31]
[SA] new: 50. a(l8(nl8),X0) = a(l8(s(X1)),X0) | ~'Sub'(X1,nl8) | a(l8(sK1(X0,X1)),X0) != a(l8(s(sK1(X0,X1))),X0) (0:27) [cnf transformation 41]
[SA] new: 49. a(l8(nl8),X0) = a(l8(s(X1)),X0) | ~'Sub'(X1,nl8) | sK1(X0,X1) != X1 (0:18) [cnf transformation 41]
[SA] new: 52. i(l8(sK2(X1))) = X1 | $less(X1,i(l8(zero))) | ~$less(X1,i(l8(nl8))) | i(l8(s(X0))) != $sum(i(l8(X0)),1) (0:26) [cnf transformation 43]
[SA] new: 51. 'Sub'(sK2(X1),nl8) | $less(X1,i(l8(zero))) | ~$less(X1,i(l8(nl8))) | i(l8(s(X0))) != $sum(i(l8(X0)),1) (0:24) [cnf transformation 43]
[SA] new: 65. i(l8(zero)) = 0 (0:5) [cnf transformation 36]
[SA] new: 64. a(l7,X0) = a(l8(zero),X0) (0:8) [cnf transformation 36]
[SA] new: 63. i(l8(X1)) = i(l10(X1)) (0:7) [cnf transformation 36]
[SA] new: 62. a(l8(X1),X2) = a(l10(X1),X2) (0:9) [cnf transformation 36]
[SA] new: 61. v = a(l11(X3),i(l10(X3))) (0:8) [cnf transformation 36]
[SA] new: 60. a(l10(X3),X4) = a(l11(X3),X4) | i(l10(X3)) = X4 (0:14) [cnf transformation 36]
[SA] new: 59. i(l10(X3)) = i(l11(X3)) (0:7) [cnf transformation 36]
[SA] new: 58. i(l8(s(X3))) = $sum(i(l11(X3)),1) (0:10) [cnf transformation 36]
[SA] new: 57. a(l11(X3),X5) = a(l8(s(X3)),X5) (0:10) [cnf transformation 36]
[SA] new: 56. $less(i(l8(X6)),alength) | ~'Sub'(X6,nl8) (0:8) [cnf transformation 36]
[SA] new: 55. ~$less(i(l8(nl8)),alength) (0:5) [cnf transformation 36]
[SA] new: 54. i(main_end) = i(l8(nl8)) (0:6) [cnf transformation 36]
[SA] new: 53. a(main_end,X7) = a(l8(nl8),X7) (0:8) [cnf transformation 36]
[SA] new: 66. 'Sub'(X0,s(X1)) | ~'Sub'(X0,X1) (0:7) [cnf transformation 37]
[SA] new: 67. 'Sub'(X0,s(X0)) (0:4) [cnf transformation 1]
[SA] passive: 22. s(X0) != s(X1) | X0 = X1 (0:8) T [term algebras injectivity]
[SA] passive: 21. zero != s(X0) (0:4) T [term algebras distinctness]
[SA] passive: 44. zero = X0 | s(p(X0)) = X0 (0:8) [cnf transformation 20]
[SA] passive: 15. $less(X0,X1) | $less(X1,X0) | X0 = X1 (0:9) T [theory axiom]
[SA] passive: 14. ~$less(X0,X1) | ~$less(X1,X2) | $less(X0,X2) (0:9) T [theory axiom]
[SA] passive: 13. ~$less(X0,X0) (0:3) T [theory axiom]
[SA] passive: 47. v != a(main_end,sK0) (0:5) [cnf transformation 39]
[SA] passive: 46. ~$less(sK0,0) (0:3) [cnf transformation 39]
[SA] passive: 45. $less(sK0,alength) (0:3) [cnf transformation 39]
[SA] passive: 48. X0 = X1 | i(l8(X0)) != i(l8(X1)) (0:10) [cnf transformation 31]
[SA] passive: 50. a(l8(nl8),X0) = a(l8(s(X1)),X0) | ~'Sub'(X1,nl8) | a(l8(sK1(X0,X1)),X0) != a(l8(s(sK1(X0,X1))),X0) (0:27) [cnf transformation 41]
[SA] passive: 49. a(l8(nl8),X0) = a(l8(s(X1)),X0) | ~'Sub'(X1,nl8) | sK1(X0,X1) != X1 (0:18) [cnf transformation 41]
[SA] passive: 52. i(l8(sK2(X1))) = X1 | $less(X1,i(l8(zero))) | ~$less(X1,i(l8(nl8))) | i(l8(s(X0))) != $sum(i(l8(X0)),1) (0:26) [cnf transformation 43]
[SA] passive: 51. 'Sub'(sK2(X1),nl8) | $less(X1,i(l8(zero))) | ~$less(X1,i(l8(nl8))) | i(l8(s(X0))) != $sum(i(l8(X0)),1) (0:24) [cnf transformation 43]
[SA] new: 68. $less(X1,0) | 'Sub'(sK2(X1),nl8) | ~$less(X1,i(l8(nl8))) | i(l8(s(X0))) != $sum(i(l8(X0)),1) (0:22) [backward demodulation 65,51]
[SA] new: 69. $less(X1,0) | i(l8(sK2(X1))) = X1 | ~$less(X1,i(l8(nl8))) | i(l8(s(X0))) != $sum(i(l8(X0)),1) (0:24) [backward demodulation 65,52]
[SA] passive: 65. i(l8(zero)) = 0 (0:5) [cnf transformation 36]
[SA] passive: 68. $less(X1,0) | 'Sub'(sK2(X1),nl8) | ~$less(X1,i(l8(nl8))) | i(l8(s(X0))) != $sum(i(l8(X0)),1) (0:22) [backward demodulation 65,51]
[SA] passive: 69. $less(X1,0) | i(l8(sK2(X1))) = X1 | ~$less(X1,i(l8(nl8))) | i(l8(s(X0))) != $sum(i(l8(X0)),1) (0:24) [backward demodulation 65,52]
[SA] passive: 64. a(l7,X0) = a(l8(zero),X0) (0:8) [cnf transformation 36]
[SA] passive: 63. i(l8(X1)) = i(l10(X1)) (0:7) [cnf transformation 36]
[SA] passive: 62. a(l8(X1),X2) = a(l10(X1),X2) (0:9) [cnf transformation 36]
[SA] new: 70. v = a(l11(X3),i(l8(X3))) (0:8) [forward demodulation 61,63]
[SA] passive: 70. v = a(l11(X3),i(l8(X3))) (0:8) [forward demodulation 61,63]
[SA] new: 71. a(l8(X3),X4) = a(l11(X3),X4) | i(l10(X3)) = X4 (0:14) [forward demodulation 60,62]
[SA] new: 72. i(l8(X3)) = X4 | a(l8(X3),X4) = a(l11(X3),X4) (0:14) [forward demodulation 71,63]
[SA] passive: 72. i(l8(X3)) = X4 | a(l8(X3),X4) = a(l11(X3),X4) (0:14) [forward demodulation 71,63]
[SA] new: 73. i(l8(X3)) = i(l11(X3)) (0:7) [forward demodulation 59,63]
[SA] passive: 73. i(l8(X3)) = i(l11(X3)) (0:7) [forward demodulation 59,63]
[SA] new: 74. i(l8(s(X3))) = $sum(i(l8(X3)),1) (0:10) [forward demodulation 58,73]
[SA] new: 75. i(l8(s(X0))) != i(l8(s(X0))) | $less(X1,0) | i(l8(sK2(X1))) = X1 | ~$less(X1,i(l8(nl8))) (0:23) [backward demodulation 74,69]
[SA] new: 76. i(l8(s(X0))) != i(l8(s(X0))) | $less(X1,0) | 'Sub'(sK2(X1),nl8) | ~$less(X1,i(l8(nl8))) (0:21) [backward demodulation 74,68]
[SA] passive: 74. i(l8(s(X3))) = $sum(i(l8(X3)),1) (0:10) [forward demodulation 58,73]
[SA] new: 77. $less(X1,0) | 'Sub'(sK2(X1),nl8) | ~$less(X1,i(l8(nl8))) (0:12) [trivial inequality removal 76]
[SA] new: 78. $less(X1,0) | i(l8(sK2(X1))) = X1 | ~$less(X1,i(l8(nl8))) (0:14) [trivial inequality removal 75]
[SA] passive: 78. $less(X1,0) | i(l8(sK2(X1))) = X1 | ~$less(X1,i(l8(nl8))) (0:14) [trivial inequality removal 75]
[SA] passive: 77. $less(X1,0) | 'Sub'(sK2(X1),nl8) | ~$less(X1,i(l8(nl8))) (0:12) [trivial inequality removal 76]
[SA] new: 79. a(l8(nl8),X0) = a(l11(X1),X0) | ~'Sub'(X1,nl8) | sK1(X0,X1) != X1 (0:17) [backward demodulation 57,49]
[SA] new: 80. a(l8(nl8),X0) = a(l11(X1),X0) | ~'Sub'(X1,nl8) | a(l8(sK1(X0,X1)),X0) != a(l8(s(sK1(X0,X1))),X0) (0:26) [backward demodulation 57,50]
[SA] passive: 57. a(l11(X3),X5) = a(l8(s(X3)),X5) (0:10) [cnf transformation 36]
[SA] passive: 79. a(l8(nl8),X0) = a(l11(X1),X0) | ~'Sub'(X1,nl8) | sK1(X0,X1) != X1 (0:17) [backward demodulation 57,49]
[SA] new: 81. a(l8(sK1(X0,X1)),X0) != a(l11(sK1(X0,X1)),X0) | a(l8(nl8),X0) = a(l11(X1),X0) | ~'Sub'(X1,nl8) (0:25) [forward demodulation 80,57]
[SA] passive: 81. a(l8(sK1(X0,X1)),X0) != a(l11(sK1(X0,X1)),X0) | a(l8(nl8),X0) = a(l11(X1),X0) | ~'Sub'(X1,nl8) (0:25) [forward demodulation 80,57]
[SA] passive: 56. $less(i(l8(X6)),alength) | ~'Sub'(X6,nl8) (0:8) [cnf transformation 36]
[SA] passive: 55. ~$less(i(l8(nl8)),alength) (0:5) [cnf transformation 36]
[SA] new: 82. ~$less(i(main_end),alength) (0:4) [backward demodulation 54,55]
[SA] new: 83. ~$less(X1,i(main_end)) | $less(X1,0) | 'Sub'(sK2(X1),nl8) (0:11) [backward demodulation 54,77]
[SA] new: 84. ~$less(X1,i(main_end)) | $less(X1,0) | i(l8(sK2(X1))) = X1 (0:13) [backward demodulation 54,78]
[SA] passive: 54. i(main_end) = i(l8(nl8)) (0:6) [cnf transformation 36]
[SA] passive: 82. ~$less(i(main_end),alength) (0:4) [backward demodulation 54,55]
[SA] passive: 83. ~$less(X1,i(main_end)) | $less(X1,0) | 'Sub'(sK2(X1),nl8) (0:11) [backward demodulation 54,77]
[SA] passive: 84. ~$less(X1,i(main_end)) | $less(X1,0) | i(l8(sK2(X1))) = X1 (0:13) [backward demodulation 54,78]
[SA] new: 85. a(main_end,X0) = a(l11(X1),X0) | a(l8(sK1(X0,X1)),X0) != a(l11(sK1(X0,X1)),X0) | ~'Sub'(X1,nl8) (0:24) [backward demodulation 53,81]
[SA] new: 86. a(main_end,X0) = a(l11(X1),X0) | ~'Sub'(X1,nl8) | sK1(X0,X1) != X1 (0:16) [backward demodulation 53,79]
[SA] passive: 53. a(main_end,X7) = a(l8(nl8),X7) (0:8) [cnf transformation 36]
[SA] passive: 85. a(main_end,X0) = a(l11(X1),X0) | a(l8(sK1(X0,X1)),X0) != a(l11(sK1(X0,X1)),X0) | ~'Sub'(X1,nl8) (0:24) [backward demodulation 53,81]
[SA] passive: 86. a(main_end,X0) = a(l11(X1),X0) | ~'Sub'(X1,nl8) | sK1(X0,X1) != X1 (0:16) [backward demodulation 53,79]
[SA] passive: 66. 'Sub'(X0,s(X1)) | ~'Sub'(X0,X1) (0:7) [cnf transformation 37]
[SA] passive: 67. 'Sub'(X0,s(X0)) (0:4) [cnf transformation 1]
[SA] active: 45. $less(sK0,alength) (0:3:1) [cnf transformation 39]
[SA] active: 46. ~$less(sK0,0) (0:3:1) [cnf transformation 39]
[SA] active: 13. ~$less(X0,X0) (0:3:1) T [theory axiom]
[SA] active: 67. 'Sub'(X0,s(X0)) (0:4:1) [cnf transformation 1]
[SA] active: 82. ~$less(i(main_end),alength) (0:4:1) [backward demodulation 54,55]
[SA] active: 21. zero != s(X0) (0:4:1) T [term algebras distinctness]
[SA] active: 47. v != a(main_end,sK0) (0:5:1) [cnf transformation 39]
[SA] active: 65. i(l8(zero)) = 0 (0:5:1) [cnf transformation 36]
[SA] active: 54. i(main_end) = i(l8(nl8)) (0:6:1) [cnf transformation 36]
[SA] active: 63. i(l8(X1)) = i(l10(X1)) (0:7:1) [cnf transformation 36]
[SA] active: 66. 'Sub'(X0,s(X1)) | ~'Sub'(X0,X1) (0:7:1) [cnf transformation 37]
[SA] active: 73. i(l8(X3)) = i(l11(X3)) (0:7:1) [forward demodulation 59,63]
[SA] active: 53. a(main_end,X7) = a(l8(nl8),X7) (0:8:1) [cnf transformation 36]
[SA] active: 56. ~'Sub'(X6,nl8) | $less(i(l8(X6)),alength) (0:8:1) [cnf transformation 36]
[SA] active: 64. a(l7,X0) = a(l8(zero),X0) (0:8:1) [cnf transformation 36]
[SA] active: 70. v = a(l11(X3),i(l8(X3))) (0:8:1) [forward demodulation 61,63]
[SA] new: 87. v = a(l11(zero),0) (1:6) [superposition 70,65]
[SA] new: 88. v = a(l11(nl8),i(main_end)) (1:7) [superposition 70,54]
[SA] passive: 87. v = a(l11(zero),0) (1:6) [superposition 70,65]
[SA] passive: 88. v = a(l11(nl8),i(main_end)) (1:7) [superposition 70,54]
[SA] active: 87. v = a(l11(zero),0) (1:6:1) [superposition 70,65]
[SA] active: 22. s(X0) != s(X1) | X0 = X1 (0:8:1) T [term algebras injectivity]
[SA] new: 89. X0 = X0 (1:3) T [equality resolution 22]
[SA] active: 88. v = a(l11(nl8),i(main_end)) (1:7:1) [superposition 70,54]
[SA] active: 44. s(p(X0)) = X0 | zero = X0 (0:8:1) [cnf transformation 20]
[SA] new: 90. s(X1) != X0 | p(X0) = X1 | zero = X0 (1:11) [superposition 22,44]
[SA] new: 91. s(X3) != X2 | p(X2) = X3 | zero = X2 (1:11) [superposition 22,44]
[SA] new: 92. 'Sub'(X5,X4) | ~'Sub'(X5,p(X4)) | zero = X4 (1:10) [superposition 66,44]
[SA] new: 93. zero != X6 | zero = X6 (1:6) [superposition 21,44]
[SA] new: 94. 'Sub'(p(X7),X7) | zero = X7 (1:7) [superposition 67,44]
[SA] passive: 90. s(X1) != X0 | p(X0) = X1 | zero = X0 (1:11) [superposition 22,44]
[SA] passive: 92. 'Sub'(X5,X4) | ~'Sub'(X5,p(X4)) | zero = X4 (1:10) [superposition 66,44]
[SA] passive: 94. 'Sub'(p(X7),X7) | zero = X7 (1:7) [superposition 67,44]
[SA] active: 94. 'Sub'(p(X7),X7) | zero = X7 (1:7:1) [superposition 67,44]
[SA] new: 95. zero = nl8 | $less(i(l8(p(nl8))),alength) (2:9) [resolution 94,56]
[SA] passive: 95. zero = nl8 | $less(i(l8(p(nl8))),alength) (2:9) [resolution 94,56]
[SA] active: 62. a(l8(X1),X2) = a(l10(X1),X2) (0:9:1) [cnf transformation 36]
[SA] active: 14. ~$less(X1,X2) | ~$less(X0,X1) | $less(X0,X2) (0:9:1) T [theory axiom]
[SA] new: 96. ~$less(X0,sK0) | $less(X0,alength) (1:6) [resolution 14,45]
[SA] passive: 96. ~$less(X0,sK0) | $less(X0,alength) (1:6) [resolution 14,45]
[SA] active: 15. $less(X1,X0) | $less(X0,X1) | X0 = X1 (0:9:2) T [theory axiom]
[SA] new: 97. $less(0,sK0) | 0 = sK0 (1:6) [resolution 15,46]
[SA] new: 98. $less(X0,X0) | X0 = X0 (1:6) T [resolution 15,13]
[SA] new: 99. $less(X1,X2) | X1 = X2 | ~$less(X3,X2) | $less(X3,X1) (1:12) T [resolution 15,14]
[SA] new: 100. $less(alength,i(main_end)) | alength = i(main_end) (1:8) [resolution 15,82]
[SA] new: 101. $less(0,sK0) | 0 = sK0 (1:6) [resolution 15,46]
[SA] new: 102. $less(X0,X0) | X0 = X0 (1:6) T [resolution 15,13]
[SA] new: 103. $less(X1,X2) | X1 = X2 | ~$less(X3,X2) | $less(X3,X1) (1:12) T [resolution 15,14]
[SA] new: 104. $less(alength,i(main_end)) | alength = i(main_end) (1:8) [resolution 15,82]
[SA] new: 105. $less(X0,X0) | X0 = X0 (1:6) T [factoring 15]
[SA] passive: 97. $less(0,sK0) | 0 = sK0 (1:6) [resolution 15,46]
[SA] passive: 99. $less(X1,X2) | X1 = X2 | ~$less(X3,X2) | $less(X3,X1) (1:12) T [resolution 15,14]
[SA] passive: 100. $less(alength,i(main_end)) | alength = i(main_end) (1:8) [resolution 15,82]
[SA] active: 96. ~$less(X0,sK0) | $less(X0,alength) (1:6:1) [resolution 14,45]
[SA] new: 106. $less(X0,alength) | $less(sK0,X0) | sK0 = X0 (2:9) [resolution 96,15]
[SA] new: 107. $less(X1,alength) | $less(sK0,X1) | sK0 = X1 (2:9) [resolution 96,15]
[SA] passive: 106. $less(X0,alength) | $less(sK0,X0) | sK0 = X0 (2:9) [resolution 96,15]
[SA] active: 48. i(l8(X0)) != i(l8(X1)) | X0 = X1 (0:10:1) [cnf transformation 31]
[SA] new: 108. i(l8(X0)) != 0 | zero = X0 (1:8) [superposition 48,65]
[SA] new: 109. i(main_end) != i(l8(X1)) | nl8 = X1 (1:9) [superposition 48,54]
[SA] new: 110. i(l8(X0)) != 0 | zero = X0 (1:8) [superposition 48,65]
[SA] new: 111. i(main_end) != i(l8(X1)) | nl8 = X1 (1:9) [superposition 48,54]
[SA] new: 112. X0 = X0 (1:3) [equality resolution 48]
[SA] passive: 108. i(l8(X0)) != 0 | zero = X0 (1:8) [superposition 48,65]
[SA] passive: 109. i(main_end) != i(l8(X1)) | nl8 = X1 (1:9) [superposition 48,54]
[SA] active: 97. $less(0,sK0) | 0 = sK0 (1:6:1) [resolution 15,46]
[SA] new: 113. 0 = sK0 | $less(0,alength) (2:6) [resolution 97,96]
[SA] new: 114. 0 = sK0 | ~$less(X0,0) | $less(X0,sK0) (2:9) [resolution 97,14]
[SA] passive: 113. 0 = sK0 | $less(0,alength) (2:6) [resolution 97,96]
[SA] passive: 114. 0 = sK0 | ~$less(X0,0) | $less(X0,sK0) (2:9) [resolution 97,14]
[SA] active: 57. a(l11(X3),X5) = a(l8(s(X3)),X5) (0:10:1) [cnf transformation 36]
[SA] new: 115. a(l8(X0),X1) = a(l11(p(X0)),X1) | zero = X0 (1:13) [superposition 57,44]
[SA] passive: 115. a(l8(X0),X1) = a(l11(p(X0)),X1) | zero = X0 (1:13) [superposition 57,44]
[SA] active: 113. $less(0,alength) | 0 = sK0 (2:6:1) [resolution 97,96]
[SA] new: 116. 0 = sK0 | ~$less(X0,0) | $less(X0,alength) (3:9) [resolution 113,14]
[SA] passive: 116. 0 = sK0 | ~$less(X0,0) | $less(X0,alength) (3:9) [resolution 113,14]
[SA] active: 74. i(l8(s(X3))) = $sum(i(l8(X3)),1) (0:10:1) [forward demodulation 58,73]
[SA] new: 117. i(l8(s(zero))) = $sum(0,1) (1:8) [superposition 74,65]
[SA] new: 118. i(l8(s(nl8))) = $sum(i(main_end),1) (1:9) [superposition 74,54]
[SA] new: 119. i(l8(s(zero))) = 1 (1:6) [evaluation 117]
[SA] passive: 119. i(l8(s(zero))) = 1 (1:6) [evaluation 117]
[SA] passive: 118. i(l8(s(nl8))) = $sum(i(main_end),1) (1:9) [superposition 74,54]
[SA] active: 119. i(l8(s(zero))) = 1 (1:6:1) [evaluation 117]
[SA] new: 120. i(l8(s(s(zero)))) = $sum(1,1) (2:9) [superposition 74,119]
[SA] new: 121. i(l8(X0)) != 1 | s(zero) = X0 (2:9) [superposition 48,119]
[SA] new: 122. i(l8(X1)) != 1 | s(zero) = X1 (2:9) [superposition 48,119]
[SA] new: 123. v = a(l11(s(zero)),1) (2:7) [superposition 70,119]
[SA] new: 124. i(l8(s(s(zero)))) = 2 (2:7) [evaluation 120]
[SA] passive: 124. i(l8(s(s(zero)))) = 2 (2:7) [evaluation 120]
[SA] passive: 121. i(l8(X0)) != 1 | s(zero) = X0 (2:9) [superposition 48,119]
[SA] passive: 123. v = a(l11(s(zero)),1) (2:7) [superposition 70,119]
[SA] active: 83. ~$less(X1,i(main_end)) | $less(X1,0) | 'Sub'(sK2(X1),nl8) (0:11:1) [backward demodulation 54,77]
[SA] new: 125. $less(X0,0) | 'Sub'(sK2(X0),nl8) | $less(i(main_end),X0) | i(main_end) = X0 (1:15) [resolution 83,15]
[SA] new: 126. $less(X1,0) | 'Sub'(sK2(X1),nl8) | $less(i(main_end),X1) | i(main_end) = X1 (1:15) [resolution 83,15]
[SA] passive: 125. $less(X0,0) | 'Sub'(sK2(X0),nl8) | $less(i(main_end),X0) | i(main_end) = X0 (1:15) [resolution 83,15]
[SA] active: 123. v = a(l11(s(zero)),1) (2:7:1) [superposition 70,119]
[SA] active: 84. ~$less(X1,i(main_end)) | $less(X1,0) | i(l8(sK2(X1))) = X1 (0:13:1) [backward demodulation 54,78]
[SA] new: 127. $less(X0,0) | i(l8(sK2(X0))) = X0 | $less(i(main_end),X0) | i(main_end) = X0 (1:17) [resolution 84,15]
[SA] new: 128. $less(X1,0) | i(l8(sK2(X1))) = X1 | $less(i(main_end),X1) | i(main_end) = X1 (1:17) [resolution 84,15]
[SA] passive: 127. $less(X0,0) | i(l8(sK2(X0))) = X0 | $less(i(main_end),X0) | i(main_end) = X0 (1:17) [resolution 84,15]
[SA] active: 124. i(l8(s(s(zero)))) = 2 (2:7:1) [evaluation 120]
[SA] new: 129. i(l8(s(s(s(zero))))) = $sum(2,1) (3:10) [superposition 74,124]
[SA] new: 130. i(l8(X0)) != 2 | s(s(zero)) = X0 (3:10) [superposition 48,124]
[SA] new: 131. i(l8(X1)) != 2 | s(s(zero)) = X1 (3:10) [superposition 48,124]
[SA] new: 132. v = a(l11(s(s(zero))),2) (3:8) [superposition 70,124]
[SA] new: 133. i(l8(s(s(s(zero))))) = 3 (3:8) [evaluation 129]
[SA] passive: 133. i(l8(s(s(s(zero))))) = 3 (3:8) [evaluation 129]
[SA] passive: 130. i(l8(X0)) != 2 | s(s(zero)) = X0 (3:10) [superposition 48,124]
[SA] passive: 132. v = a(l11(s(s(zero))),2) (3:8) [superposition 70,124]
[SA] active: 72. a(l8(X3),X4) = a(l11(X3),X4) | i(l8(X3)) = X4 (0:14:1) [forward demodulation 71,63]
[SA] new: 135. v = a(l8(nl8),i(main_end)) | i(main_end) = i(l8(nl8)) (2:13) [superposition 72,88]
[SA] new: 136. v = a(l8(s(zero)),1) | i(l8(s(zero))) = 1 (3:13) [superposition 72,123]
[SA] new: 137. v = a(l8(zero),0) | i(l8(zero)) = 0 (2:11) [superposition 72,87]
[SA] new: 139. v = a(l8(nl8),i(main_end)) | i(main_end) = i(l8(nl8)) (2:13) [superposition 88,72]
[SA] new: 140. v = a(l8(s(zero)),1) | i(l8(s(zero))) = 1 (3:13) [superposition 123,72]
[SA] new: 141. v = a(l8(zero),0) | i(l8(zero)) = 0 (2:11) [superposition 87,72]
[SA] active: 100. $less(alength,i(main_end)) | alength = i(main_end) (1:8:1) [resolution 15,82]
[SA] new: 142. alength = i(main_end) | $less(alength,0) | alength = i(l8(sK2(alength))) (2:13) [resolution 100,84]
[SA] new: 143. alength = i(main_end) | $less(alength,0) | 'Sub'(sK2(alength),nl8) (2:11) [resolution 100,83]
[SA] new: 144. alength = i(main_end) | ~$less(X0,alength) | $less(X0,i(main_end)) (2:11) [resolution 100,14]
[SA] passive: 142. alength = i(main_end) | $less(alength,0) | alength = i(l8(sK2(alength))) (2:13) [resolution 100,84]
[SA] passive: 143. alength = i(main_end) | $less(alength,0) | 'Sub'(sK2(alength),nl8) (2:11) [resolution 100,83]
[SA] passive: 144. alength = i(main_end) | ~$less(X0,alength) | $less(X0,i(main_end)) (2:11) [resolution 100,14]
[SA] active: 86. sK1(X0,X1) != X1 | ~'Sub'(X1,nl8) | a(main_end,X0) = a(l11(X1),X0) (0:16:1) [backward demodulation 53,79]
[SA] active: 108. i(l8(X0)) != 0 | zero = X0 (1:8:1) [superposition 48,65]
[SA] new: 146. i(main_end) != 0 | zero = nl8 (2:7) [superposition 108,54]
[SA] new: 147. 1 != 0 | zero = s(zero) (2:7) [superposition 108,119]
[SA] new: 148. 0 != 2 | zero = s(s(zero)) (3:8) [superposition 108,124]
[SA] passive: 146. i(main_end) != 0 | zero = nl8 (2:7) [superposition 108,54]
[SA] active: 85. a(l8(sK1(X0,X1)),X0) != a(l11(sK1(X0,X1)),X0) | a(main_end,X0) = a(l11(X1),X0) | ~'Sub'(X1,nl8) (0:24:1) [backward demodulation 53,81]
[SA] new: 149. a(l8(sK1(X0,X1)),X0) != a(l8(sK1(X0,X1)),X0) | a(main_end,X0) = a(l11(X1),X0) | ~'Sub'(X1,nl8) | i(l8(sK1(X0,X1))) = X0 (1:31) [superposition 85,72]
[SA] new: 150. a(main_end,X0) = a(l11(X1),X0) | ~'Sub'(X1,nl8) | i(l8(sK1(X0,X1))) = X0 (1:18) [trivial inequality removal 149]
[SA] passive: 150. a(main_end,X0) = a(l11(X1),X0) | ~'Sub'(X1,nl8) | i(l8(sK1(X0,X1))) = X0 (1:18) [trivial inequality removal 149]
[SA] active: 146. i(main_end) != 0 | zero = nl8 (2:7:1) [superposition 108,54]
[SA] active: 109. i(main_end) != i(l8(X1)) | nl8 = X1 (1:9:1) [superposition 48,54]
[SA] new: 151. i(main_end) != 0 | zero = nl8 (2:7) [superposition 109,65]
[SA] new: 153. i(main_end) != 1 | s(zero) = nl8 (2:8) [superposition 109,119]
[SA] new: 154. i(main_end) != 2 | s(s(zero)) = nl8 (3:9) [superposition 109,124]
[SA] passive: 153. i(main_end) != 1 | s(zero) = nl8 (2:8) [superposition 109,119]
[SA] passive: 154. i(main_end) != 2 | s(s(zero)) = nl8 (3:9) [superposition 109,124]
[SA] active: 153. i(main_end) != 1 | s(zero) = nl8 (2:8:1) [superposition 109,119]
[SA] active: 118. i(l8(s(nl8))) = $sum(i(main_end),1) (1:9:1) [superposition 74,54]
[SA] new: 155. i(main_end) != $sum(i(main_end),1) | s(nl8) = nl8 (2:11) [superposition 109,118]
[SA] new: 156. $sum(i(main_end),1) != 0 | zero = s(nl8) (2:10) [superposition 108,118]
[SA] new: 157. i(l8(s(s(nl8)))) = $sum($sum(i(main_end),1),1) (2:12) [superposition 74,118]
[SA] new: 158. i(l8(X0)) != $sum(i(main_end),1) | s(nl8) = X0 (2:12) [superposition 48,118]
[SA] new: 159. i(l8(X1)) != $sum(i(main_end),1) | s(nl8) = X1 (2:12) [superposition 48,118]
[SA] new: 160. v = a(l11(s(nl8)),$sum(i(main_end),1)) (2:10) [superposition 70,118]
[SA] new: 161. i(main_end) != $uminus(1) | zero = s(nl8) (2:9) [evaluation 156]
[SA] new: 162. i(main_end) != -1 | zero = s(nl8) (2:8) [evaluation 161]
[SA] new: 163. i(main_end) != -1 (2:4) [term algebras distinctness 162]
[SA] passive: 155. i(main_end) != $sum(i(main_end),1) | s(nl8) = nl8 (2:11) [superposition 109,118]
[SA] passive: 163. i(main_end) != -1 (2:4) [term algebras distinctness 162]
[SA] passive: 157. i(l8(s(s(nl8)))) = $sum($sum(i(main_end),1),1) (2:12) [superposition 74,118]
[SA] passive: 158. i(l8(X0)) != $sum(i(main_end),1) | s(nl8) = X0 (2:12) [superposition 48,118]
[SA] passive: 160. v = a(l11(s(nl8)),$sum(i(main_end),1)) (2:10) [superposition 70,118]
[SA] active: 163. i(main_end) != -1 (2:4:1) [term algebras distinctness 162]
[SA] active: 92. ~'Sub'(X5,p(X4)) | 'Sub'(X5,X4) | zero = X4 (1:10:1) [superposition 66,44]
[SA] new: 164. 'Sub'(p(p(X0)),X0) | zero = X0 | zero = p(X0) (2:12) [resolution 92,94]
[SA] passive: 164. 'Sub'(p(p(X0)),X0) | zero = X0 | zero = p(X0) (2:12) [resolution 92,94]
[SA] active: 132. v = a(l11(s(s(zero))),2) (3:8:1) [superposition 70,124]
[SA] new: 165. v = a(l8(s(s(zero))),2) | i(l8(s(s(zero)))) = 2 (4:15) [superposition 132,72]
[SA] new: 166. v = a(l8(s(s(zero))),2) | i(l8(s(s(zero)))) = 2 (4:15) [superposition 72,132]
[SA] active: 90. s(X1) != X0 | p(X0) = X1 | zero = X0 (1:11:1) [superposition 22,44]
[SA] new: 167. X0 != X1 | p(X0) = p(X1) | zero = X1 | zero = X0 (2:14) [superposition 90,44]
[SA] new: 168. p(s(X0)) = X0 | zero = s(X0) (2:9) [equality resolution 90]
[SA] new: 169. p(s(X0)) = X0 (2:5) [term algebras distinctness 168]
[SA] passive: 167. X0 != X1 | p(X0) = p(X1) | zero = X1 | zero = X0 (2:14) [superposition 90,44]
[SA] passive: 169. p(s(X0)) = X0 (2:5) [term algebras distinctness 168]
[SA] active: 169. p(s(X0)) = X0 (2:5:1) [term algebras distinctness 168]
[SA] new: 170. ~'Sub'(X1,X0) | 'Sub'(X1,s(X0)) | zero = s(X0) (3:11) [superposition 92,169]
[SA] new: 171. 'Sub'(X2,s(X2)) | zero = s(X2) (3:8) [superposition 94,169]
[SA] new: 172. 'Sub'(X2,s(X2)) (3:4) [term algebras distinctness 171]
[SA] new: 173. ~'Sub'(X1,X0) | 'Sub'(X1,s(X0)) (3:7) [term algebras distinctness 170]
[SA] active: 99. ~$less(X3,X2) | X1 = X2 | $less(X1,X2) | $less(X3,X1) (1:12:1) T [resolution 15,14]
[SA] new: 174. X0 = X1 | $less(X0,X1) | $less(X2,X0) | $less(X1,X2) | X1 = X2 (2:15) T [resolution 99,15]
[SA] new: 175. X3 = X4 | $less(X3,X4) | $less(X5,X3) | $less(X4,X5) | X4 = X5 (2:15) T [resolution 99,15]
[SA] new: 176. i(main_end) = X6 | $less(X6,i(main_end)) | $less(alength,X6) | alength = i(main_end) (2:15) [resolution 99,100]
[SA] new: 177. sK0 = X7 | $less(X7,sK0) | $less(0,X7) | 0 = sK0 (2:12) [resolution 99,97]
[SA] new: 178. alength = X8 | $less(X8,alength) | $less(0,X8) | 0 = sK0 (3:12) [resolution 99,113]
[SA] new: 179. alength = X9 | $less(X9,alength) | $less(sK0,X9) (2:9) [resolution 99,45]
[SA] passive: 174. X0 = X1 | $less(X0,X1) | $less(X2,X0) | $less(X1,X2) | X1 = X2 (2:15) T [resolution 99,15]
[SA] passive: 176. i(main_end) = X6 | $less(X6,i(main_end)) | $less(alength,X6) | alength = i(main_end) (2:15) [resolution 99,100]
[SA] passive: 177. sK0 = X7 | $less(X7,sK0) | $less(0,X7) | 0 = sK0 (2:12) [resolution 99,97]
[SA] passive: 178. alength = X8 | $less(X8,alength) | $less(0,X8) | 0 = sK0 (3:12) [resolution 99,113]
[SA] passive: 179. alength = X9 | $less(X9,alength) | $less(sK0,X9) (2:9) [resolution 99,45]
[SA] active: 133. i(l8(s(s(s(zero))))) = 3 (3:8:1) [evaluation 129]
[SA] new: 180. i(main_end) != 3 | s(s(s(zero))) = nl8 (4:10) [superposition 109,133]
[SA] new: 181. 0 != 3 | zero = s(s(s(zero))) (4:9) [superposition 108,133]
[SA] new: 182. i(l8(s(s(s(s(zero)))))) = $sum(3,1) (4:11) [superposition 74,133]
[SA] new: 183. i(l8(X0)) != 3 | s(s(s(zero))) = X0 (4:11) [superposition 48,133]
[SA] new: 184. i(l8(X1)) != 3 | s(s(s(zero))) = X1 (4:11) [superposition 48,133]
[SA] new: 185. v = a(l11(s(s(s(zero)))),3) (4:9) [superposition 70,133]
[SA] new: 186. i(l8(s(s(s(s(zero)))))) = 4 (4:9) [evaluation 182]
[SA] passive: 180. i(main_end) != 3 | s(s(s(zero))) = nl8 (4:10) [superposition 109,133]
[SA] passive: 186. i(l8(s(s(s(s(zero)))))) = 4 (4:9) [evaluation 182]
[SA] passive: 183. i(l8(X0)) != 3 | s(s(s(zero))) = X0 (4:11) [superposition 48,133]
[SA] passive: 185. v = a(l11(s(s(s(zero)))),3) (4:9) [superposition 70,133]
[SA] active: 115. a(l8(X0),X1) = a(l11(p(X0)),X1) | zero = X0 (1:13:1) [superposition 57,44]
[SA] new: 187. a(l11(X0),X1) = a(l8(s(X0)),X1) | zero = s(X0) (3:14) [superposition 115,169]
[SA] new: 188. a(l8(X2),X3) = a(l8(p(X2)),X3) | zero = X2 | i(l8(p(X2))) = X3 (2:19) [superposition 115,72]
[SA] new: 189. v = a(l8(X4),i(l8(p(X4)))) | zero = X4 (2:12) [superposition 115,70]
[SA] new: 190. a(l8(X0),X1) = a(l8(p(X0)),X1) | i(l8(p(X0))) = X1 | zero = X0 (2:19) [superposition 72,115]
[SA] new: 191. v = a(l8(X2),i(l8(p(X2)))) | zero = X2 (2:12) [superposition 70,115]
[SA] new: 192. a(l11(X0),X1) = a(l8(s(X0)),X1) (3:10) [term algebras distinctness 187]
[SA] passive: 188. a(l8(X2),X3) = a(l8(p(X2)),X3) | zero = X2 | i(l8(p(X2))) = X3 (2:19) [superposition 115,72]
[SA] passive: 189. v = a(l8(X4),i(l8(p(X4)))) | zero = X4 (2:12) [superposition 115,70]
[SA] active: 95. $less(i(l8(p(nl8))),alength) | zero = nl8 (2:9:1) [resolution 94,56]
[SA] new: 193. zero = nl8 | alength = X0 | $less(X0,alength) | $less(i(l8(p(nl8))),X0) (3:15) [resolution 95,99]
[SA] new: 194. zero = nl8 | ~$less(X1,i(l8(p(nl8)))) | $less(X1,alength) (3:12) [resolution 95,14]
[SA] passive: 193. zero = nl8 | alength = X0 | $less(X0,alength) | $less(i(l8(p(nl8))),X0) (3:15) [resolution 95,99]
[SA] passive: 194. zero = nl8 | ~$less(X1,i(l8(p(nl8)))) | $less(X1,alength) (3:12) [resolution 95,14]
[SA] active: 125. 'Sub'(sK2(X0),nl8) | $less(X0,0) | $less(i(main_end),X0) | i(main_end) = X0 (1:15:1) [resolution 83,15]
[SA] new: 195. $less(X0,0) | $less(i(main_end),X0) | i(main_end) = X0 | $less(i(l8(sK2(X0))),alength) (2:17) [resolution 125,56]
[SA] passive: 195. $less(X0,0) | $less(i(main_end),X0) | i(main_end) = X0 | $less(i(l8(sK2(X0))),alength) (2:17) [resolution 125,56]
[SA] active: 106. $less(sK0,X0) | $less(X0,alength) | sK0 = X0 (2:9:2) [resolution 96,15]
[SA] new: 196. $less(0,alength) | 0 = sK0 (3:6) [resolution 106,46]
[SA] new: 197. $less(sK0,alength) | sK0 = sK0 (3:6) [resolution 106,13]
[SA] new: 198. $less(X0,alength) | sK0 = X0 | X0 = X1 | $less(X1,X0) | $less(sK0,X1) (3:15) [resolution 106,99]
[SA] new: 199. $less(X2,alength) | sK0 = X2 | ~$less(X3,sK0) | $less(X3,X2) (3:12) [resolution 106,14]
[SA] new: 200. $less(i(main_end),alength) | i(main_end) = sK0 | $less(sK0,0) | i(l8(sK2(sK0))) = sK0 (3:17) [resolution 106,84]
[SA] new: 201. $less(i(main_end),alength) | i(main_end) = sK0 | $less(sK0,0) | 'Sub'(sK2(sK0),nl8) (3:15) [resolution 106,83]
[SA] new: 202. $less(sK0,alength) | sK0 = sK0 | $less(sK0,alength) (3:9) [resolution 106,96]
[SA] new: 203. $less(sK0,alength) | alength = sK0 (3:6) [resolution 106,13]
[SA] new: 204. $less(sK0,X0) | sK0 = X0 | alength = X1 | $less(X1,alength) | $less(X0,X1) (3:15) [resolution 106,99]
[SA] new: 205. $less(sK0,X2) | sK0 = X2 | ~$less(X3,X2) | $less(X3,alength) (3:12) [resolution 106,14]
[SA] new: 206. $less(sK0,i(main_end)) | i(main_end) = sK0 (3:8) [resolution 106,82]
[SA] new: 207. $less(sK0,alength) | sK0 = sK0 (3:6) [duplicate literal removal 202]
[SA] passive: 198. $less(X0,alength) | sK0 = X0 | X0 = X1 | $less(X1,X0) | $less(sK0,X1) (3:15) [resolution 106,99]
[SA] passive: 199. $less(X2,alength) | sK0 = X2 | ~$less(X3,sK0) | $less(X3,X2) (3:12) [resolution 106,14]
[SA] new: 208. i(main_end) = sK0 | $less(sK0,0) | i(l8(sK2(sK0))) = sK0 (3:13) [subsumption resolution 200,82]
[SA] new: 209. i(main_end) = sK0 | i(l8(sK2(sK0))) = sK0 (3:10) [subsumption resolution 208,46]
[SA] passive: 209. i(main_end) = sK0 | i(l8(sK2(sK0))) = sK0 (3:10) [subsumption resolution 208,46]
[SA] new: 210. i(main_end) = sK0 | $less(sK0,0) | 'Sub'(sK2(sK0),nl8) (3:11) [subsumption resolution 201,82]
[SA] new: 211. i(main_end) = sK0 | 'Sub'(sK2(sK0),nl8) (3:8) [subsumption resolution 210,46]
[SA] passive: 211. i(main_end) = sK0 | 'Sub'(sK2(sK0),nl8) (3:8) [subsumption resolution 210,46]
[SA] passive: 204. $less(sK0,X0) | sK0 = X0 | alength = X1 | $less(X1,alength) | $less(X0,X1) (3:15) [resolution 106,99]
[SA] passive: 205. $less(sK0,X2) | sK0 = X2 | ~$less(X3,X2) | $less(X3,alength) (3:12) [resolution 106,14]
[SA] passive: 206. $less(sK0,i(main_end)) | i(main_end) = sK0 (3:8) [resolution 106,82]
[SA] active: 127. $less(i(main_end),X0) | i(l8(sK2(X0))) = X0 | $less(X0,0) | i(main_end) = X0 (1:17:1) [resolution 84,15]
[SA] new: 212. alength = i(l8(sK2(alength))) | $less(alength,0) | alength = i(main_end) (2:13) [resolution 127,82]
[SA] new: 213. i(main_end) = i(l8(sK2(i(main_end)))) | $less(i(main_end),0) | i(main_end) = i(main_end) (2:17) [resolution 127,13]
[SA] new: 214. i(l8(sK2(X0))) = X0 | $less(X0,0) | i(main_end) = X0 | X0 = X1 | $less(X1,X0) | $less(i(main_end),X1) (2:23) [resolution 127,99]
[SA] new: 215. i(l8(sK2(X2))) = X2 | $less(X2,0) | i(main_end) = X2 | ~$less(X3,i(main_end)) | $less(X3,X2) (2:20) [resolution 127,14]
[SA] new: 216. i(main_end) = i(l8(sK2(i(main_end)))) | $less(i(main_end),0) | i(main_end) = i(main_end) | $less(i(main_end),0) | i(main_end) = i(l8(sK2(i(main_end)))) (2:29) [resolution 127,84]
[SA] new: 217. i(main_end) = i(l8(sK2(i(main_end)))) | $less(i(main_end),0) | i(main_end) = i(main_end) | $less(i(main_end),0) | 'Sub'(sK2(i(main_end)),nl8) (2:26) [resolution 127,83]
[SA] new: 218. i(l8(sK2(sK0))) = sK0 | $less(sK0,0) | i(main_end) = sK0 | $less(i(main_end),alength) (2:17) [resolution 127,96]
[SA] new: 219. i(main_end) = i(l8(sK2(i(main_end)))) | $less(i(main_end),0) | i(main_end) = i(main_end) | 'Sub'(sK2(i(main_end)),nl8) (2:22) [duplicate literal removal 217]
[SA] new: 220. i(main_end) = i(l8(sK2(i(main_end)))) | $less(i(main_end),0) | i(main_end) = i(main_end) (2:17) [duplicate literal removal 216]
[SA] passive: 214. i(l8(sK2(X0))) = X0 | $less(X0,0) | i(main_end) = X0 | X0 = X1 | $less(X1,X0) | $less(i(main_end),X1) (2:23) [resolution 127,99]
[SA] passive: 215. i(l8(sK2(X2))) = X2 | $less(X2,0) | i(main_end) = X2 | ~$less(X3,i(main_end)) | $less(X3,X2) (2:20) [resolution 127,14]
[SA] active: 206. $less(sK0,i(main_end)) | i(main_end) = sK0 (3:8:1) [resolution 106,82]
[SA] new: 221. i(main_end) = sK0 | $less(sK0,0) | i(l8(sK2(sK0))) = sK0 (4:13) [resolution 206,84]
[SA] new: 222. i(main_end) = sK0 | $less(sK0,0) | 'Sub'(sK2(sK0),nl8) (4:11) [resolution 206,83]
[SA] new: 223. i(main_end) = sK0 | i(main_end) = X0 | $less(X0,i(main_end)) | $less(sK0,X0) (4:15) [resolution 206,99]
[SA] new: 224. i(main_end) = sK0 | ~$less(X1,sK0) | $less(X1,i(main_end)) (4:11) [resolution 206,14]
[SA] passive: 223. i(main_end) = sK0 | i(main_end) = X0 | $less(X0,i(main_end)) | $less(sK0,X0) (4:15) [resolution 206,99]
[SA] passive: 224. i(main_end) = sK0 | ~$less(X1,sK0) | $less(X1,i(main_end)) (4:11) [resolution 206,14]
[SA] active: 150. ~'Sub'(X1,nl8) | a(main_end,X0) = a(l11(X1),X0) | i(l8(sK1(X0,X1))) = X0 (1:18:1) [trivial inequality removal 149]
[SA] new: 225. a(main_end,X0) = a(l11(sK2(X1)),X0) | i(l8(sK1(X0,sK2(X1)))) = X0 | $less(X1,0) | $less(i(main_end),X1) | i(main_end) = X1 (2:28) [resolution 150,125]
[SA] new: 226. a(main_end,X2) = a(l11(p(nl8)),X2) | i(l8(sK1(X2,p(nl8)))) = X2 | zero = nl8 (2:20) [resolution 150,94]
[SA] passive: 225. a(main_end,X0) = a(l11(sK2(X1)),X0) | i(l8(sK1(X0,sK2(X1)))) = X0 | $less(X1,0) | $less(i(main_end),X1) | i(main_end) = X1 (2:28) [resolution 150,125]
[SA] passive: 226. a(main_end,X2) = a(l11(p(nl8)),X2) | i(l8(sK1(X2,p(nl8)))) = X2 | zero = nl8 (2:20) [resolution 150,94]
[SA] active: 211. 'Sub'(sK2(sK0),nl8) | i(main_end) = sK0 (3:8:1) [subsumption resolution 210,46]
[SA] new: 227. i(main_end) = sK0 | a(main_end,X0) = a(l11(sK2(sK0)),X0) | i(l8(sK1(X0,sK2(sK0)))) = X0 (4:21) [resolution 211,150]
[SA] new: 228. i(main_end) = sK0 | $less(i(l8(sK2(sK0))),alength) (4:10) [resolution 211,56]
[SA] passive: 227. i(main_end) = sK0 | a(main_end,X0) = a(l11(sK2(sK0)),X0) | i(l8(sK1(X0,sK2(sK0)))) = X0 (4:21) [resolution 211,150]
[SA] passive: 228. i(main_end) = sK0 | $less(i(l8(sK2(sK0))),alength) (4:10) [resolution 211,56]
[SA] active: 114. ~$less(X0,0) | 0 = sK0 | $less(X0,sK0) (2:9:1) [resolution 97,14]
[SA] new: 229. 0 = sK0 | $less(X0,sK0) | $less(0,X0) | 0 = X0 (3:12) [resolution 114,15]
[SA] new: 230. 0 = sK0 | $less(X1,sK0) | $less(0,X1) | 0 = X1 (3:12) [resolution 114,15]
[SA] new: 231. 0 = sK0 | $less(i(main_end),sK0) | i(l8(sK2(0))) = 0 | $less(0,0) | i(main_end) = 0 (3:20) [resolution 114,127]
[SA] new: 232. 0 = sK0 | $less(sK0,sK0) | $less(0,alength) | 0 = sK0 (3:12) [resolution 114,106]
[SA] new: 233. 0 = sK0 | $less(sK0,sK0) | $less(0,alength) (3:9) [duplicate literal removal 232]
[SA] new: 234. 0 = sK0 | $less(i(main_end),sK0) | i(l8(sK2(0))) = 0 | i(main_end) = 0 (3:17) [evaluation 231]
[SA] passive: 229. 0 = sK0 | $less(X0,sK0) | $less(0,X0) | 0 = X0 (3:12) [resolution 114,15]
[SA] passive: 234. 0 = sK0 | $less(i(main_end),sK0) | i(l8(sK2(0))) = 0 | i(main_end) = 0 (3:17) [evaluation 231]
[SA] active: 121. i(l8(X0)) != 1 | s(zero) = X0 (2:9:1) [superposition 48,119]
[SA] new: 235. 1 != 0 | zero = s(zero) (3:7) [superposition 121,65]
[SA] new: 236. i(main_end) != 1 | s(zero) = nl8 (3:8) [superposition 121,54]
[SA] new: 238. 1 != 2 | s(zero) = s(s(zero)) (3:9) [superposition 121,124]
[SA] new: 239. 1 != 3 | s(zero) = s(s(s(zero))) (4:10) [superposition 121,133]
[SA] new: 240. 1 != $sum(i(main_end),1) | s(zero) = s(nl8) (3:11) [superposition 121,118]
[SA] new: 241. i(main_end) != $sum(1,-1) | s(zero) = s(nl8) (3:11) [evaluation 240]
[SA] new: 242. i(main_end) != 0 | s(zero) = s(nl8) (3:9) [evaluation 241]
[SA] new: 243. i(main_end) != 0 | zero = nl8 (3:7) [term algebras injectivity 242]
[SA] active: 179. $less(sK0,X9) | $less(X9,alength) | alength = X9 (2:9:2) [resolution 99,45]
[SA] new: 244. $less(0,alength) | alength = 0 (3:6) [resolution 179,46]
[SA] new: 245. $less(sK0,alength) | alength = sK0 (3:6) [resolution 179,13]
[SA] new: 246. $less(X0,alength) | alength = X0 | X0 = X1 | $less(X1,X0) | $less(sK0,X1) (3:15) [resolution 179,99]
[SA] new: 247. $less(X2,alength) | alength = X2 | ~$less(X3,sK0) | $less(X3,X2) (3:12) [resolution 179,14]
[SA] new: 248. $less(i(main_end),alength) | alength = i(main_end) | $less(sK0,0) | i(l8(sK2(sK0))) = sK0 (3:17) [resolution 179,84]
[SA] new: 249. $less(i(main_end),alength) | alength = i(main_end) | $less(sK0,0) | 'Sub'(sK2(sK0),nl8) (3:15) [resolution 179,83]
[SA] new: 250. $less(0,alength) | alength = 0 | 0 = sK0 | $less(sK0,sK0) (3:12) [resolution 179,114]
[SA] new: 251. $less(sK0,alength) | alength = sK0 | $less(sK0,alength) (3:9) [resolution 179,96]
[SA] new: 252. $less(sK0,alength) | alength = alength (3:6) [resolution 179,13]
[SA] new: 253. $less(sK0,X0) | alength = X0 | alength = X1 | $less(X1,alength) | $less(X0,X1) (3:15) [resolution 179,99]
[SA] new: 254. $less(sK0,X2) | alength = X2 | ~$less(X3,X2) | $less(X3,alength) (3:12) [resolution 179,14]
[SA] new: 255. $less(sK0,i(main_end)) | alength = i(main_end) (3:8) [resolution 179,82]
[SA] new: 256. $less(sK0,alength) | alength = sK0 (3:6) [duplicate literal removal 251]
[SA] passive: 244. $less(0,alength) | alength = 0 (3:6) [resolution 179,46]
[SA] passive: 246. $less(X0,alength) | alength = X0 | X0 = X1 | $less(X1,X0) | $less(sK0,X1) (3:15) [resolution 179,99]
[SA] passive: 247. $less(X2,alength) | alength = X2 | ~$less(X3,sK0) | $less(X3,X2) (3:12) [resolution 179,14]
[SA] new: 257. alength = i(main_end) | $less(sK0,0) | i(l8(sK2(sK0))) = sK0 (3:13) [subsumption resolution 248,82]
[SA] new: 258. alength = i(main_end) | i(l8(sK2(sK0))) = sK0 (3:10) [subsumption resolution 257,46]
[SA] passive: 258. alength = i(main_end) | i(l8(sK2(sK0))) = sK0 (3:10) [subsumption resolution 257,46]
[SA] new: 259. alength = i(main_end) | $less(sK0,0) | 'Sub'(sK2(sK0),nl8) (3:11) [subsumption resolution 249,82]
[SA] new: 260. alength = i(main_end) | 'Sub'(sK2(sK0),nl8) (3:8) [subsumption resolution 259,46]
[SA] passive: 260. alength = i(main_end) | 'Sub'(sK2(sK0),nl8) (3:8) [subsumption resolution 259,46]
[SA] passive: 253. $less(sK0,X0) | alength = X0 | alength = X1 | $less(X1,alength) | $less(X0,X1) (3:15) [resolution 179,99]
[SA] passive: 254. $less(sK0,X2) | alength = X2 | ~$less(X3,X2) | $less(X3,alength) (3:12) [resolution 179,14]
[SA] passive: 255. $less(sK0,i(main_end)) | alength = i(main_end) (3:8) [resolution 179,82]
[SA] active: 244. $less(0,alength) | alength = 0 (3:6:1) [resolution 179,46]
[SA] new: 261. alength = 0 | alength = X0 | $less(X0,alength) | $less(0,X0) (4:12) [resolution 244,99]
[SA] new: 262. alength = 0 | ~$less(X1,0) | $less(X1,alength) (4:9) [resolution 244,14]
[SA] passive: 261. alength = 0 | alength = X0 | $less(X0,alength) | $less(0,X0) (4:12) [resolution 244,99]
[SA] passive: 262. alength = 0 | ~$less(X1,0) | $less(X1,alength) (4:9) [resolution 244,14]
[SA] active: 160. v = a(l11(s(nl8)),$sum(i(main_end),1)) (2:10:1) [superposition 70,118]
[SA] new: 263. v = a(l8(s(nl8)),$sum(i(main_end),1)) | i(l8(s(nl8))) = $sum(i(main_end),1) (3:19) [superposition 160,72]
[SA] new: 264. v = a(l8(s(nl8)),$sum(i(main_end),1)) | i(l8(s(nl8))) = $sum(i(main_end),1) (3:19) [superposition 72,160]
[SA] active: 255. $less(sK0,i(main_end)) | alength = i(main_end) (3:8:1) [resolution 179,82]
[SA] new: 265. alength = i(main_end) | $less(sK0,0) | i(l8(sK2(sK0))) = sK0 (4:13) [resolution 255,84]
[SA] new: 266. alength = i(main_end) | $less(sK0,0) | 'Sub'(sK2(sK0),nl8) (4:11) [resolution 255,83]
[SA] new: 267. alength = i(main_end) | i(main_end) = X0 | $less(X0,i(main_end)) | $less(sK0,X0) (4:15) [resolution 255,99]
[SA] new: 268. alength = i(main_end) | ~$less(X1,sK0) | $less(X1,i(main_end)) (4:11) [resolution 255,14]
[SA] passive: 267. alength = i(main_end) | i(main_end) = X0 | $less(X0,i(main_end)) | $less(sK0,X0) (4:15) [resolution 255,99]
[SA] passive: 268. alength = i(main_end) | ~$less(X1,sK0) | $less(X1,i(main_end)) (4:11) [resolution 255,14]
[SA] active: 143. 'Sub'(sK2(alength),nl8) | $less(alength,0) | alength = i(main_end) (2:11:1) [resolution 100,83]
[SA] new: 269. $less(alength,0) | alength = i(main_end) | a(main_end,X0) = a(l11(sK2(alength)),X0) | i(l8(sK1(X0,sK2(alength)))) = X0 (3:24) [resolution 143,150]
[SA] new: 270. $less(alength,0) | alength = i(main_end) | $less(i(l8(sK2(alength))),alength) (3:13) [resolution 143,56]
[SA] passive: 269. $less(alength,0) | alength = i(main_end) | a(main_end,X0) = a(l11(sK2(alength)),X0) | i(l8(sK1(X0,sK2(alength)))) = X0 (3:24) [resolution 143,150]
[SA] passive: 270. $less(alength,0) | alength = i(main_end) | $less(i(l8(sK2(alength))),alength) (3:13) [resolution 143,56]
[SA] active: 260. 'Sub'(sK2(sK0),nl8) | alength = i(main_end) (3:8:1) [subsumption resolution 259,46]
[SA] new: 271. alength = i(main_end) | a(main_end,X0) = a(l11(sK2(sK0)),X0) | i(l8(sK1(X0,sK2(sK0)))) = X0 (4:21) [resolution 260,150]
[SA] new: 272. alength = i(main_end) | $less(i(l8(sK2(sK0))),alength) (4:10) [resolution 260,56]
[SA] passive: 271. alength = i(main_end) | a(main_end,X0) = a(l11(sK2(sK0)),X0) | i(l8(sK1(X0,sK2(sK0)))) = X0 (4:21) [resolution 260,150]
[SA] passive: 272. alength = i(main_end) | $less(i(l8(sK2(sK0))),alength) (4:10) [resolution 260,56]
[SA] active: 144. $less(X0,i(main_end)) | ~$less(X0,alength) | alength = i(main_end) (2:11:1) [resolution 100,14]
[SA] new: 273. ~$less(X0,alength) | alength = i(main_end) | $less(X0,0) | i(l8(sK2(X0))) = X0 (3:16) [resolution 144,84]
[SA] new: 274. ~$less(X1,alength) | alength = i(main_end) | $less(X1,0) | 'Sub'(sK2(X1),nl8) (3:14) [resolution 144,83]
[SA] new: 275. ~$less(i(main_end),alength) | alength = i(main_end) (3:8) [resolution 144,13]
[SA] new: 276. ~$less(X2,alength) | alength = i(main_end) | i(main_end) = X3 | $less(X3,i(main_end)) | $less(X2,X3) (3:18) [resolution 144,99]
[SA] new: 277. ~$less(X4,alength) | alength = i(main_end) | ~$less(X5,X4) | $less(X5,i(main_end)) (3:14) [resolution 144,14]
[SA] passive: 273. ~$less(X0,alength) | alength = i(main_end) | $less(X0,0) | i(l8(sK2(X0))) = X0 (3:16) [resolution 144,84]
[SA] passive: 274. ~$less(X1,alength) | alength = i(main_end) | $less(X1,0) | 'Sub'(sK2(X1),nl8) (3:14) [resolution 144,83]
[SA] passive: 276. ~$less(X2,alength) | alength = i(main_end) | i(main_end) = X3 | $less(X3,i(main_end)) | $less(X2,X3) (3:18) [resolution 144,99]
[SA] passive: 277. ~$less(X4,alength) | alength = i(main_end) | ~$less(X5,X4) | $less(X5,i(main_end)) (3:14) [resolution 144,14]
[SA] active: 116. ~$less(X0,0) | 0 = sK0 | $less(X0,alength) (3:9:1) [resolution 113,14]
[SA] new: 278. 0 = sK0 | $less(X0,alength) | $less(0,X0) | 0 = X0 (4:12) [resolution 116,15]
[SA] new: 279. 0 = sK0 | $less(X1,alength) | $less(0,X1) | 0 = X1 (4:12) [resolution 116,15]
[SA] new: 280. 0 = sK0 | $less(i(main_end),alength) | i(l8(sK2(0))) = 0 | $less(0,0) | i(main_end) = 0 (4:20) [resolution 116,127]
[SA] new: 281. 0 = sK0 | $less(sK0,alength) | $less(0,alength) | alength = 0 (4:12) [resolution 116,179]
[SA] new: 282. 0 = sK0 | $less(sK0,alength) | $less(0,alength) | 0 = sK0 (4:12) [resolution 116,106]
[SA] new: 283. 0 = sK0 | $less(sK0,alength) | $less(0,alength) (4:9) [duplicate literal removal 282]
[SA] new: 284. 0 = sK0 | $less(i(main_end),alength) | i(l8(sK2(0))) = 0 | i(main_end) = 0 (4:17) [evaluation 280]
[SA] passive: 278. 0 = sK0 | $less(X0,alength) | $less(0,X0) | 0 = X0 (4:12) [resolution 116,15]
[SA] new: 285. 0 = sK0 | i(l8(sK2(0))) = 0 | i(main_end) = 0 (4:13) [subsumption resolution 284,82]
[SA] passive: 285. 0 = sK0 | i(l8(sK2(0))) = 0 | i(main_end) = 0 (4:13) [subsumption resolution 284,82]
[SA] active: 155. i(main_end) != $sum(i(main_end),1) | s(nl8) = nl8 (2:11:1) [superposition 109,118]
[SA] active: 154. i(main_end) != 2 | s(s(zero)) = nl8 (3:9:1) [superposition 109,124]
[SA] active: 157. i(l8(s(s(nl8)))) = $sum($sum(i(main_end),1),1) (2:12:1) [superposition 74,118]
[SA] active: 185. v = a(l11(s(s(s(zero)))),3) (4:9:1) [superposition 70,133]
[SA] new: 286. v = a(l8(s(s(s(zero)))),3) | i(l8(s(s(s(zero))))) = 3 (5:17) [superposition 185,72]
[SA] new: 287. v = a(l8(s(s(s(zero)))),3) | i(l8(s(s(s(zero))))) = 3 (5:17) [superposition 72,185]
[SA] active: 158. i(l8(X0)) != $sum(i(main_end),1) | s(nl8) = X0 (2:12:1) [superposition 48,118]
[SA] new: 288. 1 != $sum(i(main_end),1) | s(zero) = s(nl8) (3:11) [superposition 158,119]
[SA] new: 289. $sum(i(main_end),1) != 2 | s(nl8) = s(s(zero)) (3:12) [superposition 158,124]
[SA] new: 290. $sum(i(main_end),1) != 3 | s(nl8) = s(s(s(zero))) (4:13) [superposition 158,133]
[SA] new: 292. i(main_end) != $sum(3,-1) | s(nl8) = s(s(s(zero))) (4:13) [evaluation 290]
[SA] new: 293. i(main_end) != 2 | s(nl8) = s(s(s(zero))) (4:11) [evaluation 292]
[SA] new: 294. i(main_end) != 2 | s(s(zero)) = nl8 (4:9) [term algebras injectivity 293]
[SA] new: 295. i(main_end) != $sum(2,-1) | s(nl8) = s(s(zero)) (3:12) [evaluation 289]
[SA] new: 296. i(main_end) != 1 | s(nl8) = s(s(zero)) (3:10) [evaluation 295]
[SA] new: 297. i(main_end) != 1 | s(zero) = nl8 (3:8) [term algebras injectivity 296]
[SA] new: 298. i(main_end) != $sum(1,-1) | s(zero) = s(nl8) (3:11) [evaluation 288]
[SA] new: 299. i(main_end) != 0 | s(zero) = s(nl8) (3:9) [evaluation 298]
[SA] new: 300. i(main_end) != 0 | zero = nl8 (3:7) [term algebras injectivity 299]
[SA] active: 186. i(l8(s(s(s(s(zero)))))) = 4 (4:9:1) [evaluation 182]
[SA] new: 301. i(l8(X0)) != 4 | s(s(s(s(zero)))) = X0 (5:12) [superposition 48,186]
[SA] new: 302. i(l8(X1)) != 4 | s(s(s(s(zero)))) = X1 (5:12) [superposition 48,186]
[SA] new: 303. v = a(l11(s(s(s(s(zero))))),4) (5:10) [superposition 70,186]
[SA] new: 304. i(l8(s(s(s(s(s(zero))))))) = $sum(4,1) (5:12) [superposition 74,186]
[SA] new: 305. 0 != 4 | zero = s(s(s(s(zero)))) (5:10) [superposition 108,186]
[SA] new: 306. i(main_end) != 4 | s(s(s(s(zero)))) = nl8 (5:11) [superposition 109,186]
[SA] new: 307. 1 != 4 | s(zero) = s(s(s(s(zero)))) (5:11) [superposition 121,186]
[SA] new: 308. $sum(i(main_end),1) != 4 | s(nl8) = s(s(s(s(zero)))) (5:14) [superposition 158,186]
[SA] new: 309. i(main_end) != $sum(4,-1) | s(nl8) = s(s(s(s(zero)))) (5:14) [evaluation 308]
[SA] new: 310. i(main_end) != 3 | s(nl8) = s(s(s(s(zero)))) (5:12) [evaluation 309]
[SA] new: 311. i(main_end) != 3 | s(s(s(zero))) = nl8 (5:10) [term algebras injectivity 310]
[SA] new: 312. i(l8(s(s(s(s(s(zero))))))) = 5 (5:10) [evaluation 304]
[SA] passive: 301. i(l8(X0)) != 4 | s(s(s(s(zero)))) = X0 (5:12) [superposition 48,186]
[SA] passive: 303. v = a(l11(s(s(s(s(zero))))),4) (5:10) [superposition 70,186]
[SA] passive: 312. i(l8(s(s(s(s(s(zero))))))) = 5 (5:10) [evaluation 304]
[SA] passive: 306. i(main_end) != 4 | s(s(s(s(zero)))) = nl8 (5:11) [superposition 109,186]
[SA] active: 164. 'Sub'(p(p(X0)),X0) | zero = X0 | zero = p(X0) (2:12:1) [resolution 92,94]
[SA] new: 313. zero = nl8 | zero = p(nl8) | a(main_end,X0) = a(l11(p(p(nl8))),X0) | i(l8(sK1(X0,p(p(nl8))))) = X0 (3:26) [resolution 164,150]
[SA] new: 314. zero = nl8 | zero = p(nl8) | $less(i(l8(p(p(nl8)))),alength) (3:14) [resolution 164,56]
[SA] new: 315. zero = p(X1) | zero = p(p(X1)) | 'Sub'(p(p(p(X1))),X1) | zero = X1 (3:18) [resolution 164,92]
[SA] new: 316. 'Sub'(p(X0),s(X0)) | zero = s(X0) | zero = p(s(X0)) (3:14) [superposition 164,169]
[SA] new: 317. 'Sub'(p(X0),s(X0)) | zero = p(s(X0)) (3:10) [term algebras distinctness 316]
[SA] passive: 313. zero = nl8 | zero = p(nl8) | a(main_end,X0) = a(l11(p(p(nl8))),X0) | i(l8(sK1(X0,p(p(nl8))))) = X0 (3:26) [resolution 164,150]
[SA] passive: 314. zero = nl8 | zero = p(nl8) | $less(i(l8(p(p(nl8)))),alength) (3:14) [resolution 164,56]
[SA] passive: 315. zero = p(X1) | zero = p(p(X1)) | 'Sub'(p(p(p(X1))),X1) | zero = X1 (3:18) [resolution 164,92]
[SA] new: 318. zero = X0 | 'Sub'(p(X0),s(X0)) (3:8) [forward demodulation 317,169]
[SA] passive: 318. zero = X0 | 'Sub'(p(X0),s(X0)) (3:8) [forward demodulation 317,169]
[SA] active: 318. 'Sub'(p(X0),s(X0)) | zero = X0 (3:8:1) [forward demodulation 317,169]
[SA] new: 319. 'Sub'(X0,s(s(X0))) | zero = s(X0) (4:9) [superposition 318,169]
[SA] new: 320. 'Sub'(p(p(X0)),X0) | zero = p(X0) | zero = X0 (4:12) [superposition 318,44]
[SA] new: 321. 'Sub'(X0,s(s(X0))) (4:5) [term algebras distinctness 319]
[SA] passive: 321. 'Sub'(X0,s(s(X0))) (4:5) [term algebras distinctness 319]
[SA] active: 177. $less(0,X7) | $less(X7,sK0) | sK0 = X7 | 0 = sK0 (2:12:2) [resolution 99,97]
[SA] new: 322. $less(0,sK0) | 0 = sK0 | 0 = sK0 (3:9) [resolution 177,13]
[SA] new: 323. $less(X0,sK0) | sK0 = X0 | 0 = sK0 | X0 = X1 | $less(X1,X0) | $less(0,X1) (3:18) [resolution 177,99]
[SA] new: 324. $less(X2,sK0) | sK0 = X2 | 0 = sK0 | ~$less(X3,0) | $less(X3,X2) (3:15) [resolution 177,14]
[SA] new: 325. $less(i(main_end),sK0) | i(main_end) = sK0 | 0 = sK0 | $less(0,0) | i(l8(sK2(0))) = 0 (3:20) [resolution 177,84]
[SA] new: 326. $less(i(main_end),sK0) | i(main_end) = sK0 | 0 = sK0 | $less(0,0) | 'Sub'(sK2(0),nl8) (3:18) [resolution 177,83]
[SA] new: 327. $less(0,sK0) | 0 = sK0 | 0 = sK0 | 0 = sK0 | $less(0,alength) (4:15) [resolution 177,116]
[SA] new: 328. $less(0,sK0) | 0 = sK0 | 0 = sK0 | 0 = sK0 | $less(0,sK0) (3:15) [resolution 177,114]
[SA] new: 329. $less(sK0,sK0) | sK0 = sK0 | 0 = sK0 | $less(0,alength) (3:12) [resolution 177,96]
[SA] new: 330. $less(0,X0) | sK0 = X0 | 0 = sK0 | $less(X0,alength) (3:12) [resolution 177,96]
[SA] new: 331. $less(0,sK0) | sK0 = sK0 | 0 = sK0 (3:9) [resolution 177,13]
[SA] new: 332. $less(0,X1) | sK0 = X1 | 0 = sK0 | sK0 = X2 | $less(X2,sK0) | $less(X1,X2) (3:18) [resolution 177,99]
[SA] new: 333. $less(0,X3) | sK0 = X3 | 0 = sK0 | ~$less(X4,X3) | $less(X4,sK0) (3:15) [resolution 177,14]
[SA] new: 334. $less(0,sK0) | 0 = sK0 (3:6) [duplicate literal removal 328]
[SA] new: 335. $less(0,sK0) | 0 = sK0 | $less(0,alength) (4:9) [duplicate literal removal 327]
[SA] new: 336. $less(i(main_end),sK0) | i(main_end) = sK0 | 0 = sK0 | 'Sub'(sK2(0),nl8) (3:15) [evaluation 326]
[SA] new: 337. $less(i(main_end),sK0) | i(main_end) = sK0 | 0 = sK0 | i(l8(sK2(0))) = 0 (3:17) [evaluation 325]
[SA] new: 338. $less(0,sK0) | 0 = sK0 (3:6) [duplicate literal removal 322]
[SA] passive: 323. $less(X0,sK0) | sK0 = X0 | 0 = sK0 | X0 = X1 | $less(X1,X0) | $less(0,X1) (3:18) [resolution 177,99]
[SA] passive: 324. $less(X2,sK0) | sK0 = X2 | 0 = sK0 | ~$less(X3,0) | $less(X3,X2) (3:15) [resolution 177,14]
[SA] passive: 337. $less(i(main_end),sK0) | i(main_end) = sK0 | 0 = sK0 | i(l8(sK2(0))) = 0 (3:17) [evaluation 325]
[SA] passive: 336. $less(i(main_end),sK0) | i(main_end) = sK0 | 0 = sK0 | 'Sub'(sK2(0),nl8) (3:15) [evaluation 326]
[SA] passive: 330. $less(0,X0) | sK0 = X0 | 0 = sK0 | $less(X0,alength) (3:12) [resolution 177,96]
[SA] passive: 332. $less(0,X1) | sK0 = X1 | 0 = sK0 | sK0 = X2 | $less(X2,sK0) | $less(X1,X2) (3:18) [resolution 177,99]
[SA] passive: 333. $less(0,X3) | sK0 = X3 | 0 = sK0 | ~$less(X4,X3) | $less(X4,sK0) (3:15) [resolution 177,14]
[SA] active: 321. 'Sub'(X0,s(s(X0))) (4:5:1) [term algebras distinctness 319]
[SA] new: 339. 'Sub'(p(X0),s(X0)) | zero = X0 (5:8) [superposition 321,44]
[SA] active: 189. v = a(l8(X4),i(l8(p(X4)))) | zero = X4 (2:12:1) [superposition 115,70]
[SA] new: 340. v = a(l8(s(X0)),i(l8(X0))) | zero = s(X0) (3:13) [superposition 189,169]
[SA] new: 341. v = a(main_end,i(l8(p(nl8)))) | zero = nl8 (3:11) [superposition 189,53]
[SA] new: 343. v = a(l11(X1),i(l8(p(s(X1))))) | zero = s(X1) (3:14) [superposition 189,57]
[SA] new: 344. v = a(main_end,i(l8(p(nl8)))) | zero = nl8 (3:11) [superposition 53,189]
[SA] new: 346. v = a(l11(X0),i(l8(p(s(X0))))) | zero = s(X0) (3:14) [superposition 57,189]
[SA] new: 347. v = a(l11(X0),i(l8(p(s(X0))))) (3:10) [term algebras distinctness 346]
[SA] new: 348. v = a(l11(X1),i(l8(p(s(X1))))) (3:10) [term algebras distinctness 343]
[SA] new: 349. v = a(l8(s(X0)),i(l8(X0))) (3:9) [term algebras distinctness 340]
[SA] passive: 349. v = a(l8(s(X0)),i(l8(X0))) (3:9) [term algebras distinctness 340]
[SA] passive: 341. v = a(main_end,i(l8(p(nl8)))) | zero = nl8 (3:11) [superposition 189,53]
[SA] new: 350. v = a(l11(X1),i(l8(X1))) (3:8) [forward demodulation 348,169]
[SA] new: 351. v = a(l11(X0),i(l8(X0))) (3:8) [forward demodulation 347,169]
[SA] active: 349. v = a(l8(s(X0)),i(l8(X0))) (3:9:1) [term algebras distinctness 340]
[SA] new: 352. v = a(l8(X0),i(l8(p(X0)))) | zero = X0 (4:12) [superposition 349,44]
[SA] new: 353. v = a(l8(s(zero)),0) (4:7) [superposition 349,65]
[SA] new: 354. v = a(l8(s(nl8)),i(main_end)) (4:8) [superposition 349,54]
[SA] new: 355. v = a(l8(s(s(zero))),1) (4:8) [superposition 349,119]
[SA] new: 356. v = a(l8(s(s(s(zero)))),2) (4:9) [superposition 349,124]
[SA] new: 357. v = a(l8(s(s(s(s(zero))))),3) (4:10) [superposition 349,133]
[SA] new: 358. v = a(l8(s(s(s(s(s(zero)))))),4) (5:11) [superposition 349,186]
[SA] new: 359. v = a(l8(s(s(nl8))),$sum(i(main_end),1)) (4:11) [superposition 349,118]
[SA] new: 360. v = a(l11(X1),i(l8(X1))) (4:8) [superposition 349,57]
[SA] new: 361. v = a(l11(X0),i(l8(X0))) (4:8) [superposition 57,349]
[SA] passive: 353. v = a(l8(s(zero)),0) (4:7) [superposition 349,65]
[SA] passive: 354. v = a(l8(s(nl8)),i(main_end)) (4:8) [superposition 349,54]
[SA] passive: 355. v = a(l8(s(s(zero))),1) (4:8) [superposition 349,119]
[SA] passive: 356. v = a(l8(s(s(s(zero)))),2) (4:9) [superposition 349,124]
[SA] passive: 357. v = a(l8(s(s(s(s(zero))))),3) (4:10) [superposition 349,133]
[SA] passive: 358. v = a(l8(s(s(s(s(s(zero)))))),4) (5:11) [superposition 349,186]
[SA] passive: 359. v = a(l8(s(s(nl8))),$sum(i(main_end),1)) (4:11) [superposition 349,118]
[SA] active: 142. $less(alength,0) | alength = i(main_end) | alength = i(l8(sK2(alength))) (2:13:1) [resolution 100,84]
[SA] new: 362. alength = i(main_end) | alength = i(l8(sK2(alength))) | 0 = sK0 | $less(alength,alength) (4:16) [resolution 142,116]
[SA] new: 363. alength = i(main_end) | alength = i(l8(sK2(alength))) | 0 = sK0 | $less(alength,sK0) (3:16) [resolution 142,114]
[SA] new: 364. alength = i(main_end) | alength = i(l8(sK2(alength))) | 0 = X0 | $less(X0,0) | $less(alength,X0) (3:19) [resolution 142,99]
[SA] new: 365. alength = i(main_end) | alength = i(l8(sK2(alength))) | ~$less(X1,alength) | $less(X1,0) (3:16) [resolution 142,14]
[SA] new: 366. alength = i(main_end) | alength = i(l8(sK2(alength))) | 0 = sK0 (4:13) [subsumption resolution 362,13]
[SA] passive: 366. alength = i(main_end) | alength = i(l8(sK2(alength))) | 0 = sK0 (4:13) [subsumption resolution 362,13]
[SA] passive: 364. alength = i(main_end) | alength = i(l8(sK2(alength))) | 0 = X0 | $less(X0,0) | $less(alength,X0) (3:19) [resolution 142,99]
[SA] passive: 365. alength = i(main_end) | alength = i(l8(sK2(alength))) | ~$less(X1,alength) | $less(X1,0) (3:16) [resolution 142,14]
[SA] active: 353. v = a(l8(s(zero)),0) (4:7:1) [superposition 349,65]
[SA] new: 367. v = a(l11(zero),0) (5:6) [superposition 353,57]
[SA] new: 368. v = a(l11(zero),0) (5:6) [superposition 57,353]
[SA] active: 167. X0 != X1 | p(X0) = p(X1) | zero = X1 | zero = X0 (2:14:1) [superposition 90,44]
[SA] new: 369. p(X0) = p(X0) | zero = X0 | zero = X0 (3:11) [equality resolution 167]
[SA] new: 370. p(X0) = p(X0) | zero = X0 (3:8) [duplicate literal removal 369]
[SA] active: 354. v = a(l8(s(nl8)),i(main_end)) (4:8:1) [superposition 349,54]
[SA] new: 371. v = a(l11(nl8),i(main_end)) (5:7) [superposition 354,57]
[SA] new: 372. v = a(l11(nl8),i(main_end)) (5:7) [superposition 57,354]
[SA] active: 176. $less(X6,i(main_end)) | i(main_end) = X6 | $less(alength,X6) | alength = i(main_end) (2:15:1) [resolution 99,100]
[SA] new: 373. i(main_end) = X0 | $less(alength,X0) | alength = i(main_end) | $less(X0,0) | i(l8(sK2(X0))) = X0 (3:20) [resolution 176,84]
[SA] new: 374. i(main_end) = X1 | $less(alength,X1) | alength = i(main_end) | $less(X1,0) | 'Sub'(sK2(X1),nl8) (3:18) [resolution 176,83]
[SA] new: 375. i(main_end) = i(main_end) | $less(alength,i(main_end)) | alength = i(main_end) (3:13) [resolution 176,13]
[SA] new: 376. i(main_end) = X2 | $less(alength,X2) | alength = i(main_end) | i(main_end) = X3 | $less(X3,i(main_end)) | $less(X2,X3) (3:22) [resolution 176,99]
[SA] new: 377. i(main_end) = X4 | $less(alength,X4) | alength = i(main_end) | ~$less(X5,X4) | $less(X5,i(main_end)) (3:18) [resolution 176,14]
[SA] passive: 373. i(main_end) = X0 | $less(alength,X0) | alength = i(main_end) | $less(X0,0) | i(l8(sK2(X0))) = X0 (3:20) [resolution 176,84]
[SA] passive: 374. i(main_end) = X1 | $less(alength,X1) | alength = i(main_end) | $less(X1,0) | 'Sub'(sK2(X1),nl8) (3:18) [resolution 176,83]
[SA] passive: 376. i(main_end) = X2 | $less(alength,X2) | alength = i(main_end) | i(main_end) = X3 | $less(X3,i(main_end)) | $less(X2,X3) (3:22) [resolution 176,99]
[SA] passive: 377. i(main_end) = X4 | $less(alength,X4) | alength = i(main_end) | ~$less(X5,X4) | $less(X5,i(main_end)) (3:18) [resolution 176,14]
[SA] active: 355. v = a(l8(s(s(zero))),1) (4:8:1) [superposition 349,119]
[SA] new: 378. v = a(l11(s(zero)),1) (5:7) [superposition 355,57]
[SA] new: 379. v = a(l11(s(zero)),1) (5:7) [superposition 57,355]
[SA] active: 174. $less(X2,X0) | $less(X1,X2) | $less(X0,X1) | X0 = X1 | X1 = X2 (2:15:3) T [resolution 99,15]
[SA] new: 380. $less(X0,sK0) | $less(0,X0) | 0 = X0 | sK0 = X0 (3:12) [resolution 174,46]
[SA] new: 381. $less(X1,X2) | $less(X2,X1) | X1 = X2 | X1 = X2 (3:12) T [resolution 174,13]
[SA] new: 382. $less(X3,X4) | $less(X5,X3) | X3 = X5 | X3 = X4 | X5 = X6 | $less(X6,X5) | $less(X4,X6) (3:21) T [resolution 174,99]
[SA] new: 383. $less(X7,X8) | $less(X9,X7) | X7 = X9 | X7 = X8 | ~$less(X10,X8) | $less(X10,X9) (3:18) T [resolution 174,14]
[SA] new: 384. $less(X11,X12) | $less(i(main_end),X11) | i(main_end) = X11 | X11 = X12 | $less(X12,0) | i(l8(sK2(X12))) = X12 (3:23) [resolution 174,84]
[SA] new: 385. $less(X13,X14) | $less(i(main_end),X13) | i(main_end) = X13 | X13 = X14 | $less(X14,0) | 'Sub'(sK2(X14),nl8) (3:21) [resolution 174,83]
[SA] new: 386. $less(X15,X16) | $less(0,X15) | 0 = X15 | X15 = X16 | 0 = sK0 | $less(X16,alength) (4:18) [resolution 174,116]
[SA] new: 387. $less(X17,X18) | $less(0,X17) | 0 = X17 | X17 = X18 | 0 = sK0 | $less(X18,sK0) (3:18) [resolution 174,114]
[SA] new: 388. $less(X19,X20) | $less(sK0,X19) | sK0 = X19 | X19 = X20 | $less(X20,alength) (3:15) [resolution 174,96]
[SA] new: 389. $less(X21,i(main_end)) | $less(alength,X21) | alength = X21 | i(main_end) = X21 (3:14) [resolution 174,82]
[SA] new: 390. $less(0,X0) | $less(X0,sK0) | sK0 = X0 | 0 = sK0 (3:12) [resolution 174,46]
[SA] new: 391. $less(X1,X2) | $less(X2,X1) | X1 = X2 | X1 = X1 (3:12) T [resolution 174,13]
[SA] new: 392. $less(X3,X4) | $less(X4,X5) | X4 = X5 | X3 = X5 | X3 = X6 | $less(X6,X3) | $less(X5,X6) (3:21) T [resolution 174,99]
[SA] new: 393. $less(X7,X8) | $less(X8,X9) | X8 = X9 | X7 = X9 | ~$less(X10,X9) | $less(X10,X7) (3:18) T [resolution 174,14]
[SA] new: 394. $less(i(main_end),X11) | $less(X11,X12) | X11 = X12 | i(main_end) = X12 | $less(X12,0) | i(l8(sK2(X12))) = X12 (3:23) [resolution 174,84]
[SA] new: 395. $less(i(main_end),X13) | $less(X13,X14) | X13 = X14 | i(main_end) = X14 | $less(X14,0) | 'Sub'(sK2(X14),nl8) (3:21) [resolution 174,83]
[SA] new: 396. $less(0,X15) | $less(X15,X16) | X15 = X16 | 0 = X16 | 0 = sK0 | $less(X16,alength) (4:18) [resolution 174,116]
[SA] new: 397. $less(0,X17) | $less(X17,X18) | X17 = X18 | 0 = X18 | 0 = sK0 | $less(X18,sK0) (3:18) [resolution 174,114]
[SA] new: 398. $less(sK0,X19) | $less(X19,X20) | X19 = X20 | sK0 = X20 | $less(X20,alength) (3:15) [resolution 174,96]
[SA] new: 399. $less(alength,X21) | $less(X21,i(main_end)) | i(main_end) = X21 | alength = i(main_end) (3:15) [resolution 174,82]
[SA] new: 400. $less(X0,sK0) | $less(0,X0) | 0 = sK0 | 0 = X0 (3:12) [resolution 174,46]
[SA] new: 401. $less(X1,X2) | $less(X2,X1) | X2 = X2 | X1 = X2 (3:12) T [resolution 174,13]
[SA] new: 402. $less(X3,X4) | $less(X5,X3) | X4 = X5 | X3 = X5 | X5 = X6 | $less(X6,X5) | $less(X4,X6) (3:21) T [resolution 174,99]
[SA] new: 403. $less(X7,X8) | $less(X9,X7) | X8 = X9 | X7 = X9 | ~$less(X10,X8) | $less(X10,X9) (3:18) T [resolution 174,14]
[SA] new: 404. $less(X11,X12) | $less(i(main_end),X11) | i(main_end) = X12 | i(main_end) = X11 | $less(X12,0) | i(l8(sK2(X12))) = X12 (3:24) [resolution 174,84]
[SA] new: 405. $less(X13,X14) | $less(i(main_end),X13) | i(main_end) = X14 | i(main_end) = X13 | $less(X14,0) | 'Sub'(sK2(X14),nl8) (3:22) [resolution 174,83]
[SA] new: 406. $less(X15,X16) | $less(0,X15) | 0 = X16 | 0 = X15 | 0 = sK0 | $less(X16,alength) (4:18) [resolution 174,116]
[SA] new: 407. $less(X17,X18) | $less(0,X17) | 0 = X18 | 0 = X17 | 0 = sK0 | $less(X18,sK0) (3:18) [resolution 174,114]
[SA] new: 408. $less(X19,X20) | $less(sK0,X19) | sK0 = X20 | sK0 = X19 | $less(X20,alength) (3:15) [resolution 174,96]
[SA] new: 409. $less(X21,i(main_end)) | $less(alength,X21) | alength = i(main_end) | alength = X21 (3:14) [resolution 174,82]
[SA] new: 410. $less(X0,X0) | $less(X0,X0) | X0 = X0 | X0 = X0 (3:12) T [factoring 174]
[SA] new: 411. $less(X1,X1) | $less(X1,X1) | X1 = X1 | X1 = X1 (3:12) T [factoring 174]
[SA] new: 412. $less(X2,X2) | $less(X2,X2) | X2 = X2 | X2 = X2 (3:12) T [factoring 174]
[SA] new: 413. $less(X2,X2) | X2 = X2 (3:6) T [duplicate literal removal 412]
[SA] new: 414. $less(X1,X1) | X1 = X1 (3:6) T [duplicate literal removal 411]
[SA] new: 415. $less(X0,X0) | X0 = X0 (3:6) T [duplicate literal removal 410]
[SA] new: 416. $less(X1,X2) | $less(X2,X1) | X1 = X2 (3:9) T [duplicate literal removal 381]
[SA] passive: 380. $less(X0,sK0) | $less(0,X0) | 0 = X0 | sK0 = X0 (3:12) [resolution 174,46]
[SA] passive: 382. $less(X3,X4) | $less(X5,X3) | X3 = X5 | X3 = X4 | X5 = X6 | $less(X6,X5) | $less(X4,X6) (3:21) T [resolution 174,99]
[SA] passive: 383. $less(X7,X8) | $less(X9,X7) | X7 = X9 | X7 = X8 | ~$less(X10,X8) | $less(X10,X9) (3:18) T [resolution 174,14]
[SA] passive: 384. $less(X11,X12) | $less(i(main_end),X11) | i(main_end) = X11 | X11 = X12 | $less(X12,0) | i(l8(sK2(X12))) = X12 (3:23) [resolution 174,84]
[SA] passive: 385. $less(X13,X14) | $less(i(main_end),X13) | i(main_end) = X13 | X13 = X14 | $less(X14,0) | 'Sub'(sK2(X14),nl8) (3:21) [resolution 174,83]
[SA] passive: 386. $less(X15,X16) | $less(0,X15) | 0 = X15 | X15 = X16 | 0 = sK0 | $less(X16,alength) (4:18) [resolution 174,116]
[SA] passive: 387. $less(X17,X18) | $less(0,X17) | 0 = X17 | X17 = X18 | 0 = sK0 | $less(X18,sK0) (3:18) [resolution 174,114]
[SA] passive: 388. $less(X19,X20) | $less(sK0,X19) | sK0 = X19 | X19 = X20 | $less(X20,alength) (3:15) [resolution 174,96]
[SA] passive: 389. $less(X21,i(main_end)) | $less(alength,X21) | alength = X21 | i(main_end) = X21 (3:14) [resolution 174,82]
[SA] passive: 392. $less(X3,X4) | $less(X4,X5) | X4 = X5 | X3 = X5 | X3 = X6 | $less(X6,X3) | $less(X5,X6) (3:21) T [resolution 174,99]
[SA] passive: 393. $less(X7,X8) | $less(X8,X9) | X8 = X9 | X7 = X9 | ~$less(X10,X9) | $less(X10,X7) (3:18) T [resolution 174,14]
[SA] passive: 395. $less(i(main_end),X13) | $less(X13,X14) | X13 = X14 | i(main_end) = X14 | $less(X14,0) | 'Sub'(sK2(X14),nl8) (3:21) [resolution 174,83]
[SA] passive: 396. $less(0,X15) | $less(X15,X16) | X15 = X16 | 0 = X16 | 0 = sK0 | $less(X16,alength) (4:18) [resolution 174,116]
[SA] passive: 397. $less(0,X17) | $less(X17,X18) | X17 = X18 | 0 = X18 | 0 = sK0 | $less(X18,sK0) (3:18) [resolution 174,114]
[SA] passive: 402. $less(X3,X4) | $less(X5,X3) | X4 = X5 | X3 = X5 | X5 = X6 | $less(X6,X5) | $less(X4,X6) (3:21) T [resolution 174,99]
[SA] passive: 403. $less(X7,X8) | $less(X9,X7) | X8 = X9 | X7 = X9 | ~$less(X10,X8) | $less(X10,X9) (3:18) T [resolution 174,14]
[SA] passive: 404. $less(X11,X12) | $less(i(main_end),X11) | i(main_end) = X12 | i(main_end) = X11 | $less(X12,0) | i(l8(sK2(X12))) = X12 (3:24) [resolution 174,84]
[SA] passive: 405. $less(X13,X14) | $less(i(main_end),X13) | i(main_end) = X14 | i(main_end) = X13 | $less(X14,0) | 'Sub'(sK2(X14),nl8) (3:22) [resolution 174,83]
[SA] passive: 406. $less(X15,X16) | $less(0,X15) | 0 = X16 | 0 = X15 | 0 = sK0 | $less(X16,alength) (4:18) [resolution 174,116]
[SA] passive: 407. $less(X17,X18) | $less(0,X17) | 0 = X18 | 0 = X17 | 0 = sK0 | $less(X18,sK0) (3:18) [resolution 174,114]
[SA] passive: 408. $less(X19,X20) | $less(sK0,X19) | sK0 = X20 | sK0 = X19 | $less(X20,alength) (3:15) [resolution 174,96]
[SA] passive: 409. $less(X21,i(main_end)) | $less(alength,X21) | alength = i(main_end) | alength = X21 (3:14) [resolution 174,82]
[SA] active: 262. ~$less(X1,0) | alength = 0 | $less(X1,alength) (4:9:1) [resolution 244,14]
[SA] new: 417. alength = 0 | $less(X0,alength) | $less(X1,X0) | $less(0,X1) | 0 = X0 | 0 = X1 (5:18) [resolution 262,174]
[SA] new: 418. alength = 0 | $less(X2,alength) | $less(0,X3) | $less(X3,X2) | X2 = X3 | 0 = X2 (5:18) [resolution 262,174]
[SA] new: 419. alength = 0 | $less(X4,alength) | $less(X5,X4) | $less(0,X5) | 0 = X5 | X4 = X5 (5:18) [resolution 262,174]
[SA] new: 420. alength = 0 | $less(X6,alength) | $less(0,X6) | 0 = X6 (5:12) [resolution 262,15]
[SA] new: 421. alength = 0 | $less(X7,alength) | $less(0,X7) | 0 = X7 (5:12) [resolution 262,15]
[SA] new: 422. alength = 0 | $less(alength,alength) | alength = i(main_end) | alength = i(l8(sK2(alength))) (5:16) [resolution 262,142]
[SA] new: 423. alength = 0 | $less(i(main_end),alength) | i(l8(sK2(0))) = 0 | $less(0,0) | i(main_end) = 0 (5:20) [resolution 262,127]
[SA] new: 424. alength = 0 | $less(0,alength) | $less(0,sK0) | 0 = sK0 | 0 = sK0 (5:15) [resolution 262,177]
[SA] new: 425. alength = 0 | $less(sK0,alength) | $less(0,alength) | alength = 0 (5:12) [resolution 262,179]
[SA] new: 426. alength = 0 | $less(sK0,alength) | $less(0,alength) | 0 = sK0 (5:12) [resolution 262,106]
[SA] new: 427. alength = 0 | $less(sK0,alength) | $less(0,alength) (5:9) [duplicate literal removal 425]
[SA] new: 428. alength = 0 | $less(0,alength) | $less(0,sK0) | 0 = sK0 (5:12) [duplicate literal removal 424]
[SA] new: 429. alength = 0 | $less(i(main_end),alength) | i(l8(sK2(0))) = 0 | i(main_end) = 0 (5:17) [evaluation 423]
[SA] passive: 417. alength = 0 | $less(X0,alength) | $less(X1,X0) | $less(0,X1) | 0 = X0 | 0 = X1 (5:18) [resolution 262,174]
[SA] passive: 418. alength = 0 | $less(X2,alength) | $less(0,X3) | $less(X3,X2) | X2 = X3 | 0 = X2 (5:18) [resolution 262,174]
[SA] passive: 419. alength = 0 | $less(X4,alength) | $less(X5,X4) | $less(0,X5) | 0 = X5 | X4 = X5 (5:18) [resolution 262,174]
[SA] passive: 420. alength = 0 | $less(X6,alength) | $less(0,X6) | 0 = X6 (5:12) [resolution 262,15]
[SA] new: 430. alength = 0 | alength = i(main_end) | alength = i(l8(sK2(alength))) (5:13) [subsumption resolution 422,13]
[SA] passive: 430. alength = 0 | alength = i(main_end) | alength = i(l8(sK2(alength))) (5:13) [subsumption resolution 422,13]
[SA] new: 431. alength = 0 | i(l8(sK2(0))) = 0 | i(main_end) = 0 (5:13) [subsumption resolution 429,82]
[SA] passive: 431. alength = 0 | i(l8(sK2(0))) = 0 | i(main_end) = 0 (5:13) [subsumption resolution 429,82]
[SA] active: 195. $less(i(l8(sK2(X0))),alength) | $less(i(main_end),X0) | i(main_end) = X0 | $less(X0,0) (2:17:1) [resolution 125,56]
[SA] new: 432. $less(i(main_end),X0) | i(main_end) = X0 | $less(X0,0) | alength = X1 | $less(X1,alength) | $less(i(l8(sK2(X0))),X1) (3:23) [resolution 195,99]
[SA] new: 433. $less(i(main_end),X2) | i(main_end) = X2 | $less(X2,0) | ~$less(X3,i(l8(sK2(X2)))) | $less(X3,alength) (3:20) [resolution 195,14]
[SA] passive: 432. $less(i(main_end),X0) | i(main_end) = X0 | $less(X0,0) | alength = X1 | $less(X1,alength) | $less(i(l8(sK2(X0))),X1) (3:23) [resolution 195,99]
[SA] passive: 433. $less(i(main_end),X2) | i(main_end) = X2 | $less(X2,0) | ~$less(X3,i(l8(sK2(X2)))) | $less(X3,alength) (3:20) [resolution 195,14]
[SA] active: 356. v = a(l8(s(s(s(zero)))),2) (4:9:1) [superposition 349,124]
[SA] new: 434. v = a(l11(s(s(zero))),2) (5:8) [superposition 356,57]
[SA] new: 435. v = a(l11(s(s(zero))),2) (5:8) [superposition 57,356]
[SA] active: 188. a(l8(X2),X3) = a(l8(p(X2)),X3) | zero = X2 | i(l8(p(X2))) = X3 (2:19:1) [superposition 115,72]
[SA] new: 436. a(l8(X0),X1) = a(l8(s(X0)),X1) | zero = s(X0) | i(l8(p(s(X0)))) = X1 (3:21) [superposition 188,169]
[SA] new: 437. v = a(l8(X2),i(l8(p(p(X2))))) | zero = X2 | i(l8(p(X2))) = i(l8(p(p(X2)))) | zero = p(X2) (3:27) [superposition 188,189]
[SA] new: 438. v = a(l8(X0),i(l8(p(p(X0))))) | zero = p(X0) | zero = X0 | i(l8(p(X0))) = i(l8(p(p(X0)))) (3:27) [superposition 189,188]
[SA] new: 439. a(l8(X0),X1) = a(l8(s(X0)),X1) | i(l8(p(s(X0)))) = X1 (3:17) [term algebras distinctness 436]
[SA] new: 440. i(l8(X0)) = X1 | a(l8(X0),X1) = a(l8(s(X0)),X1) (3:15) [forward demodulation 439,169]
[SA] passive: 440. i(l8(X0)) = X1 | a(l8(X0),X1) = a(l8(s(X0)),X1) (3:15) [forward demodulation 439,169]
[SA] passive: 437. v = a(l8(X2),i(l8(p(p(X2))))) | zero = X2 | i(l8(p(X2))) = i(l8(p(p(X2)))) | zero = p(X2) (3:27) [superposition 188,189]
[SA] active: 130. i(l8(X0)) != 2 | s(s(zero)) = X0 (3:10:1) [superposition 48,124]
[SA] new: 441. 0 != 2 | zero = s(s(zero)) (4:8) [superposition 130,65]
[SA] new: 442. i(main_end) != 2 | s(s(zero)) = nl8 (4:9) [superposition 130,54]
[SA] new: 443. 1 != 2 | s(zero) = s(s(zero)) (4:9) [superposition 130,119]
[SA] new: 445. 2 != 3 | s(s(zero)) = s(s(s(zero))) (4:11) [superposition 130,133]
[SA] new: 446. 2 != 4 | s(s(zero)) = s(s(s(s(zero)))) (5:12) [superposition 130,186]
[SA] new: 447. $sum(i(main_end),1) != 2 | s(nl8) = s(s(zero)) (4:12) [superposition 130,118]
[SA] new: 448. i(main_end) != $sum(2,-1) | s(nl8) = s(s(zero)) (4:12) [evaluation 447]
[SA] new: 449. i(main_end) != 1 | s(nl8) = s(s(zero)) (4:10) [evaluation 448]
[SA] new: 450. i(main_end) != 1 | s(zero) = nl8 (4:8) [term algebras injectivity 449]
[SA] active: 215. ~$less(X3,i(main_end)) | $less(X2,0) | i(main_end) = X2 | i(l8(sK2(X2))) = X2 | $less(X3,X2) (2:20:1) [resolution 127,14]
[SA] new: 451. $less(X0,0) | i(main_end) = X0 | i(l8(sK2(X0))) = X0 | $less(X1,X0) | i(main_end) = X1 | $less(alength,X1) | alength = i(main_end) (3:27) [resolution 215,176]
[SA] new: 452. $less(X2,0) | i(main_end) = X2 | i(l8(sK2(X2))) = X2 | $less(X3,X2) | ~$less(X3,alength) | alength = i(main_end) (3:23) [resolution 215,144]
[SA] new: 453. $less(X4,0) | i(main_end) = X4 | i(l8(sK2(X4))) = X4 | $less(X5,X4) | $less(X6,X5) | $less(i(main_end),X6) | i(main_end) = X5 | i(main_end) = X6 (3:31) [resolution 215,174]
[SA] new: 454. $less(X7,0) | i(main_end) = X7 | i(l8(sK2(X7))) = X7 | $less(X8,X7) | $less(i(main_end),X9) | $less(X9,X8) | X8 = X9 | i(main_end) = X8 (3:30) [resolution 215,174]
[SA] new: 455. $less(X10,0) | i(main_end) = X10 | i(l8(sK2(X10))) = X10 | $less(X11,X10) | $less(X12,X11) | $less(i(main_end),X12) | i(main_end) = X12 | X11 = X12 (3:30) [resolution 215,174]
[SA] new: 456. $less(X13,0) | i(main_end) = X13 | i(l8(sK2(X13))) = X13 | $less(X14,X13) | $less(i(main_end),X14) | i(main_end) = X14 (3:24) [resolution 215,15]
[SA] new: 457. $less(X15,0) | i(main_end) = X15 | i(l8(sK2(X15))) = X15 | $less(X16,X15) | $less(i(main_end),X16) | i(main_end) = X16 (3:24) [resolution 215,15]
[SA] new: 458. $less(X17,0) | i(main_end) = X17 | i(l8(sK2(X17))) = X17 | $less(alength,X17) | alength = i(main_end) (3:20) [resolution 215,100]
[SA] new: 459. $less(X18,0) | i(main_end) = X18 | i(l8(sK2(X18))) = X18 | $less(i(main_end),X18) | i(main_end) = i(l8(sK2(i(main_end)))) | $less(i(main_end),0) | i(main_end) = i(main_end) (3:34) [resolution 215,127]
[SA] new: 460. $less(X19,0) | i(main_end) = X19 | i(l8(sK2(X19))) = X19 | $less(0,X19) | $less(i(main_end),sK0) | i(main_end) = sK0 | 0 = sK0 (3:27) [resolution 215,177]
[SA] new: 461. $less(X20,0) | i(main_end) = X20 | i(l8(sK2(X20))) = X20 | $less(sK0,X20) | alength = i(main_end) (4:20) [resolution 215,255]
[SA] new: 462. $less(X21,0) | i(main_end) = X21 | i(l8(sK2(X21))) = X21 | $less(sK0,X21) | i(main_end) = sK0 (4:20) [resolution 215,206]
[SA] new: 463. $less(X22,0) | i(main_end) = X22 | i(l8(sK2(X22))) = X22 | $less(sK0,X22) | $less(i(main_end),alength) | alength = i(main_end) (3:24) [resolution 215,179]
[SA] new: 464. $less(X23,0) | i(main_end) = X23 | i(l8(sK2(X23))) = X23 | $less(sK0,X23) | $less(i(main_end),alength) | i(main_end) = sK0 (3:24) [resolution 215,106]
[SA] passive: 451. $less(X0,0) | i(main_end) = X0 | i(l8(sK2(X0))) = X0 | $less(X1,X0) | i(main_end) = X1 | $less(alength,X1) | alength = i(main_end) (3:27) [resolution 215,176]
[SA] passive: 452. $less(X2,0) | i(main_end) = X2 | i(l8(sK2(X2))) = X2 | $less(X3,X2) | ~$less(X3,alength) | alength = i(main_end) (3:23) [resolution 215,144]
[SA] passive: 453. $less(X4,0) | i(main_end) = X4 | i(l8(sK2(X4))) = X4 | $less(X5,X4) | $less(X6,X5) | $less(i(main_end),X6) | i(main_end) = X5 | i(main_end) = X6 (3:31) [resolution 215,174]
[SA] passive: 454. $less(X7,0) | i(main_end) = X7 | i(l8(sK2(X7))) = X7 | $less(X8,X7) | $less(i(main_end),X9) | $less(X9,X8) | X8 = X9 | i(main_end) = X8 (3:30) [resolution 215,174]
[SA] passive: 455. $less(X10,0) | i(main_end) = X10 | i(l8(sK2(X10))) = X10 | $less(X11,X10) | $less(X12,X11) | $less(i(main_end),X12) | i(main_end) = X12 | X11 = X12 (3:30) [resolution 215,174]
[SA] passive: 460. $less(X19,0) | i(main_end) = X19 | i(l8(sK2(X19))) = X19 | $less(0,X19) | $less(i(main_end),sK0) | i(main_end) = sK0 | 0 = sK0 (3:27) [resolution 215,177]
[SA] passive: 461. $less(X20,0) | i(main_end) = X20 | i(l8(sK2(X20))) = X20 | $less(sK0,X20) | alength = i(main_end) (4:20) [resolution 215,255]
[SA] passive: 462. $less(X21,0) | i(main_end) = X21 | i(l8(sK2(X21))) = X21 | $less(sK0,X21) | i(main_end) = sK0 (4:20) [resolution 215,206]
[SA] active: 209. i(l8(sK2(sK0))) = sK0 | i(main_end) = sK0 (3:10:1) [subsumption resolution 208,46]
[SA] new: 465. $less(sK0,alength) | $less(i(main_end),sK0) | i(main_end) = sK0 | $less(sK0,0) | i(main_end) = sK0 (4:18) [superposition 195,209]
[SA] new: 466. i(l8(X0)) != sK0 | sK2(sK0) = X0 | i(main_end) = sK0 (4:13) [superposition 48,209]
[SA] new: 467. i(l8(X1)) != sK0 | sK2(sK0) = X1 | i(main_end) = sK0 (4:13) [superposition 48,209]
[SA] new: 468. v = a(l11(sK2(sK0)),sK0) | i(main_end) = sK0 (4:11) [superposition 70,209]
[SA] new: 469. i(l8(s(sK2(sK0)))) = $sum(sK0,1) | i(main_end) = sK0 (4:13) [superposition 74,209]
[SA] new: 470. 0 != sK0 | zero = sK2(sK0) | i(main_end) = sK0 (4:11) [superposition 108,209]
[SA] new: 471. i(main_end) != sK0 | nl8 = sK2(sK0) | i(main_end) = sK0 (4:12) [superposition 109,209]
[SA] new: 472. 1 != sK0 | s(zero) = sK2(sK0) | i(main_end) = sK0 (4:12) [superposition 121,209]
[SA] new: 473. sK0 != 2 | s(s(zero)) = sK2(sK0) | i(main_end) = sK0 (4:13) [superposition 130,209]
[SA] new: 474. $sum(i(main_end),1) != sK0 | s(nl8) = sK2(sK0) | i(main_end) = sK0 (4:15) [superposition 158,209]
[SA] new: 475. v = a(l8(s(sK2(sK0))),sK0) | i(main_end) = sK0 (4:12) [superposition 349,209]
[SA] new: 476. $less(sK0,alength) | $less(i(main_end),sK0) | i(main_end) = sK0 | $less(sK0,0) (4:14) [duplicate literal removal 465]
[SA] passive: 466. i(l8(X0)) != sK0 | sK2(sK0) = X0 | i(main_end) = sK0 (4:13) [superposition 48,209]
[SA] passive: 468. v = a(l11(sK2(sK0)),sK0) | i(main_end) = sK0 (4:11) [superposition 70,209]
[SA] passive: 469. i(l8(s(sK2(sK0)))) = $sum(sK0,1) | i(main_end) = sK0 (4:13) [superposition 74,209]
[SA] passive: 470. 0 != sK0 | zero = sK2(sK0) | i(main_end) = sK0 (4:11) [superposition 108,209]
[SA] passive: 472. 1 != sK0 | s(zero) = sK2(sK0) | i(main_end) = sK0 (4:12) [superposition 121,209]
[SA] passive: 473. sK0 != 2 | s(s(zero)) = sK2(sK0) | i(main_end) = sK0 (4:13) [superposition 130,209]
[SA] passive: 474. $sum(i(main_end),1) != sK0 | s(nl8) = sK2(sK0) | i(main_end) = sK0 (4:15) [superposition 158,209]
[SA] passive: 475. v = a(l8(s(sK2(sK0))),sK0) | i(main_end) = sK0 (4:12) [superposition 349,209]
[SA] active: 226. i(l8(sK1(X2,p(nl8)))) = X2 | a(main_end,X2) = a(l11(p(nl8)),X2) | zero = nl8 (2:20:1) [resolution 150,94]
[SA] new: 477. i(l8(X1)) != X0 | sK1(X0,p(nl8)) = X1 | a(main_end,X0) = a(l11(p(nl8)),X0) | zero = nl8 (3:23) [superposition 48,226]
[SA] new: 478. i(l8(X3)) != X2 | sK1(X2,p(nl8)) = X3 | a(main_end,X2) = a(l11(p(nl8)),X2) | zero = nl8 (3:23) [superposition 48,226]
[SA] new: 479. v = a(l11(sK1(X4,p(nl8))),X4) | a(main_end,X4) = a(l11(p(nl8)),X4) | zero = nl8 (3:21) [superposition 70,226]
[SA] new: 480. i(l8(s(sK1(X5,p(nl8))))) = $sum(X5,1) | a(main_end,X5) = a(l11(p(nl8)),X5) | zero = nl8 (3:23) [superposition 74,226]
[SA] new: 481. 0 != X6 | zero = sK1(X6,p(nl8)) | a(main_end,X6) = a(l11(p(nl8)),X6) | zero = nl8 (3:21) [superposition 108,226]
[SA] new: 482. i(main_end) != X7 | nl8 = sK1(X7,p(nl8)) | a(main_end,X7) = a(l11(p(nl8)),X7) | zero = nl8 (3:22) [superposition 109,226]
[SA] new: 483. 1 != X8 | s(zero) = sK1(X8,p(nl8)) | a(main_end,X8) = a(l11(p(nl8)),X8) | zero = nl8 (3:22) [superposition 121,226]
[SA] new: 484. 2 != X9 | s(s(zero)) = sK1(X9,p(nl8)) | a(main_end,X9) = a(l11(p(nl8)),X9) | zero = nl8 (4:23) [superposition 130,226]
[SA] new: 485. $sum(i(main_end),1) != X10 | s(nl8) = sK1(X10,p(nl8)) | a(main_end,X10) = a(l11(p(nl8)),X10) | zero = nl8 (3:25) [superposition 158,226]
[SA] new: 486. v = a(l8(s(sK1(X11,p(nl8)))),X11) | a(main_end,X11) = a(l11(p(nl8)),X11) | zero = nl8 (4:22) [superposition 349,226]
[SA] passive: 477. i(l8(X1)) != X0 | sK1(X0,p(nl8)) = X1 | a(main_end,X0) = a(l11(p(nl8)),X0) | zero = nl8 (3:23) [superposition 48,226]
[SA] passive: 479. v = a(l11(sK1(X4,p(nl8))),X4) | a(main_end,X4) = a(l11(p(nl8)),X4) | zero = nl8 (3:21) [superposition 70,226]
[SA] passive: 480. i(l8(s(sK1(X5,p(nl8))))) = $sum(X5,1) | a(main_end,X5) = a(l11(p(nl8)),X5) | zero = nl8 (3:23) [superposition 74,226]
[SA] passive: 481. 0 != X6 | zero = sK1(X6,p(nl8)) | a(main_end,X6) = a(l11(p(nl8)),X6) | zero = nl8 (3:21) [superposition 108,226]
[SA] passive: 482. i(main_end) != X7 | nl8 = sK1(X7,p(nl8)) | a(main_end,X7) = a(l11(p(nl8)),X7) | zero = nl8 (3:22) [superposition 109,226]
[SA] passive: 483. 1 != X8 | s(zero) = sK1(X8,p(nl8)) | a(main_end,X8) = a(l11(p(nl8)),X8) | zero = nl8 (3:22) [superposition 121,226]
[SA] passive: 484. 2 != X9 | s(s(zero)) = sK1(X9,p(nl8)) | a(main_end,X9) = a(l11(p(nl8)),X9) | zero = nl8 (4:23) [superposition 130,226]
[SA] passive: 485. $sum(i(main_end),1) != X10 | s(nl8) = sK1(X10,p(nl8)) | a(main_end,X10) = a(l11(p(nl8)),X10) | zero = nl8 (3:25) [superposition 158,226]
[SA] passive: 486. v = a(l8(s(sK1(X11,p(nl8)))),X11) | a(main_end,X11) = a(l11(p(nl8)),X11) | zero = nl8 (4:22) [superposition 349,226]
[SA] active: 258. i(l8(sK2(sK0))) = sK0 | alength = i(main_end) (3:10:1) [subsumption resolution 257,46]
[SA] new: 487. $less(sK0,alength) | $less(i(main_end),sK0) | i(main_end) = sK0 | $less(sK0,0) | alength = i(main_end) (4:18) [superposition 195,258]
[SA] new: 488. i(l8(X0)) != sK0 | sK2(sK0) = X0 | alength = i(main_end) (4:13) [superposition 48,258]
[SA] new: 489. i(l8(X1)) != sK0 | sK2(sK0) = X1 | alength = i(main_end) (4:13) [superposition 48,258]
[SA] new: 490. v = a(l11(sK2(sK0)),sK0) | alength = i(main_end) (4:11) [superposition 70,258]
[SA] new: 491. i(l8(s(sK2(sK0)))) = $sum(sK0,1) | alength = i(main_end) (4:13) [superposition 74,258]
[SA] new: 492. 0 != sK0 | zero = sK2(sK0) | alength = i(main_end) (4:11) [superposition 108,258]
[SA] new: 493. i(main_end) != sK0 | nl8 = sK2(sK0) | alength = i(main_end) (4:12) [superposition 109,258]
[SA] new: 494. 1 != sK0 | s(zero) = sK2(sK0) | alength = i(main_end) (4:12) [superposition 121,258]
[SA] new: 495. sK0 != 2 | s(s(zero)) = sK2(sK0) | alength = i(main_end) (4:13) [superposition 130,258]
[SA] new: 496. $sum(i(main_end),1) != sK0 | s(nl8) = sK2(sK0) | alength = i(main_end) (4:15) [superposition 158,258]
[SA] new: 497. v = a(l8(s(sK2(sK0))),sK0) | alength = i(main_end) (4:12) [superposition 349,258]
[SA] passive: 488. i(l8(X0)) != sK0 | sK2(sK0) = X0 | alength = i(main_end) (4:13) [superposition 48,258]
[SA] passive: 490. v = a(l11(sK2(sK0)),sK0) | alength = i(main_end) (4:11) [superposition 70,258]
[SA] passive: 491. i(l8(s(sK2(sK0)))) = $sum(sK0,1) | alength = i(main_end) (4:13) [superposition 74,258]
[SA] passive: 492. 0 != sK0 | zero = sK2(sK0) | alength = i(main_end) (4:11) [superposition 108,258]
[SA] passive: 493. i(main_end) != sK0 | nl8 = sK2(sK0) | alength = i(main_end) (4:12) [superposition 109,258]
[SA] passive: 494. 1 != sK0 | s(zero) = sK2(sK0) | alength = i(main_end) (4:12) [superposition 121,258]
[SA] passive: 495. sK0 != 2 | s(s(zero)) = sK2(sK0) | alength = i(main_end) (4:13) [superposition 130,258]
[SA] passive: 496. $sum(i(main_end),1) != sK0 | s(nl8) = sK2(sK0) | alength = i(main_end) (4:15) [superposition 158,258]
[SA] passive: 497. v = a(l8(s(sK2(sK0))),sK0) | alength = i(main_end) (4:12) [superposition 349,258]
[SA] active: 214. $less(i(main_end),X1) | $less(X1,X0) | $less(X0,0) | X0 = X1 | i(main_end) = X0 | i(l8(sK2(X0))) = X0 (2:23:3) [resolution 127,99]
[SA] new: 498. $less(alength,X0) | $less(X0,0) | alength = X0 | i(main_end) = X0 | i(l8(sK2(X0))) = X0 (3:19) [resolution 214,82]
[SA] new: 499. $less(i(main_end),X1) | $less(X1,0) | i(main_end) = X1 | i(main_end) = X1 | i(l8(sK2(X1))) = X1 (3:21) [resolution 214,13]
[SA] new: 500. $less(X2,X3) | $less(X3,0) | X2 = X3 | i(main_end) = X3 | i(l8(sK2(X3))) = X3 | X2 = X4 | $less(X4,X2) | $less(i(main_end),X4) (3:29) [resolution 214,99]
[SA] new: 501. $less(X5,X6) | $less(X6,0) | X5 = X6 | i(main_end) = X6 | i(l8(sK2(X6))) = X6 | ~$less(X7,i(main_end)) | $less(X7,X5) (3:26) [resolution 214,14]
[SA] new: 502. $less(i(main_end),X8) | $less(X8,0) | i(main_end) = X8 | i(main_end) = X8 | i(l8(sK2(X8))) = X8 | $less(X9,0) | i(main_end) = X9 | i(l8(sK2(X9))) = X9 | $less(i(main_end),X9) (3:38) [resolution 214,215]
[SA] new: 503. $less(i(main_end),X10) | $less(X10,0) | i(main_end) = X10 | i(main_end) = X10 | i(l8(sK2(X10))) = X10 | $less(i(main_end),0) | i(main_end) = i(l8(sK2(i(main_end)))) (3:33) [resolution 214,84]
[SA] new: 504. $less(i(main_end),X11) | $less(X11,0) | i(main_end) = X11 | i(main_end) = X11 | i(l8(sK2(X11))) = X11 | $less(i(main_end),0) | 'Sub'(sK2(i(main_end)),nl8) (3:30) [resolution 214,83]
[SA] new: 505. $less(0,X12) | $less(X12,0) | 0 = X12 | i(main_end) = X12 | i(l8(sK2(X12))) = X12 | alength = 0 | $less(i(main_end),alength) (5:26) [resolution 214,262]
[SA] new: 506. $less(0,X13) | $less(X13,0) | 0 = X13 | i(main_end) = X13 | i(l8(sK2(X13))) = X13 | 0 = sK0 | $less(i(main_end),alength) (4:26) [resolution 214,116]
[SA] new: 507. $less(0,X14) | $less(X14,0) | 0 = X14 | i(main_end) = X14 | i(l8(sK2(X14))) = X14 | 0 = sK0 | $less(i(main_end),sK0) (3:26) [resolution 214,114]
[SA] new: 508. $less(sK0,X15) | $less(X15,0) | sK0 = X15 | i(main_end) = X15 | i(l8(sK2(X15))) = X15 | $less(i(main_end),alength) (3:23) [resolution 214,96]
[SA] new: 509. $less(i(main_end),sK0) | $less(0,0) | 0 = sK0 | i(main_end) = 0 | i(l8(sK2(0))) = 0 (3:20) [resolution 214,46]
[SA] new: 510. $less(i(main_end),X0) | $less(X0,0) | X0 = X0 | i(main_end) = X0 | i(l8(sK2(X0))) = X0 (3:20) [resolution 214,13]
[SA] new: 511. $less(i(main_end),X1) | $less(X2,0) | X1 = X2 | i(main_end) = X2 | i(l8(sK2(X2))) = X2 | X2 = X3 | $less(X3,X2) | $less(X1,X3) (3:29) [resolution 214,99]
[SA] new: 512. $less(i(main_end),X4) | $less(X5,0) | X4 = X5 | i(main_end) = X5 | i(l8(sK2(X5))) = X5 | ~$less(X6,X4) | $less(X6,X5) (3:26) [resolution 214,14]
[SA] new: 513. $less(i(main_end),X7) | $less(i(main_end),0) | i(main_end) = X7 | i(main_end) = i(main_end) | i(main_end) = i(l8(sK2(i(main_end)))) | $less(X8,0) | i(main_end) = X8 | i(l8(sK2(X8))) = X8 | $less(X7,X8) (3:41) [resolution 214,215]
[SA] new: 514. $less(i(main_end),X9) | $less(i(main_end),0) | i(main_end) = X9 | i(main_end) = i(main_end) | i(main_end) = i(l8(sK2(i(main_end)))) | $less(X9,0) | i(l8(sK2(X9))) = X9 (3:34) [resolution 214,84]
[SA] new: 515. $less(i(main_end),X10) | $less(i(main_end),0) | i(main_end) = X10 | i(main_end) = i(main_end) | i(main_end) = i(l8(sK2(i(main_end)))) | $less(X10,0) | 'Sub'(sK2(X10),nl8) (3:32) [resolution 214,83]
[SA] new: 516. $less(i(main_end),X11) | $less(0,0) | 0 = X11 | i(main_end) = 0 | i(l8(sK2(0))) = 0 | alength = 0 | $less(X11,alength) (5:26) [resolution 214,262]
[SA] new: 517. $less(i(main_end),X12) | $less(0,0) | 0 = X12 | i(main_end) = 0 | i(l8(sK2(0))) = 0 | 0 = sK0 | $less(X12,alength) (4:26) [resolution 214,116]
[SA] new: 518. $less(i(main_end),X13) | $less(0,0) | 0 = X13 | i(main_end) = 0 | i(l8(sK2(0))) = 0 | 0 = sK0 | $less(X13,sK0) (3:26) [resolution 214,114]
[SA] new: 519. $less(i(main_end),X14) | $less(sK0,0) | sK0 = X14 | i(main_end) = sK0 | i(l8(sK2(sK0))) = sK0 | $less(X14,alength) (3:23) [resolution 214,96]
[SA] new: 520. $less(i(main_end),i(main_end)) | $less(alength,0) | alength = i(main_end) | alength = i(main_end) | alength = i(l8(sK2(alength))) (3:22) [resolution 214,82]
[SA] new: 521. $less(i(main_end),X0) | $less(X0,sK0) | sK0 = X0 | i(main_end) = sK0 | i(l8(sK2(sK0))) = sK0 (3:20) [resolution 214,46]
[SA] new: 522. $less(i(main_end),X1) | $less(X1,X2) | X1 = X2 | i(main_end) = X2 | i(l8(sK2(X2))) = X2 | alength = 0 | $less(X2,alength) (5:26) [resolution 214,262]
[SA] new: 523. $less(i(main_end),X3) | $less(X3,X4) | X3 = X4 | i(main_end) = X4 | i(l8(sK2(X4))) = X4 | 0 = sK0 | $less(X4,alength) (4:26) [resolution 214,116]
[SA] new: 524. $less(i(main_end),X5) | $less(X5,X6) | X5 = X6 | i(main_end) = X6 | i(l8(sK2(X6))) = X6 | 0 = sK0 | $less(X6,sK0) (3:26) [resolution 214,114]
[SA] new: 525. $less(i(main_end),X7) | $less(X7,0) | 0 = X7 | i(main_end) = 0 | i(l8(sK2(0))) = 0 (3:20) [resolution 214,13]
[SA] new: 526. $less(i(main_end),X8) | $less(X8,X9) | X8 = X9 | i(main_end) = X9 | i(l8(sK2(X9))) = X9 | 0 = X10 | $less(X10,0) | $less(X9,X10) (3:29) [resolution 214,99]
[SA] new: 527. $less(i(main_end),X11) | $less(X11,X12) | X11 = X12 | i(main_end) = X12 | i(l8(sK2(X12))) = X12 | ~$less(X13,X12) | $less(X13,0) (3:26) [resolution 214,14]
[SA] new: 528. $less(i(main_end),i(main_end)) | $less(i(main_end),0) | i(main_end) = i(main_end) | i(main_end) = i(main_end) | i(main_end) = i(l8(sK2(i(main_end)))) (3:27) [factoring 214]
[SA] new: 529. $less(i(main_end),0) | $less(0,i(main_end)) | i(main_end) = 0 | i(main_end) = i(main_end) | i(main_end) = i(l8(sK2(i(main_end)))) (3:25) [factoring 214]
[SA] new: 530. $less(i(main_end),0) | $less(0,0) | 0 = 0 | i(main_end) = 0 | i(l8(sK2(0))) = 0 (3:20) [factoring 214]
[SA] new: 531. $less(i(main_end),i(main_end)) | $less(i(main_end),0) | i(main_end) = i(main_end) | i(main_end) = i(l8(sK2(i(main_end)))) (3:22) [duplicate literal removal 528]
[SA] new: 532. $less(i(main_end),i(main_end)) | $less(alength,0) | alength = i(main_end) | alength = i(l8(sK2(alength))) (3:18) [duplicate literal removal 520]
[SA] new: 533. $less(i(main_end),X13) | 0 = X13 | i(main_end) = 0 | i(l8(sK2(0))) = 0 | 0 = sK0 | $less(X13,sK0) (3:23) [evaluation 518]
[SA] new: 534. $less(i(main_end),X12) | 0 = X12 | i(main_end) = 0 | i(l8(sK2(0))) = 0 | 0 = sK0 | $less(X12,alength) (4:23) [evaluation 517]
[SA] new: 535. $less(i(main_end),X11) | 0 = X11 | i(main_end) = 0 | i(l8(sK2(0))) = 0 | alength = 0 | $less(X11,alength) (5:23) [evaluation 516]
[SA] new: 536. $less(i(main_end),sK0) | 0 = sK0 | i(main_end) = 0 | i(l8(sK2(0))) = 0 (3:17) [evaluation 509]
[SA] new: 537. $less(i(main_end),X11) | $less(X11,0) | i(main_end) = X11 | i(l8(sK2(X11))) = X11 | $less(i(main_end),0) | 'Sub'(sK2(i(main_end)),nl8) (3:26) [duplicate literal removal 504]
[SA] new: 538. $less(i(main_end),X10) | $less(X10,0) | i(main_end) = X10 | i(l8(sK2(X10))) = X10 | $less(i(main_end),0) | i(main_end) = i(l8(sK2(i(main_end)))) (3:29) [duplicate literal removal 503]
[SA] new: 539. $less(i(main_end),X8) | $less(X8,0) | i(main_end) = X8 | i(l8(sK2(X8))) = X8 | $less(X9,0) | i(main_end) = X9 | i(l8(sK2(X9))) = X9 | $less(i(main_end),X9) (3:34) [duplicate literal removal 502]
[SA] new: 540. $less(i(main_end),X1) | $less(X1,0) | i(main_end) = X1 | i(l8(sK2(X1))) = X1 (3:17) [duplicate literal removal 499]
[SA] passive: 498. $less(alength,X0) | $less(X0,0) | alength = X0 | i(main_end) = X0 | i(l8(sK2(X0))) = X0 (3:19) [resolution 214,82]
[SA] passive: 500. $less(X2,X3) | $less(X3,0) | X2 = X3 | i(main_end) = X3 | i(l8(sK2(X3))) = X3 | X2 = X4 | $less(X4,X2) | $less(i(main_end),X4) (3:29) [resolution 214,99]
[SA] passive: 501. $less(X5,X6) | $less(X6,0) | X5 = X6 | i(main_end) = X6 | i(l8(sK2(X6))) = X6 | ~$less(X7,i(main_end)) | $less(X7,X5) (3:26) [resolution 214,14]
[SA] new: 541. $less(sK0,X15) | $less(X15,0) | sK0 = X15 | i(main_end) = X15 | i(l8(sK2(X15))) = X15 (3:19) [subsumption resolution 508,82]
[SA] passive: 541. $less(sK0,X15) | $less(X15,0) | sK0 = X15 | i(main_end) = X15 | i(l8(sK2(X15))) = X15 (3:19) [subsumption resolution 508,82]
[SA] passive: 511. $less(i(main_end),X1) | $less(X2,0) | X1 = X2 | i(main_end) = X2 | i(l8(sK2(X2))) = X2 | X2 = X3 | $less(X3,X2) | $less(X1,X3) (3:29) [resolution 214,99]
[SA] passive: 512. $less(i(main_end),X4) | $less(X5,0) | X4 = X5 | i(main_end) = X5 | i(l8(sK2(X5))) = X5 | ~$less(X6,X4) | $less(X6,X5) (3:26) [resolution 214,14]
[SA] passive: 522. $less(i(main_end),X1) | $less(X1,X2) | X1 = X2 | i(main_end) = X2 | i(l8(sK2(X2))) = X2 | alength = 0 | $less(X2,alength) (5:26) [resolution 214,262]
[SA] passive: 523. $less(i(main_end),X3) | $less(X3,X4) | X3 = X4 | i(main_end) = X4 | i(l8(sK2(X4))) = X4 | 0 = sK0 | $less(X4,alength) (4:26) [resolution 214,116]
[SA] passive: 524. $less(i(main_end),X5) | $less(X5,X6) | X5 = X6 | i(main_end) = X6 | i(l8(sK2(X6))) = X6 | 0 = sK0 | $less(X6,sK0) (3:26) [resolution 214,114]
[SA] passive: 525. $less(i(main_end),X7) | $less(X7,0) | 0 = X7 | i(main_end) = 0 | i(l8(sK2(0))) = 0 (3:20) [resolution 214,13]
[SA] passive: 526. $less(i(main_end),X8) | $less(X8,X9) | X8 = X9 | i(main_end) = X9 | i(l8(sK2(X9))) = X9 | 0 = X10 | $less(X10,0) | $less(X9,X10) (3:29) [resolution 214,99]
[SA] passive: 527. $less(i(main_end),X11) | $less(X11,X12) | X11 = X12 | i(main_end) = X12 | i(l8(sK2(X12))) = X12 | ~$less(X13,X12) | $less(X13,0) (3:26) [resolution 214,14]
[SA] active: 180. i(main_end) != 3 | s(s(s(zero))) = nl8 (4:10:1) [superposition 109,133]
[SA] active: 225. $less(i(main_end),X1) | i(l8(sK1(X0,sK2(X1)))) = X0 | $less(X1,0) | a(main_end,X0) = a(l11(sK2(X1)),X0) | i(main_end) = X1 (2:28:1) [resolution 150,125]
[SA] new: 542. i(l8(sK1(X0,sK2(alength)))) = X0 | $less(alength,0) | a(main_end,X0) = a(l11(sK2(alength)),X0) | alength = i(main_end) (3:24) [resolution 225,82]
[SA] new: 543. i(l8(sK1(X1,sK2(i(main_end))))) = X1 | $less(i(main_end),0) | a(main_end,X1) = a(l11(sK2(i(main_end))),X1) | i(main_end) = i(main_end) (3:28) [resolution 225,13]
[SA] new: 544. i(l8(sK1(X2,sK2(X3)))) = X2 | $less(X3,0) | a(main_end,X2) = a(l11(sK2(X3)),X2) | i(main_end) = X3 | X3 = X4 | $less(X4,X3) | $less(i(main_end),X4) (3:34) [resolution 225,99]
[SA] new: 545. i(l8(sK1(X5,sK2(X6)))) = X5 | $less(X6,0) | a(main_end,X5) = a(l11(sK2(X6)),X5) | i(main_end) = X6 | ~$less(X7,i(main_end)) | $less(X7,X6) (3:31) [resolution 225,14]
[SA] new: 546. i(l8(sK1(X8,sK2(i(main_end))))) = X8 | $less(i(main_end),0) | a(main_end,X8) = a(l11(sK2(i(main_end))),X8) | i(main_end) = i(main_end) | $less(X9,0) | i(main_end) = X9 | i(l8(sK2(X9))) = X9 | $less(i(main_end),X9) (3:45) [resolution 225,215]
[SA] new: 547. i(l8(sK1(X10,sK2(i(main_end))))) = X10 | $less(i(main_end),0) | a(main_end,X10) = a(l11(sK2(i(main_end))),X10) | i(main_end) = i(main_end) | $less(i(main_end),0) | i(main_end) = i(l8(sK2(i(main_end)))) (3:40) [resolution 225,84]
[SA] new: 548. i(l8(sK1(X11,sK2(i(main_end))))) = X11 | $less(i(main_end),0) | a(main_end,X11) = a(l11(sK2(i(main_end))),X11) | i(main_end) = i(main_end) | $less(i(main_end),0) | 'Sub'(sK2(i(main_end)),nl8) (3:37) [resolution 225,83]
[SA] new: 549. i(l8(sK1(X12,sK2(0)))) = X12 | $less(0,0) | a(main_end,X12) = a(l11(sK2(0)),X12) | i(main_end) = 0 | alength = 0 | $less(i(main_end),alength) (5:31) [resolution 225,262]
[SA] new: 550. i(l8(sK1(X13,sK2(0)))) = X13 | $less(0,0) | a(main_end,X13) = a(l11(sK2(0)),X13) | i(main_end) = 0 | 0 = sK0 | $less(i(main_end),alength) (4:31) [resolution 225,116]
[SA] new: 551. i(l8(sK1(X14,sK2(0)))) = X14 | $less(0,0) | a(main_end,X14) = a(l11(sK2(0)),X14) | i(main_end) = 0 | 0 = sK0 | $less(i(main_end),sK0) (3:31) [resolution 225,114]
[SA] new: 552. i(l8(sK1(X15,sK2(sK0)))) = X15 | $less(sK0,0) | a(main_end,X15) = a(l11(sK2(sK0)),X15) | i(main_end) = sK0 | $less(i(main_end),alength) (3:28) [resolution 225,96]
[SA] new: 553. i(l8(sK1(X14,sK2(0)))) = X14 | a(main_end,X14) = a(l11(sK2(0)),X14) | i(main_end) = 0 | 0 = sK0 | $less(i(main_end),sK0) (3:28) [evaluation 551]
[SA] new: 554. i(l8(sK1(X13,sK2(0)))) = X13 | a(main_end,X13) = a(l11(sK2(0)),X13) | i(main_end) = 0 | 0 = sK0 | $less(i(main_end),alength) (4:28) [evaluation 550]
[SA] new: 555. i(l8(sK1(X12,sK2(0)))) = X12 | a(main_end,X12) = a(l11(sK2(0)),X12) | i(main_end) = 0 | alength = 0 | $less(i(main_end),alength) (5:28) [evaluation 549]
[SA] new: 556. i(l8(sK1(X11,sK2(i(main_end))))) = X11 | $less(i(main_end),0) | a(main_end,X11) = a(l11(sK2(i(main_end))),X11) | i(main_end) = i(main_end) | 'Sub'(sK2(i(main_end)),nl8) (3:33) [duplicate literal removal 548]
[SA] new: 557. i(l8(sK1(X10,sK2(i(main_end))))) = X10 | $less(i(main_end),0) | a(main_end,X10) = a(l11(sK2(i(main_end))),X10) | i(main_end) = i(main_end) | i(main_end) = i(l8(sK2(i(main_end)))) (3:36) [duplicate literal removal 547]
[SA] passive: 544. i(l8(sK1(X2,sK2(X3)))) = X2 | $less(X3,0) | a(main_end,X2) = a(l11(sK2(X3)),X2) | i(main_end) = X3 | X3 = X4 | $less(X4,X3) | $less(i(main_end),X4) (3:34) [resolution 225,99]
[SA] passive: 545. i(l8(sK1(X5,sK2(X6)))) = X5 | $less(X6,0) | a(main_end,X5) = a(l11(sK2(X6)),X5) | i(main_end) = X6 | ~$less(X7,i(main_end)) | $less(X7,X6) (3:31) [resolution 225,14]
[SA] new: 558. i(l8(sK1(X12,sK2(0)))) = X12 | a(main_end,X12) = a(l11(sK2(0)),X12) | i(main_end) = 0 | alength = 0 (5:24) [subsumption resolution 555,82]
[SA] passive: 558. i(l8(sK1(X12,sK2(0)))) = X12 | a(main_end,X12) = a(l11(sK2(0)),X12) | i(main_end) = 0 | alength = 0 (5:24) [subsumption resolution 555,82]
[SA] new: 559. i(l8(sK1(X13,sK2(0)))) = X13 | a(main_end,X13) = a(l11(sK2(0)),X13) | i(main_end) = 0 | 0 = sK0 (4:24) [subsumption resolution 554,82]
[SA] passive: 559. i(l8(sK1(X13,sK2(0)))) = X13 | a(main_end,X13) = a(l11(sK2(0)),X13) | i(main_end) = 0 | 0 = sK0 (4:24) [subsumption resolution 554,82]
[SA] active: 228. $less(i(l8(sK2(sK0))),alength) | i(main_end) = sK0 (4:10:1) [resolution 211,56]
[SA] new: 560. i(main_end) = sK0 | alength = X0 | $less(X0,alength) | $less(i(l8(sK2(sK0))),X0) (5:16) [resolution 228,99]
[SA] new: 561. i(main_end) = sK0 | ~$less(X1,i(l8(sK2(sK0)))) | $less(X1,alength) (5:13) [resolution 228,14]
[SA] new: 562. $less(sK0,alength) | i(main_end) = sK0 | alength = i(main_end) (5:11) [superposition 228,258]
[SA] new: 563. $less(sK0,alength) | i(main_end) = sK0 | i(main_end) = sK0 (5:11) [superposition 228,209]
[SA] new: 564. $less(sK0,alength) | i(main_end) = sK0 (5:7) [duplicate literal removal 563]
[SA] passive: 560. i(main_end) = sK0 | alength = X0 | $less(X0,alength) | $less(i(l8(sK2(sK0))),X0) (5:16) [resolution 228,99]
[SA] passive: 561. i(main_end) = sK0 | ~$less(X1,i(l8(sK2(sK0)))) | $less(X1,alength) (5:13) [resolution 228,14]
[SA] active: 341. v = a(main_end,i(l8(p(nl8)))) | zero = nl8 (3:11:1) [superposition 189,53]
[SA] active: 272. $less(i(l8(sK2(sK0))),alength) | alength = i(main_end) (4:10:1) [resolution 260,56]
[SA] new: 565. alength = i(main_end) | alength = X0 | $less(X0,alength) | $less(i(l8(sK2(sK0))),X0) (5:16) [resolution 272,99]
[SA] new: 566. alength = i(main_end) | ~$less(X1,i(l8(sK2(sK0)))) | $less(X1,alength) (5:13) [resolution 272,14]
[SA] new: 567. $less(sK0,alength) | alength = i(main_end) | alength = i(main_end) (5:11) [superposition 272,258]
[SA] new: 568. $less(sK0,alength) | alength = i(main_end) | i(main_end) = sK0 (5:11) [superposition 272,209]
[SA] new: 569. $less(sK0,alength) | alength = i(main_end) (5:7) [duplicate literal removal 567]
[SA] passive: 565. alength = i(main_end) | alength = X0 | $less(X0,alength) | $less(i(l8(sK2(sK0))),X0) (5:16) [resolution 272,99]
[SA] passive: 566. alength = i(main_end) | ~$less(X1,i(l8(sK2(sK0)))) | $less(X1,alength) (5:13) [resolution 272,14]
[SA] active: 178. $less(0,X8) | $less(X8,alength) | alength = X8 | 0 = sK0 (3:12:2) [resolution 99,113]
[SA] new: 570. $less(0,alength) | alength = 0 | 0 = sK0 (4:9) [resolution 178,13]
[SA] new: 571. $less(X0,alength) | alength = X0 | 0 = sK0 | X0 = X1 | $less(X1,X0) | $less(0,X1) (4:18) [resolution 178,99]
[SA] new: 572. $less(X2,alength) | alength = X2 | 0 = sK0 | ~$less(X3,0) | $less(X3,X2) (4:15) [resolution 178,14]
[SA] new: 573. $less(i(main_end),alength) | alength = i(main_end) | 0 = sK0 | $less(X4,0) | i(main_end) = X4 | i(l8(sK2(X4))) = X4 | $less(0,X4) (4:27) [resolution 178,215]
[SA] new: 574. $less(i(main_end),alength) | alength = i(main_end) | 0 = sK0 | $less(0,0) | i(l8(sK2(0))) = 0 (4:20) [resolution 178,84]
[SA] new: 575. $less(i(main_end),alength) | alength = i(main_end) | 0 = sK0 | $less(0,0) | 'Sub'(sK2(0),nl8) (4:18) [resolution 178,83]
[SA] new: 576. $less(0,alength) | alength = 0 | 0 = sK0 | alength = 0 | $less(0,alength) (5:15) [resolution 178,262]
[SA] new: 577. $less(0,alength) | alength = 0 | 0 = sK0 | 0 = sK0 | $less(0,alength) (4:15) [resolution 178,116]
[SA] new: 578. $less(0,alength) | alength = 0 | 0 = sK0 | 0 = sK0 | $less(0,sK0) (4:15) [resolution 178,114]
[SA] new: 579. $less(sK0,alength) | alength = sK0 | 0 = sK0 | $less(0,alength) (4:12) [resolution 178,96]
[SA] new: 580. $less(0,alength) | alength = alength | 0 = sK0 (4:9) [resolution 178,13]
[SA] new: 581. $less(0,X0) | alength = X0 | 0 = sK0 | alength = X1 | $less(X1,alength) | $less(X0,X1) (4:18) [resolution 178,99]
[SA] new: 582. $less(0,X2) | alength = X2 | 0 = sK0 | ~$less(X3,X2) | $less(X3,alength) (4:15) [resolution 178,14]
[SA] new: 583. $less(0,i(main_end)) | alength = i(main_end) | 0 = sK0 (4:11) [resolution 178,82]
[SA] new: 584. $less(0,alength) | alength = 0 | 0 = sK0 | $less(0,sK0) (4:12) [duplicate literal removal 578]
[SA] new: 585. $less(0,alength) | alength = 0 | 0 = sK0 (4:9) [duplicate literal removal 577]
[SA] new: 586. $less(0,alength) | alength = 0 | 0 = sK0 (5:9) [duplicate literal removal 576]
[SA] new: 587. $less(i(main_end),alength) | alength = i(main_end) | 0 = sK0 | 'Sub'(sK2(0),nl8) (4:15) [evaluation 575]
[SA] new: 588. $less(i(main_end),alength) | alength = i(main_end) | 0 = sK0 | i(l8(sK2(0))) = 0 (4:17) [evaluation 574]
[SA] passive: 571. $less(X0,alength) | alength = X0 | 0 = sK0 | X0 = X1 | $less(X1,X0) | $less(0,X1) (4:18) [resolution 178,99]
[SA] passive: 572. $less(X2,alength) | alength = X2 | 0 = sK0 | ~$less(X3,0) | $less(X3,X2) (4:15) [resolution 178,14]
[SA] new: 589. alength = i(main_end) | 0 = sK0 | $less(X4,0) | i(main_end) = X4 | i(l8(sK2(X4))) = X4 | $less(0,X4) (4:23) [subsumption resolution 573,82]
[SA] passive: 589. alength = i(main_end) | 0 = sK0 | $less(X4,0) | i(main_end) = X4 | i(l8(sK2(X4))) = X4 | $less(0,X4) (4:23) [subsumption resolution 573,82]
[SA] new: 590. alength = i(main_end) | 0 = sK0 | i(l8(sK2(0))) = 0 (4:13) [subsumption resolution 588,82]
[SA] passive: 590. alength = i(main_end) | 0 = sK0 | i(l8(sK2(0))) = 0 (4:13) [subsumption resolution 588,82]
[SA] new: 591. alength = i(main_end) | 0 = sK0 | 'Sub'(sK2(0),nl8) (4:11) [subsumption resolution 587,82]
[SA] passive: 591. alength = i(main_end) | 0 = sK0 | 'Sub'(sK2(0),nl8) (4:11) [subsumption resolution 587,82]
[SA] passive: 581. $less(0,X0) | alength = X0 | 0 = sK0 | alength = X1 | $less(X1,alength) | $less(X0,X1) (4:18) [resolution 178,99]
[SA] passive: 582. $less(0,X2) | alength = X2 | 0 = sK0 | ~$less(X3,X2) | $less(X3,alength) (4:15) [resolution 178,14]
[SA] passive: 583. $less(0,i(main_end)) | alength = i(main_end) | 0 = sK0 (4:11) [resolution 178,82]
[SA] active: 357. v = a(l8(s(s(s(s(zero))))),3) (4:10:1) [superposition 349,133]
[SA] new: 592. v = a(l11(s(s(s(zero)))),3) (5:9) [superposition 357,57]
[SA] new: 593. v = a(l11(s(s(s(zero)))),3) (5:9) [superposition 57,357]
[SA] active: 194. ~$less(X1,i(l8(p(nl8)))) | zero = nl8 | $less(X1,alength) (3:12:1) [resolution 95,14]
[SA] new: 594. zero = nl8 | $less(X0,alength) | $less(i(main_end),X0) | $less(i(l8(p(nl8))),0) | i(l8(p(nl8))) = X0 | i(main_end) = i(l8(p(nl8))) | i(l8(p(nl8))) = i(l8(sK2(i(l8(p(nl8)))))) (4:41) [resolution 194,214]
[SA] new: 595. zero = nl8 | $less(X1,alength) | $less(X2,X1) | $less(i(l8(p(nl8))),X2) | i(l8(p(nl8))) = X1 | i(l8(p(nl8))) = X2 (4:27) [resolution 194,174]
[SA] new: 596. zero = nl8 | $less(X3,alength) | $less(i(l8(p(nl8))),X4) | $less(X4,X3) | X3 = X4 | i(l8(p(nl8))) = X3 (4:24) [resolution 194,174]
[SA] new: 597. zero = nl8 | $less(X5,alength) | $less(X6,X5) | $less(i(l8(p(nl8))),X6) | i(l8(p(nl8))) = X6 | X5 = X6 (4:24) [resolution 194,174]
[SA] new: 598. zero = nl8 | $less(X7,alength) | $less(i(l8(p(nl8))),X7) | i(l8(p(nl8))) = X7 (4:18) [resolution 194,15]
[SA] new: 599. zero = nl8 | $less(X8,alength) | $less(i(l8(p(nl8))),X8) | i(l8(p(nl8))) = X8 (4:18) [resolution 194,15]
[SA] new: 600. zero = nl8 | $less(i(main_end),alength) | i(l8(sK1(X9,sK2(i(l8(p(nl8))))))) = X9 | $less(i(l8(p(nl8))),0) | a(main_end,X9) = a(l11(sK2(i(l8(p(nl8))))),X9) | i(main_end) = i(l8(p(nl8))) (4:43) [resolution 194,225]
[SA] new: 601. zero = nl8 | $less(i(main_end),alength) | $less(i(l8(p(nl8))),X10) | $less(X10,0) | i(l8(p(nl8))) = X10 | i(main_end) = X10 | i(l8(sK2(X10))) = X10 (4:32) [resolution 194,214]
[SA] new: 602. zero = nl8 | $less(i(main_end),alength) | i(l8(p(nl8))) = i(l8(sK2(i(l8(p(nl8)))))) | $less(i(l8(p(nl8))),0) | i(main_end) = i(l8(p(nl8))) (4:32) [resolution 194,127]
[SA] new: 603. zero = nl8 | $less(0,alength) | $less(i(l8(p(nl8))),alength) | alength = i(l8(p(nl8))) | 0 = sK0 (4:21) [resolution 194,178]
[SA] new: 604. zero = nl8 | $less(0,alength) | $less(i(l8(p(nl8))),sK0) | i(l8(p(nl8))) = sK0 | 0 = sK0 (4:21) [resolution 194,177]
[SA] new: 605. zero = nl8 | $less(sK0,alength) | $less(i(l8(p(nl8))),alength) | alength = i(l8(p(nl8))) (4:18) [resolution 194,179]
[SA] new: 606. zero = nl8 | $less(sK0,alength) | $less(i(l8(p(nl8))),alength) | i(l8(p(nl8))) = sK0 (4:18) [resolution 194,106]
[SA] passive: 594. zero = nl8 | $less(X0,alength) | $less(i(main_end),X0) | $less(i(l8(p(nl8))),0) | i(l8(p(nl8))) = X0 | i(main_end) = i(l8(p(nl8))) | i(l8(p(nl8))) = i(l8(sK2(i(l8(p(nl8)))))) (4:41) [resolution 194,214]
[SA] passive: 595. zero = nl8 | $less(X1,alength) | $less(X2,X1) | $less(i(l8(p(nl8))),X2) | i(l8(p(nl8))) = X1 | i(l8(p(nl8))) = X2 (4:27) [resolution 194,174]
[SA] passive: 596. zero = nl8 | $less(X3,alength) | $less(i(l8(p(nl8))),X4) | $less(X4,X3) | X3 = X4 | i(l8(p(nl8))) = X3 (4:24) [resolution 194,174]
[SA] passive: 597. zero = nl8 | $less(X5,alength) | $less(X6,X5) | $less(i(l8(p(nl8))),X6) | i(l8(p(nl8))) = X6 | X5 = X6 (4:24) [resolution 194,174]
[SA] passive: 598. zero = nl8 | $less(X7,alength) | $less(i(l8(p(nl8))),X7) | i(l8(p(nl8))) = X7 (4:18) [resolution 194,15]
[SA] new: 607. zero = nl8 | i(l8(sK1(X9,sK2(i(l8(p(nl8))))))) = X9 | $less(i(l8(p(nl8))),0) | a(main_end,X9) = a(l11(sK2(i(l8(p(nl8))))),X9) | i(main_end) = i(l8(p(nl8))) (4:39) [subsumption resolution 600,82]
[SA] passive: 607. zero = nl8 | i(l8(sK1(X9,sK2(i(l8(p(nl8))))))) = X9 | $less(i(l8(p(nl8))),0) | a(main_end,X9) = a(l11(sK2(i(l8(p(nl8))))),X9) | i(main_end) = i(l8(p(nl8))) (4:39) [subsumption resolution 600,82]
[SA] new: 608. zero = nl8 | $less(i(l8(p(nl8))),X10) | $less(X10,0) | i(l8(p(nl8))) = X10 | i(main_end) = X10 | i(l8(sK2(X10))) = X10 (4:28) [subsumption resolution 601,82]
[SA] passive: 608. zero = nl8 | $less(i(l8(p(nl8))),X10) | $less(X10,0) | i(l8(p(nl8))) = X10 | i(main_end) = X10 | i(l8(sK2(X10))) = X10 (4:28) [subsumption resolution 601,82]
[SA] new: 609. zero = nl8 | i(l8(p(nl8))) = i(l8(sK2(i(l8(p(nl8)))))) | $less(i(l8(p(nl8))),0) | i(main_end) = i(l8(p(nl8))) (4:28) [subsumption resolution 602,82]
[SA] passive: 609. zero = nl8 | i(l8(p(nl8))) = i(l8(sK2(i(l8(p(nl8)))))) | $less(i(l8(p(nl8))),0) | i(main_end) = i(l8(p(nl8))) (4:28) [subsumption resolution 602,82]
[SA] active: 303. v = a(l11(s(s(s(s(zero))))),4) (5:10:1) [superposition 70,186]
[SA] new: 610. v = a(l8(s(s(s(s(zero))))),4) | i(l8(s(s(s(s(zero)))))) = 4 (6:19) [superposition 303,72]
[SA] new: 611. v = a(l8(s(s(s(s(zero))))),4) | i(l8(s(s(s(s(zero)))))) = 4 (6:19) [superposition 72,303]
[SA] active: 199. ~$less(X3,sK0) | sK0 = X2 | $less(X2,alength) | $less(X3,X2) (3:12:1) [resolution 106,14]
[SA] new: 612. sK0 = X0 | $less(X0,alength) | $less(X1,X0) | $less(0,X1) | sK0 = X1 | 0 = sK0 (4:18) [resolution 199,177]
[SA] new: 613. sK0 = X2 | $less(X2,alength) | $less(X3,X2) | $less(i(main_end),X3) | $less(sK0,0) | sK0 = X3 | i(main_end) = sK0 | i(l8(sK2(sK0))) = sK0 (4:29) [resolution 199,214]
[SA] new: 614. sK0 = X4 | $less(X4,alength) | $less(X5,X4) | $less(X6,X5) | $less(sK0,X6) | sK0 = X5 | sK0 = X6 (4:21) [resolution 199,174]
[SA] new: 615. sK0 = X7 | $less(X7,alength) | $less(X8,X7) | $less(sK0,X9) | $less(X9,X8) | X8 = X9 | sK0 = X8 (4:21) [resolution 199,174]
[SA] new: 616. sK0 = X10 | $less(X10,alength) | $less(X11,X10) | $less(X12,X11) | $less(sK0,X12) | sK0 = X12 | X11 = X12 (4:21) [resolution 199,174]
[SA] new: 617. sK0 = X13 | $less(X13,alength) | $less(X14,X13) | $less(sK0,X14) | sK0 = X14 (4:15) [resolution 199,15]
[SA] new: 618. sK0 = X15 | $less(X15,alength) | $less(X16,X15) | $less(sK0,X16) | sK0 = X16 (4:15) [resolution 199,15]
[SA] new: 619. sK0 = X17 | $less(X17,alength) | $less(i(main_end),X17) | i(l8(sK1(X18,sK2(sK0)))) = X18 | $less(sK0,0) | a(main_end,X18) = a(l11(sK2(sK0)),X18) | i(main_end) = sK0 (4:34) [resolution 199,225]
[SA] new: 620. sK0 = X19 | $less(X19,alength) | $less(i(main_end),X19) | $less(sK0,X20) | $less(X20,0) | sK0 = X20 | i(main_end) = X20 | i(l8(sK2(X20))) = X20 (4:29) [resolution 199,214]
[SA] new: 621. sK0 = X21 | $less(X21,alength) | $less(i(main_end),X21) | i(l8(sK2(sK0))) = sK0 | $less(sK0,0) | i(main_end) = sK0 (4:23) [resolution 199,127]
[SA] new: 622. sK0 = X22 | $less(X22,alength) | $less(0,X22) | 0 = sK0 (4:12) [resolution 199,97]
[SA] new: 623. sK0 = X23 | $less(X23,alength) | $less(0,X23) | $less(sK0,alength) | alength = sK0 | 0 = sK0 (4:18) [resolution 199,178]
[SA] new: 624. sK0 = X24 | $less(X24,alength) | $less(0,X24) | $less(sK0,sK0) | sK0 = sK0 | 0 = sK0 (4:18) [resolution 199,177]
[SA] new: 625. sK0 = X25 | $less(X25,alength) | $less(sK0,X25) | $less(sK0,alength) | alength = sK0 (4:15) [resolution 199,179]
[SA] new: 626. sK0 = X26 | $less(X26,alength) | $less(sK0,X26) | $less(sK0,alength) | sK0 = sK0 (4:15) [resolution 199,106]
[SA] passive: 612. sK0 = X0 | $less(X0,alength) | $less(X1,X0) | $less(0,X1) | sK0 = X1 | 0 = sK0 (4:18) [resolution 199,177]
[SA] passive: 614. sK0 = X4 | $less(X4,alength) | $less(X5,X4) | $less(X6,X5) | $less(sK0,X6) | sK0 = X5 | sK0 = X6 (4:21) [resolution 199,174]
[SA] passive: 615. sK0 = X7 | $less(X7,alength) | $less(X8,X7) | $less(sK0,X9) | $less(X9,X8) | X8 = X9 | sK0 = X8 (4:21) [resolution 199,174]
[SA] passive: 616. sK0 = X10 | $less(X10,alength) | $less(X11,X10) | $less(X12,X11) | $less(sK0,X12) | sK0 = X12 | X11 = X12 (4:21) [resolution 199,174]
[SA] active: 312. i(l8(s(s(s(s(s(zero))))))) = 5 (5:10:1) [evaluation 304]
[SA] new: 627. i(l8(X0)) != 5 | s(s(s(s(s(zero))))) = X0 (6:13) [superposition 48,312]
[SA] new: 628. i(l8(X1)) != 5 | s(s(s(s(s(zero))))) = X1 (6:13) [superposition 48,312]
[SA] new: 629. v = a(l11(s(s(s(s(s(zero)))))),5) (6:11) [superposition 70,312]
[SA] new: 630. i(l8(s(s(s(s(s(s(zero)))))))) = $sum(5,1) (6:13) [superposition 74,312]
[SA] new: 631. 0 != 5 | zero = s(s(s(s(s(zero))))) (6:11) [superposition 108,312]
[SA] new: 632. i(main_end) != 5 | s(s(s(s(s(zero))))) = nl8 (6:12) [superposition 109,312]
[SA] new: 633. 1 != 5 | s(zero) = s(s(s(s(s(zero))))) (6:12) [superposition 121,312]
[SA] new: 634. 2 != 5 | s(s(zero)) = s(s(s(s(s(zero))))) (6:13) [superposition 130,312]
[SA] new: 635. $sum(i(main_end),1) != 5 | s(nl8) = s(s(s(s(s(zero))))) (6:15) [superposition 158,312]
[SA] new: 636. v = a(l8(s(s(s(s(s(s(zero))))))),5) (6:12) [superposition 349,312]
[SA] new: 637. i(main_end) != $sum(5,-1) | s(nl8) = s(s(s(s(s(zero))))) (6:15) [evaluation 635]
[SA] new: 638. i(main_end) != 4 | s(nl8) = s(s(s(s(s(zero))))) (6:13) [evaluation 637]
[SA] new: 639. i(main_end) != 4 | s(s(s(s(zero)))) = nl8 (6:11) [term algebras injectivity 638]
[SA] new: 640. i(l8(s(s(s(s(s(s(zero)))))))) = 6 (6:11) [evaluation 630]
[SA] passive: 627. i(l8(X0)) != 5 | s(s(s(s(s(zero))))) = X0 (6:13) [superposition 48,312]
[SA] passive: 629. v = a(l11(s(s(s(s(s(zero)))))),5) (6:11) [superposition 70,312]
[SA] passive: 640. i(l8(s(s(s(s(s(s(zero)))))))) = 6 (6:11) [evaluation 630]
[SA] passive: 632. i(main_end) != 5 | s(s(s(s(s(zero))))) = nl8 (6:12) [superposition 109,312]
[SA] passive: 636. v = a(l8(s(s(s(s(s(s(zero))))))),5) (6:12) [superposition 349,312]
[SA] active: 205. ~$less(X3,X2) | sK0 = X2 | $less(sK0,X2) | $less(X3,alength) (3:12:1) [resolution 106,14]
[SA] new: 641. sK0 = X0 | $less(sK0,X0) | $less(X1,alength) | $less(i(main_end),X1) | $less(X0,0) | X0 = X1 | i(main_end) = X0 | i(l8(sK2(X0))) = X0 (4:29) [resolution 205,214]
[SA] new: 642. sK0 = X2 | $less(sK0,X2) | $less(X3,alength) | $less(X4,X3) | $less(X2,X4) | X2 = X3 | X2 = X4 (4:21) [resolution 205,174]
[SA] new: 643. sK0 = X5 | $less(sK0,X5) | $less(X6,alength) | $less(X5,X7) | $less(X7,X6) | X6 = X7 | X5 = X6 (4:21) [resolution 205,174]
[SA] new: 644. sK0 = X8 | $less(sK0,X8) | $less(X9,alength) | $less(X10,X9) | $less(X8,X10) | X8 = X10 | X9 = X10 (4:21) [resolution 205,174]
[SA] new: 645. sK0 = X11 | $less(sK0,X11) | $less(X12,alength) | $less(X11,X12) | X11 = X12 (4:15) [resolution 205,15]
[SA] new: 646. sK0 = X13 | $less(sK0,X13) | $less(X14,alength) | $less(X13,X14) | X13 = X14 (4:15) [resolution 205,15]
[SA] new: 647. alength = sK0 | $less(sK0,alength) | $less(X15,alength) | $less(0,X15) | alength = X15 | 0 = sK0 (4:18) [resolution 205,178]
[SA] new: 648. alength = sK0 | $less(sK0,alength) | $less(X16,alength) | $less(sK0,X16) | alength = X16 (4:15) [resolution 205,179]
[SA] new: 649. alength = sK0 | $less(sK0,alength) | $less(X17,alength) | $less(sK0,X17) | sK0 = X17 (4:15) [resolution 205,106]
[SA] new: 650. i(main_end) = sK0 | $less(sK0,i(main_end)) | $less(X18,alength) | i(main_end) = X18 | $less(alength,X18) | alength = i(main_end) (4:22) [resolution 205,176]
[SA] new: 651. i(main_end) = sK0 | $less(sK0,i(main_end)) | $less(X19,alength) | ~$less(X19,alength) | alength = i(main_end) (4:18) [resolution 205,144]
[SA] new: 652. 0 = sK0 | $less(sK0,0) | $less(X20,alength) | $less(i(main_end),X21) | $less(X21,X20) | X20 = X21 | i(main_end) = X20 | i(l8(sK2(X20))) = X20 (4:29) [resolution 205,214]
[SA] new: 653. sK0 = sK0 | $less(sK0,sK0) | $less(X22,alength) | $less(0,X22) | sK0 = X22 | 0 = sK0 (4:18) [resolution 205,177]
[SA] new: 654. i(main_end) = sK0 | $less(sK0,i(main_end)) | $less(alength,alength) | alength = i(main_end) (4:15) [resolution 205,100]
[SA] new: 655. 0 = sK0 | $less(sK0,0) | $less(alength,alength) | alength = i(main_end) | alength = i(l8(sK2(alength))) (4:19) [resolution 205,142]
[SA] new: 656. alength = sK0 | $less(sK0,alength) | $less(i(l8(p(nl8))),alength) | zero = nl8 (4:15) [resolution 205,95]
[SA] new: 657. alength = sK0 | $less(sK0,alength) | $less(i(l8(sK2(X23))),alength) | $less(i(main_end),X23) | i(main_end) = X23 | $less(X23,0) (4:23) [resolution 205,195]
[SA] new: 658. alength = sK0 | $less(sK0,alength) | $less(i(l8(sK2(sK0))),alength) | alength = i(main_end) (5:16) [resolution 205,272]
[SA] new: 659. alength = sK0 | $less(sK0,alength) | $less(i(l8(sK2(sK0))),alength) | i(main_end) = sK0 (5:16) [resolution 205,228]
[SA] new: 660. sK0 = X24 | $less(sK0,X24) | $less(i(main_end),alength) | i(l8(sK1(X25,sK2(X24)))) = X25 | $less(X24,0) | a(main_end,X25) = a(l11(sK2(X24)),X25) | i(main_end) = X24 (4:34) [resolution 205,225]
[SA] new: 661. sK0 = X26 | $less(sK0,X26) | $less(i(main_end),alength) | $less(X26,X27) | $less(X27,0) | X26 = X27 | i(main_end) = X27 | i(l8(sK2(X27))) = X27 (4:29) [resolution 205,214]
[SA] new: 662. sK0 = X28 | $less(sK0,X28) | $less(i(main_end),alength) | i(l8(sK2(X28))) = X28 | $less(X28,0) | i(main_end) = X28 (4:23) [resolution 205,127]
[SA] new: 663. sK0 = sK0 | $less(sK0,sK0) | $less(0,alength) | 0 = sK0 (4:12) [resolution 205,97]
[SA] new: 664. alength = sK0 | $less(sK0,alength) | $less(0,alength) | alength = 0 (4:12) [resolution 205,244]
[SA] new: 665. alength = sK0 | $less(sK0,alength) | $less(0,alength) | 0 = sK0 (4:12) [resolution 205,113]
[SA] new: 666. sK0 = X29 | $less(sK0,X29) | $less(0,alength) | $less(X29,alength) | alength = X29 | 0 = sK0 (4:18) [resolution 205,178]
[SA] new: 667. sK0 = X30 | $less(sK0,X30) | $less(0,alength) | $less(X30,sK0) | sK0 = X30 | 0 = sK0 (4:18) [resolution 205,177]
[SA] new: 668. alength = sK0 | $less(sK0,alength) | $less(sK0,alength) (4:9) [resolution 205,45]
[SA] new: 669. sK0 = X31 | $less(sK0,X31) | $less(sK0,alength) | $less(X31,alength) | alength = X31 (4:15) [resolution 205,179]
[SA] new: 670. sK0 = X32 | $less(sK0,X32) | $less(sK0,alength) | $less(X32,alength) | sK0 = X32 (4:15) [resolution 205,106]
[SA] new: 671. i(main_end) = sK0 | $less(sK0,i(main_end)) | $less(sK0,alength) | alength = i(main_end) (4:15) [resolution 205,255]
[SA] new: 672. i(main_end) = sK0 | $less(sK0,i(main_end)) | $less(sK0,alength) | i(main_end) = sK0 (4:15) [resolution 205,206]
[SA] new: 673. i(main_end) = sK0 | $less(sK0,i(main_end)) | $less(sK0,alength) (4:11) [duplicate literal removal 672]
[SA] new: 674. sK0 = X32 | $less(sK0,X32) | $less(sK0,alength) | $less(X32,alength) (4:12) [duplicate literal removal 670]
[SA] new: 675. alength = sK0 | $less(sK0,alength) (4:6) [duplicate literal removal 668]
[SA] new: 676. sK0 = X30 | $less(sK0,X30) | $less(0,alength) | $less(X30,sK0) | 0 = sK0 (4:15) [duplicate literal removal 667]
[SA] passive: 642. sK0 = X2 | $less(sK0,X2) | $less(X3,alength) | $less(X4,X3) | $less(X2,X4) | X2 = X3 | X2 = X4 (4:21) [resolution 205,174]
[SA] passive: 643. sK0 = X5 | $less(sK0,X5) | $less(X6,alength) | $less(X5,X7) | $less(X7,X6) | X6 = X7 | X5 = X6 (4:21) [resolution 205,174]
[SA] passive: 644. sK0 = X8 | $less(sK0,X8) | $less(X9,alength) | $less(X10,X9) | $less(X8,X10) | X8 = X10 | X9 = X10 (4:21) [resolution 205,174]
[SA] new: 677. sK0 = X24 | $less(sK0,X24) | i(l8(sK1(X25,sK2(X24)))) = X25 | $less(X24,0) | a(main_end,X25) = a(l11(sK2(X24)),X25) | i(main_end) = X24 (4:30) [subsumption resolution 660,82]
[SA] passive: 677. sK0 = X24 | $less(sK0,X24) | i(l8(sK1(X25,sK2(X24)))) = X25 | $less(X24,0) | a(main_end,X25) = a(l11(sK2(X24)),X25) | i(main_end) = X24 (4:30) [subsumption resolution 660,82]
[SA] new: 678. sK0 = X26 | $less(sK0,X26) | $less(X26,X27) | $less(X27,0) | X26 = X27 | i(main_end) = X27 | i(l8(sK2(X27))) = X27 (4:25) [subsumption resolution 661,82]
[SA] passive: 678. sK0 = X26 | $less(sK0,X26) | $less(X26,X27) | $less(X27,0) | X26 = X27 | i(main_end) = X27 | i(l8(sK2(X27))) = X27 (4:25) [subsumption resolution 661,82]
[SA] active: 183. i(l8(X0)) != 3 | s(s(s(zero))) = X0 (4:11:1) [superposition 48,133]
[SA] new: 679. 0 != 3 | zero = s(s(s(zero))) (5:9) [superposition 183,65]
[SA] new: 680. 1 != 3 | s(zero) = s(s(s(zero))) (5:10) [superposition 183,119]
[SA] new: 681. 2 != 3 | s(s(zero)) = s(s(s(zero))) (5:11) [superposition 183,124]
[SA] new: 683. 3 != 4 | s(s(s(zero))) = s(s(s(s(zero)))) (5:13) [superposition 183,186]
[SA] new: 684. 3 != 5 | s(s(s(zero))) = s(s(s(s(s(zero))))) (6:14) [superposition 183,312]
[SA] new: 685. $sum(i(main_end),1) != 3 | s(nl8) = s(s(s(zero))) (5:13) [superposition 183,118]
[SA] new: 686. i(main_end) != 3 | s(s(s(zero))) = nl8 (5:10) [superposition 183,54]
[SA] new: 687. 3 != X0 | s(s(s(zero))) = sK1(X0,p(nl8)) | a(main_end,X0) = a(l11(p(nl8)),X0) | zero = nl8 (5:24) [superposition 183,226]
[SA] new: 688. sK0 != 3 | s(s(s(zero))) = sK2(sK0) | alength = i(main_end) (5:14) [superposition 183,258]
[SA] new: 689. sK0 != 3 | s(s(s(zero))) = sK2(sK0) | i(main_end) = sK0 (5:14) [superposition 183,209]
[SA] new: 690. i(main_end) != $sum(3,-1) | s(nl8) = s(s(s(zero))) (5:13) [evaluation 685]
[SA] new: 691. i(main_end) != 2 | s(nl8) = s(s(s(zero))) (5:11) [evaluation 690]
[SA] new: 692. i(main_end) != 2 | s(s(zero)) = nl8 (5:9) [term algebras injectivity 691]
[SA] passive: 687. 3 != X0 | s(s(s(zero))) = sK1(X0,p(nl8)) | a(main_end,X0) = a(l11(p(nl8)),X0) | zero = nl8 (5:24) [superposition 183,226]
[SA] passive: 688. sK0 != 3 | s(s(s(zero))) = sK2(sK0) | alength = i(main_end) (5:14) [superposition 183,258]
[SA] passive: 689. sK0 != 3 | s(s(s(zero))) = sK2(sK0) | i(main_end) = sK0 (5:14) [superposition 183,209]
[SA] active: 229. $less(0,X0) | $less(X0,sK0) | 0 = sK0 | 0 = X0 (3:12:2) [resolution 114,15]
[SA] new: 693. $less(0,sK0) | 0 = sK0 | 0 = 0 (4:9) [resolution 229,13]
[SA] new: 694. $less(X0,sK0) | 0 = sK0 | 0 = X0 | sK0 = X0 | $less(sK0,X0) | $less(0,alength) (4:18) [resolution 229,205]
[SA] new: 695. $less(X1,sK0) | 0 = sK0 | 0 = X1 | X1 = X2 | $less(X2,X1) | $less(0,X2) (4:18) [resolution 229,99]
[SA] new: 696. $less(X3,sK0) | 0 = sK0 | 0 = X3 | ~$less(X4,0) | $less(X4,X3) (4:15) [resolution 229,14]
[SA] new: 697. $less(i(main_end),sK0) | 0 = sK0 | i(main_end) = 0 | $less(X5,0) | i(main_end) = X5 | i(l8(sK2(X5))) = X5 | $less(0,X5) (4:27) [resolution 229,215]
[SA] new: 698. $less(i(main_end),sK0) | 0 = sK0 | i(main_end) = 0 | $less(0,0) | i(l8(sK2(0))) = 0 (4:20) [resolution 229,84]
[SA] new: 699. $less(i(main_end),sK0) | 0 = sK0 | i(main_end) = 0 | $less(0,0) | 'Sub'(sK2(0),nl8) (4:18) [resolution 229,83]
[SA] new: 700. $less(i(l8(p(nl8))),sK0) | 0 = sK0 | i(l8(p(nl8))) = 0 | zero = nl8 | $less(0,alength) (4:21) [resolution 229,194]
[SA] new: 701. $less(0,sK0) | 0 = sK0 | 0 = 0 | alength = 0 | $less(0,alength) (5:15) [resolution 229,262]
[SA] new: 702. $less(0,sK0) | 0 = sK0 | 0 = 0 | 0 = sK0 | $less(0,alength) (4:15) [resolution 229,116]
[SA] new: 703. $less(0,sK0) | 0 = sK0 | 0 = 0 | 0 = sK0 | $less(0,sK0) (4:15) [resolution 229,114]
[SA] new: 704. $less(sK0,sK0) | 0 = sK0 | 0 = sK0 | sK0 = X6 | $less(X6,alength) | $less(0,X6) (4:18) [resolution 229,199]
[SA] new: 705. $less(sK0,sK0) | 0 = sK0 | 0 = sK0 | $less(0,alength) (4:12) [resolution 229,96]
[SA] new: 706. $less(0,X0) | 0 = sK0 | 0 = X0 | sK0 = X1 | $less(X1,alength) | $less(X0,X1) (4:18) [resolution 229,199]
[SA] new: 707. $less(0,X2) | 0 = sK0 | 0 = X2 | $less(X2,alength) (4:12) [resolution 229,96]
[SA] new: 708. $less(0,sK0) | 0 = sK0 | 0 = sK0 (4:9) [resolution 229,13]
[SA] new: 709. $less(0,X3) | 0 = sK0 | 0 = X3 | sK0 = sK0 | $less(sK0,sK0) | $less(X3,alength) (4:18) [resolution 229,205]
[SA] new: 710. $less(0,X4) | 0 = sK0 | 0 = X4 | sK0 = X5 | $less(X5,sK0) | $less(X4,X5) (4:18) [resolution 229,99]
[SA] new: 711. $less(0,X6) | 0 = sK0 | 0 = X6 | ~$less(X7,X6) | $less(X7,sK0) (4:15) [resolution 229,14]
[SA] new: 712. $less(0,sK0) | 0 = sK0 (4:6) [duplicate literal removal 708]
[SA] new: 713. $less(sK0,sK0) | 0 = sK0 | $less(0,alength) (4:9) [duplicate literal removal 705]
[SA] new: 714. $less(sK0,sK0) | 0 = sK0 | sK0 = X6 | $less(X6,alength) | $less(0,X6) (4:15) [duplicate literal removal 704]
[SA] new: 715. $less(0,sK0) | 0 = sK0 | 0 = 0 (4:9) [duplicate literal removal 703]
[SA] new: 716. $less(0,sK0) | 0 = sK0 | 0 = 0 | $less(0,alength) (4:12) [duplicate literal removal 702]
[SA] new: 717. $less(i(main_end),sK0) | 0 = sK0 | i(main_end) = 0 | 'Sub'(sK2(0),nl8) (4:15) [evaluation 699]
[SA] new: 718. $less(i(main_end),sK0) | 0 = sK0 | i(main_end) = 0 | i(l8(sK2(0))) = 0 (4:17) [evaluation 698]
[SA] passive: 696. $less(X3,sK0) | 0 = sK0 | 0 = X3 | ~$less(X4,0) | $less(X4,X3) (4:15) [resolution 229,14]
[SA] passive: 697. $less(i(main_end),sK0) | 0 = sK0 | i(main_end) = 0 | $less(X5,0) | i(main_end) = X5 | i(l8(sK2(X5))) = X5 | $less(0,X5) (4:27) [resolution 229,215]
[SA] passive: 717. $less(i(main_end),sK0) | 0 = sK0 | i(main_end) = 0 | 'Sub'(sK2(0),nl8) (4:15) [evaluation 699]
[SA] passive: 706. $less(0,X0) | 0 = sK0 | 0 = X0 | sK0 = X1 | $less(X1,alength) | $less(X0,X1) (4:18) [resolution 229,199]
[SA] passive: 710. $less(0,X4) | 0 = sK0 | 0 = X4 | sK0 = X5 | $less(X5,sK0) | $less(X4,X5) (4:18) [resolution 229,99]
[SA] passive: 711. $less(0,X6) | 0 = sK0 | 0 = X6 | ~$less(X7,X6) | $less(X7,sK0) (4:15) [resolution 229,14]
[SA] active: 224. $less(X1,i(main_end)) | ~$less(X1,sK0) | i(main_end) = sK0 (4:11:1) [resolution 206,14]
[SA] new: 719. ~$less(X0,sK0) | i(main_end) = sK0 | $less(X1,0) | i(main_end) = X1 | i(l8(sK2(X1))) = X1 | $less(X0,X1) (5:23) [resolution 224,215]
[SA] new: 720. ~$less(X2,sK0) | i(main_end) = sK0 | $less(X2,0) | i(l8(sK2(X2))) = X2 (5:16) [resolution 224,84]
[SA] new: 721. ~$less(X3,sK0) | i(main_end) = sK0 | $less(X3,0) | 'Sub'(sK2(X3),nl8) (5:14) [resolution 224,83]
[SA] new: 722. ~$less(i(main_end),sK0) | i(main_end) = sK0 (5:8) [resolution 224,13]
[SA] new: 723. ~$less(X4,sK0) | i(main_end) = sK0 | i(main_end) = sK0 | $less(sK0,i(main_end)) | $less(X4,alength) (5:18) [resolution 224,205]
[SA] new: 724. ~$less(X5,sK0) | i(main_end) = sK0 | i(main_end) = X6 | $less(X6,i(main_end)) | $less(X5,X6) (5:18) [resolution 224,99]
[SA] new: 725. ~$less(X7,sK0) | i(main_end) = sK0 | ~$less(X8,X7) | $less(X8,i(main_end)) (5:14) [resolution 224,14]
[SA] new: 726. ~$less(X4,sK0) | i(main_end) = sK0 | $less(sK0,i(main_end)) | $less(X4,alength) (5:14) [duplicate literal removal 723]
[SA] passive: 719. ~$less(X0,sK0) | i(main_end) = sK0 | $less(X1,0) | i(main_end) = X1 | i(l8(sK2(X1))) = X1 | $less(X0,X1) (5:23) [resolution 224,215]
[SA] passive: 720. ~$less(X2,sK0) | i(main_end) = sK0 | $less(X2,0) | i(l8(sK2(X2))) = X2 (5:16) [resolution 224,84]
[SA] passive: 721. ~$less(X3,sK0) | i(main_end) = sK0 | $less(X3,0) | 'Sub'(sK2(X3),nl8) (5:14) [resolution 224,83]
[SA] passive: 722. ~$less(i(main_end),sK0) | i(main_end) = sK0 (5:8) [resolution 224,13]
[SA] passive: 724. ~$less(X5,sK0) | i(main_end) = sK0 | i(main_end) = X6 | $less(X6,i(main_end)) | $less(X5,X6) (5:18) [resolution 224,99]
[SA] passive: 725. ~$less(X7,sK0) | i(main_end) = sK0 | ~$less(X8,X7) | $less(X8,i(main_end)) (5:14) [resolution 224,14]
[SA] active: 247. ~$less(X3,sK0) | alength = X2 | $less(X2,alength) | $less(X3,X2) (3:12:1) [resolution 179,14]
[SA] new: 727. alength = X0 | $less(X0,alength) | $less(X1,X0) | $less(0,X1) | 0 = sK0 | 0 = X1 (4:18) [resolution 247,229]
[SA] new: 728. alength = X2 | $less(X2,alength) | $less(X3,X2) | $less(0,X3) | sK0 = X3 | 0 = sK0 (4:18) [resolution 247,177]
[SA] new: 729. alength = X4 | $less(X4,alength) | $less(X5,X4) | $less(i(main_end),X5) | $less(sK0,0) | sK0 = X5 | i(main_end) = sK0 | i(l8(sK2(sK0))) = sK0 (4:29) [resolution 247,214]
[SA] new: 730. alength = X6 | $less(X6,alength) | $less(X7,X6) | $less(X8,X7) | $less(sK0,X8) | sK0 = X7 | sK0 = X8 (4:21) [resolution 247,174]
[SA] new: 731. alength = X9 | $less(X9,alength) | $less(X10,X9) | $less(sK0,X11) | $less(X11,X10) | X10 = X11 | sK0 = X10 (4:21) [resolution 247,174]
[SA] new: 732. alength = X12 | $less(X12,alength) | $less(X13,X12) | $less(X14,X13) | $less(sK0,X14) | sK0 = X14 | X13 = X14 (4:21) [resolution 247,174]
[SA] new: 733. alength = X15 | $less(X15,alength) | $less(X16,X15) | $less(sK0,X16) | sK0 = X16 (4:15) [resolution 247,15]
[SA] new: 734. alength = X17 | $less(X17,alength) | $less(X18,X17) | $less(sK0,X18) | sK0 = X18 (4:15) [resolution 247,15]
[SA] new: 735. alength = X19 | $less(X19,alength) | $less(i(main_end),X19) | i(l8(sK1(X20,sK2(sK0)))) = X20 | $less(sK0,0) | a(main_end,X20) = a(l11(sK2(sK0)),X20) | i(main_end) = sK0 (4:34) [resolution 247,225]
[SA] new: 736. alength = X21 | $less(X21,alength) | $less(i(main_end),X21) | $less(sK0,X22) | $less(X22,0) | sK0 = X22 | i(main_end) = X22 | i(l8(sK2(X22))) = X22 (4:29) [resolution 247,214]
[SA] new: 737. alength = X23 | $less(X23,alength) | $less(i(main_end),X23) | i(l8(sK2(sK0))) = sK0 | $less(sK0,0) | i(main_end) = sK0 (4:23) [resolution 247,127]
[SA] new: 738. alength = X24 | $less(X24,alength) | $less(0,X24) | 0 = sK0 (4:12) [resolution 247,97]
[SA] new: 739. alength = X25 | $less(X25,alength) | $less(0,X25) | $less(sK0,sK0) | 0 = sK0 | 0 = sK0 (4:18) [resolution 247,229]
[SA] new: 740. alength = X26 | $less(X26,alength) | $less(0,X26) | $less(sK0,alength) | alength = sK0 | 0 = sK0 (4:18) [resolution 247,178]
[SA] new: 741. alength = X27 | $less(X27,alength) | $less(0,X27) | $less(sK0,sK0) | sK0 = sK0 | 0 = sK0 (4:18) [resolution 247,177]
[SA] new: 742. alength = X28 | $less(X28,alength) | $less(sK0,X28) | $less(sK0,alength) | alength = sK0 (4:15) [resolution 247,179]
[SA] new: 743. alength = X29 | $less(X29,alength) | $less(sK0,X29) | $less(sK0,alength) | sK0 = sK0 (4:15) [resolution 247,106]
[SA] new: 744. alength = X25 | $less(X25,alength) | $less(0,X25) | $less(sK0,sK0) | 0 = sK0 (4:15) [duplicate literal removal 739]
[SA] passive: 727. alength = X0 | $less(X0,alength) | $less(X1,X0) | $less(0,X1) | 0 = sK0 | 0 = X1 (4:18) [resolution 247,229]
[SA] passive: 728. alength = X2 | $less(X2,alength) | $less(X3,X2) | $less(0,X3) | sK0 = X3 | 0 = sK0 (4:18) [resolution 247,177]
[SA] passive: 730. alength = X6 | $less(X6,alength) | $less(X7,X6) | $less(X8,X7) | $less(sK0,X8) | sK0 = X7 | sK0 = X8 (4:21) [resolution 247,174]
[SA] passive: 731. alength = X9 | $less(X9,alength) | $less(X10,X9) | $less(sK0,X11) | $less(X11,X10) | X10 = X11 | sK0 = X10 (4:21) [resolution 247,174]
[SA] passive: 732. alength = X12 | $less(X12,alength) | $less(X13,X12) | $less(X14,X13) | $less(sK0,X14) | sK0 = X14 | X13 = X14 (4:21) [resolution 247,174]
[SA] active: 722. ~$less(i(main_end),sK0) | i(main_end) = sK0 (5:8:1) [resolution 224,13]
[SA] new: 745. i(main_end) = sK0 | i(l8(sK1(X0,sK2(sK0)))) = X0 | $less(sK0,0) | a(main_end,X0) = a(l11(sK2(sK0)),X0) | i(main_end) = sK0 (6:28) [resolution 722,225]
[SA] new: 746. i(main_end) = sK0 | $less(sK0,X1) | $less(X1,0) | sK0 = X1 | i(main_end) = X1 | i(l8(sK2(X1))) = X1 (6:23) [resolution 722,214]
[SA] new: 747. i(main_end) = sK0 | i(l8(sK2(sK0))) = sK0 | $less(sK0,0) | i(main_end) = sK0 (6:17) [resolution 722,127]
[SA] new: 748. i(main_end) = sK0 | $less(0,i(main_end)) | 0 = sK0 | i(main_end) = 0 (6:15) [resolution 722,229]
[SA] new: 749. i(main_end) = sK0 | $less(0,i(main_end)) | i(main_end) = sK0 | 0 = sK0 (6:15) [resolution 722,177]
[SA] new: 750. i(main_end) = sK0 | $less(i(main_end),i(main_end)) | $less(sK0,0) | i(main_end) = sK0 | i(main_end) = sK0 | i(l8(sK2(sK0))) = sK0 (6:26) [resolution 722,214]
[SA] new: 751. i(main_end) = sK0 | $less(X2,i(main_end)) | $less(sK0,X2) | i(main_end) = sK0 | sK0 = X2 (6:18) [resolution 722,174]
[SA] new: 752. i(main_end) = sK0 | $less(sK0,X3) | $less(X3,i(main_end)) | i(main_end) = X3 | i(main_end) = sK0 (6:19) [resolution 722,174]
[SA] new: 753. i(main_end) = sK0 | $less(X4,i(main_end)) | $less(sK0,X4) | sK0 = X4 | i(main_end) = X4 (6:18) [resolution 722,174]
[SA] new: 754. i(main_end) = sK0 | $less(sK0,i(main_end)) | i(main_end) = sK0 (6:12) [resolution 722,15]
[SA] new: 755. i(main_end) = sK0 | $less(sK0,i(main_end)) | i(main_end) = sK0 (6:12) [resolution 722,15]
[SA] new: 756. i(main_end) = sK0 | $less(sK0,i(main_end)) (6:8) [duplicate literal removal 755]
[SA] new: 757. i(main_end) = sK0 | $less(sK0,i(main_end)) (6:8) [duplicate literal removal 754]
[SA] new: 758. i(main_end) = sK0 | $less(sK0,X3) | $less(X3,i(main_end)) | i(main_end) = X3 (6:15) [duplicate literal removal 752]
[SA] new: 759. i(main_end) = sK0 | $less(X2,i(main_end)) | $less(sK0,X2) | sK0 = X2 (6:14) [duplicate literal removal 751]
[SA] new: 760. i(main_end) = sK0 | $less(i(main_end),i(main_end)) | $less(sK0,0) | i(l8(sK2(sK0))) = sK0 (6:18) [duplicate literal removal 750]
[SA] new: 761. i(main_end) = sK0 | $less(0,i(main_end)) | 0 = sK0 (6:11) [duplicate literal removal 749]
[SA] new: 762. i(main_end) = sK0 | i(l8(sK2(sK0))) = sK0 | $less(sK0,0) (6:13) [duplicate literal removal 747]
[SA] new: 763. i(main_end) = sK0 | i(l8(sK1(X0,sK2(sK0)))) = X0 | $less(sK0,0) | a(main_end,X0) = a(l11(sK2(sK0)),X0) (6:24) [duplicate literal removal 745]
[SA] passive: 748. i(main_end) = sK0 | $less(0,i(main_end)) | 0 = sK0 | i(main_end) = 0 (6:15) [resolution 722,229]
[SA] passive: 761. i(main_end) = sK0 | $less(0,i(main_end)) | 0 = sK0 (6:11) [duplicate literal removal 749]
[SA] passive: 759. i(main_end) = sK0 | $less(X2,i(main_end)) | $less(sK0,X2) | sK0 = X2 (6:14) [duplicate literal removal 751]
[SA] active: 254. ~$less(X3,X2) | alength = X2 | $less(sK0,X2) | $less(X3,alength) (3:12:1) [resolution 179,14]
[SA] new: 764. alength = X0 | $less(sK0,X0) | $less(X1,alength) | $less(i(main_end),X1) | $less(X0,0) | X0 = X1 | i(main_end) = X0 | i(l8(sK2(X0))) = X0 (4:29) [resolution 254,214]
[SA] new: 765. alength = X2 | $less(sK0,X2) | $less(X3,alength) | $less(X4,X3) | $less(X2,X4) | X2 = X3 | X2 = X4 (4:21) [resolution 254,174]
[SA] new: 766. alength = X5 | $less(sK0,X5) | $less(X6,alength) | $less(X5,X7) | $less(X7,X6) | X6 = X7 | X5 = X6 (4:21) [resolution 254,174]
[SA] new: 767. alength = X8 | $less(sK0,X8) | $less(X9,alength) | $less(X10,X9) | $less(X8,X10) | X8 = X10 | X9 = X10 (4:21) [resolution 254,174]
[SA] new: 768. alength = X11 | $less(sK0,X11) | $less(X12,alength) | $less(X11,X12) | X11 = X12 (4:15) [resolution 254,15]
[SA] new: 769. alength = X13 | $less(sK0,X13) | $less(X14,alength) | $less(X13,X14) | X13 = X14 (4:15) [resolution 254,15]
[SA] new: 770. alength = alength | $less(sK0,alength) | $less(X15,alength) | $less(0,X15) | alength = X15 | 0 = sK0 (4:18) [resolution 254,178]
[SA] new: 771. alength = alength | $less(sK0,alength) | $less(X16,alength) | $less(sK0,X16) | alength = X16 (4:15) [resolution 254,179]
[SA] new: 772. alength = alength | $less(sK0,alength) | $less(X17,alength) | $less(sK0,X17) | sK0 = X17 (4:15) [resolution 254,106]
[SA] new: 773. alength = i(main_end) | $less(sK0,i(main_end)) | $less(X18,alength) | ~$less(X18,sK0) | i(main_end) = sK0 (5:18) [resolution 254,224]
[SA] new: 774. alength = i(main_end) | $less(sK0,i(main_end)) | $less(X19,alength) | i(main_end) = X19 | $less(alength,X19) | alength = i(main_end) (4:22) [resolution 254,176]
[SA] new: 775. alength = i(main_end) | $less(sK0,i(main_end)) | $less(X20,alength) | ~$less(X20,alength) | alength = i(main_end) (4:18) [resolution 254,144]
[SA] new: 776. alength = 0 | $less(sK0,0) | $less(X21,alength) | $less(i(main_end),X22) | $less(X22,X21) | X21 = X22 | i(main_end) = X21 | i(l8(sK2(X21))) = X21 (4:29) [resolution 254,214]
[SA] new: 777. alength = sK0 | $less(sK0,sK0) | $less(X23,alength) | $less(0,X23) | 0 = sK0 | 0 = X23 (4:18) [resolution 254,229]
[SA] new: 778. alength = sK0 | $less(sK0,sK0) | $less(X24,alength) | $less(0,X24) | sK0 = X24 | 0 = sK0 (4:18) [resolution 254,177]
[SA] new: 779. alength = i(main_end) | $less(sK0,i(main_end)) | $less(alength,alength) | alength = i(main_end) (4:15) [resolution 254,100]
[SA] new: 780. alength = 0 | $less(sK0,0) | $less(alength,alength) | alength = i(main_end) | alength = i(l8(sK2(alength))) (4:19) [resolution 254,142]
[SA] new: 781. alength = alength | $less(sK0,alength) | $less(i(l8(p(nl8))),alength) | zero = nl8 (4:15) [resolution 254,95]
[SA] new: 782. alength = alength | $less(sK0,alength) | $less(i(l8(sK2(X25))),alength) | $less(i(main_end),X25) | i(main_end) = X25 | $less(X25,0) (4:23) [resolution 254,195]
[SA] new: 783. alength = alength | $less(sK0,alength) | $less(i(l8(sK2(sK0))),alength) | alength = i(main_end) (5:16) [resolution 254,272]
[SA] new: 784. alength = alength | $less(sK0,alength) | $less(i(l8(sK2(sK0))),alength) | i(main_end) = sK0 (5:16) [resolution 254,228]
[SA] new: 785. alength = X26 | $less(sK0,X26) | $less(i(main_end),alength) | i(l8(sK1(X27,sK2(X26)))) = X27 | $less(X26,0) | a(main_end,X27) = a(l11(sK2(X26)),X27) | i(main_end) = X26 (4:34) [resolution 254,225]
[SA] new: 786. alength = X28 | $less(sK0,X28) | $less(i(main_end),alength) | $less(X28,X29) | $less(X29,0) | X28 = X29 | i(main_end) = X29 | i(l8(sK2(X29))) = X29 (4:29) [resolution 254,214]
[SA] new: 787. alength = X30 | $less(sK0,X30) | $less(i(main_end),alength) | i(l8(sK2(X30))) = X30 | $less(X30,0) | i(main_end) = X30 (4:23) [resolution 254,127]
[SA] new: 788. alength = sK0 | $less(sK0,sK0) | $less(0,alength) | 0 = sK0 (4:12) [resolution 254,97]
[SA] new: 789. alength = alength | $less(sK0,alength) | $less(0,alength) | alength = 0 (4:12) [resolution 254,244]
[SA] new: 790. alength = alength | $less(sK0,alength) | $less(0,alength) | 0 = sK0 (4:12) [resolution 254,113]
[SA] new: 791. alength = X31 | $less(sK0,X31) | $less(0,alength) | $less(X31,sK0) | 0 = sK0 | 0 = X31 (4:18) [resolution 254,229]
[SA] new: 792. alength = X32 | $less(sK0,X32) | $less(0,alength) | $less(X32,alength) | alength = X32 | 0 = sK0 (4:18) [resolution 254,178]
[SA] new: 793. alength = X33 | $less(sK0,X33) | $less(0,alength) | $less(X33,sK0) | sK0 = X33 | 0 = sK0 (4:18) [resolution 254,177]
[SA] new: 794. alength = alength | $less(sK0,alength) | $less(sK0,alength) (4:9) [resolution 254,45]
[SA] new: 795. alength = X34 | $less(sK0,X34) | $less(sK0,alength) | $less(X34,alength) | alength = X34 (4:15) [resolution 254,179]
[SA] new: 796. alength = X35 | $less(sK0,X35) | $less(sK0,alength) | $less(X35,alength) | sK0 = X35 (4:15) [resolution 254,106]
[SA] new: 797. alength = i(main_end) | $less(sK0,i(main_end)) | $less(sK0,alength) | alength = i(main_end) (4:15) [resolution 254,255]
[SA] new: 798. alength = i(main_end) | $less(sK0,i(main_end)) | $less(sK0,alength) | i(main_end) = sK0 (4:15) [resolution 254,206]
[SA] new: 799. alength = i(main_end) | $less(sK0,i(main_end)) | $less(sK0,alength) (4:11) [duplicate literal removal 797]
[SA] new: 800. alength = X34 | $less(sK0,X34) | $less(sK0,alength) | $less(X34,alength) (4:12) [duplicate literal removal 795]
[SA] new: 801. alength = alength | $less(sK0,alength) (4:6) [duplicate literal removal 794]
[SA] new: 802. alength = X32 | $less(sK0,X32) | $less(0,alength) | $less(X32,alength) | 0 = sK0 (4:15) [duplicate literal removal 792]
[SA] new: 803. alength = i(main_end) | $less(sK0,i(main_end)) | $less(alength,alength) (4:11) [duplicate literal removal 779]
[SA] new: 804. alength = i(main_end) | $less(sK0,i(main_end)) | $less(X20,alength) | ~$less(X20,alength) (4:14) [duplicate literal removal 775]
[SA] new: 805. alength = i(main_end) | $less(sK0,i(main_end)) | $less(X19,alength) | i(main_end) = X19 | $less(alength,X19) (4:18) [duplicate literal removal 774]
[SA] passive: 764. alength = X0 | $less(sK0,X0) | $less(X1,alength) | $less(i(main_end),X1) | $less(X0,0) | X0 = X1 | i(main_end) = X0 | i(l8(sK2(X0))) = X0 (4:29) [resolution 254,214]
[SA] passive: 765. alength = X2 | $less(sK0,X2) | $less(X3,alength) | $less(X4,X3) | $less(X2,X4) | X2 = X3 | X2 = X4 (4:21) [resolution 254,174]
[SA] passive: 766. alength = X5 | $less(sK0,X5) | $less(X6,alength) | $less(X5,X7) | $less(X7,X6) | X6 = X7 | X5 = X6 (4:21) [resolution 254,174]
[SA] passive: 767. alength = X8 | $less(sK0,X8) | $less(X9,alength) | $less(X10,X9) | $less(X8,X10) | X8 = X10 | X9 = X10 (4:21) [resolution 254,174]
[SA] passive: 768. alength = X11 | $less(sK0,X11) | $less(X12,alength) | $less(X11,X12) | X11 = X12 (4:15) [resolution 254,15]
[SA] new: 806. alength = X26 | $less(sK0,X26) | i(l8(sK1(X27,sK2(X26)))) = X27 | $less(X26,0) | a(main_end,X27) = a(l11(sK2(X26)),X27) | i(main_end) = X26 (4:30) [subsumption resolution 785,82]
[SA] passive: 806. alength = X26 | $less(sK0,X26) | i(l8(sK1(X27,sK2(X26)))) = X27 | $less(X26,0) | a(main_end,X27) = a(l11(sK2(X26)),X27) | i(main_end) = X26 (4:30) [subsumption resolution 785,82]
[SA] new: 807. alength = X28 | $less(sK0,X28) | $less(X28,X29) | $less(X29,0) | X28 = X29 | i(main_end) = X29 | i(l8(sK2(X29))) = X29 (4:25) [subsumption resolution 786,82]
[SA] passive: 807. alength = X28 | $less(sK0,X28) | $less(X28,X29) | $less(X29,0) | X28 = X29 | i(main_end) = X29 | i(l8(sK2(X29))) = X29 (4:25) [subsumption resolution 786,82]
[SA] new: 808. alength = X30 | $less(sK0,X30) | i(l8(sK2(X30))) = X30 | $less(X30,0) | i(main_end) = X30 (4:19) [subsumption resolution 787,82]
[SA] passive: 808. alength = X30 | $less(sK0,X30) | i(l8(sK2(X30))) = X30 | $less(X30,0) | i(main_end) = X30 (4:19) [subsumption resolution 787,82]
[SA] active: 268. $less(X1,i(main_end)) | ~$less(X1,sK0) | alength = i(main_end) (4:11:1) [resolution 255,14]
[SA] new: 809. ~$less(X0,sK0) | alength = i(main_end) | $less(X1,0) | i(main_end) = X1 | i(l8(sK2(X1))) = X1 | $less(X0,X1) (5:23) [resolution 268,215]
[SA] new: 810. ~$less(X2,sK0) | alength = i(main_end) | $less(X2,0) | i(l8(sK2(X2))) = X2 (5:16) [resolution 268,84]
[SA] new: 811. ~$less(X3,sK0) | alength = i(main_end) | $less(X3,0) | 'Sub'(sK2(X3),nl8) (5:14) [resolution 268,83]
[SA] new: 812. ~$less(i(main_end),sK0) | alength = i(main_end) (5:8) [resolution 268,13]
[SA] new: 813. ~$less(X4,sK0) | alength = i(main_end) | alength = i(main_end) | $less(sK0,i(main_end)) | $less(X4,alength) (5:18) [resolution 268,254]
[SA] new: 814. ~$less(X5,sK0) | alength = i(main_end) | i(main_end) = sK0 | $less(sK0,i(main_end)) | $less(X5,alength) (5:18) [resolution 268,205]
[SA] new: 815. ~$less(X6,sK0) | alength = i(main_end) | i(main_end) = X7 | $less(X7,i(main_end)) | $less(X6,X7) (5:18) [resolution 268,99]
[SA] new: 816. ~$less(X8,sK0) | alength = i(main_end) | ~$less(X9,X8) | $less(X9,i(main_end)) (5:14) [resolution 268,14]
[SA] new: 817. ~$less(X4,sK0) | alength = i(main_end) | $less(sK0,i(main_end)) | $less(X4,alength) (5:14) [duplicate literal removal 813]
[SA] passive: 809. ~$less(X0,sK0) | alength = i(main_end) | $less(X1,0) | i(main_end) = X1 | i(l8(sK2(X1))) = X1 | $less(X0,X1) (5:23) [resolution 268,215]
[SA] passive: 810. ~$less(X2,sK0) | alength = i(main_end) | $less(X2,0) | i(l8(sK2(X2))) = X2 (5:16) [resolution 268,84]
[SA] passive: 811. ~$less(X3,sK0) | alength = i(main_end) | $less(X3,0) | 'Sub'(sK2(X3),nl8) (5:14) [resolution 268,83]
[SA] passive: 812. ~$less(i(main_end),sK0) | alength = i(main_end) (5:8) [resolution 268,13]
[SA] passive: 815. ~$less(X6,sK0) | alength = i(main_end) | i(main_end) = X7 | $less(X7,i(main_end)) | $less(X6,X7) (5:18) [resolution 268,99]
[SA] passive: 816. ~$less(X8,sK0) | alength = i(main_end) | ~$less(X9,X8) | $less(X9,i(main_end)) (5:14) [resolution 268,14]
[SA] active: 330. $less(0,X0) | $less(X0,alength) | 0 = sK0 | sK0 = X0 (3:12:2) [resolution 177,96]
[SA] new: 818. $less(0,alength) | 0 = sK0 | 0 = sK0 (4:9) [resolution 330,13]
[SA] new: 819. $less(X0,alength) | 0 = sK0 | sK0 = X0 | alength = X0 | $less(sK0,X0) | $less(0,alength) (4:18) [resolution 330,254]
[SA] new: 820. $less(X1,alength) | 0 = sK0 | sK0 = X1 | sK0 = X1 | $less(sK0,X1) | $less(0,alength) (4:18) [resolution 330,205]
[SA] new: 821. $less(X2,alength) | 0 = sK0 | sK0 = X2 | X2 = X3 | $less(X3,X2) | $less(0,X3) (4:18) [resolution 330,99]
[SA] new: 822. $less(X4,alength) | 0 = sK0 | sK0 = X4 | ~$less(X5,0) | $less(X5,X4) (4:15) [resolution 330,14]
[SA] new: 823. $less(i(main_end),alength) | 0 = sK0 | i(main_end) = sK0 | $less(X6,0) | i(main_end) = X6 | i(l8(sK2(X6))) = X6 | $less(0,X6) (4:27) [resolution 330,215]
[SA] new: 824. $less(i(main_end),alength) | 0 = sK0 | i(main_end) = sK0 | $less(0,0) | i(l8(sK2(0))) = 0 (4:20) [resolution 330,84]
[SA] new: 825. $less(i(main_end),alength) | 0 = sK0 | i(main_end) = sK0 | $less(0,0) | 'Sub'(sK2(0),nl8) (4:18) [resolution 330,83]
[SA] new: 826. $less(i(l8(p(nl8))),alength) | 0 = sK0 | i(l8(p(nl8))) = sK0 | zero = nl8 | $less(0,alength) (4:21) [resolution 330,194]
[SA] new: 827. $less(0,alength) | 0 = sK0 | 0 = sK0 | alength = 0 | $less(0,alength) (5:15) [resolution 330,262]
[SA] new: 828. $less(0,alength) | 0 = sK0 | 0 = sK0 | 0 = sK0 | $less(0,alength) (4:15) [resolution 330,116]
[SA] new: 829. $less(0,alength) | 0 = sK0 | 0 = sK0 | 0 = sK0 | $less(0,sK0) (4:15) [resolution 330,114]
[SA] new: 830. $less(sK0,alength) | 0 = sK0 | sK0 = sK0 | alength = X7 | $less(X7,alength) | $less(0,X7) (4:18) [resolution 330,247]
[SA] new: 831. $less(sK0,alength) | 0 = sK0 | sK0 = sK0 | sK0 = X8 | $less(X8,alength) | $less(0,X8) (4:18) [resolution 330,199]
[SA] new: 832. $less(sK0,alength) | 0 = sK0 | sK0 = sK0 | $less(0,alength) (4:12) [resolution 330,96]
[SA] new: 833. $less(0,alength) | 0 = sK0 | alength = sK0 (4:9) [resolution 330,13]
[SA] new: 834. $less(0,X0) | 0 = sK0 | sK0 = X0 | alength = alength | $less(sK0,alength) | $less(X0,alength) (4:18) [resolution 330,254]
[SA] new: 835. $less(0,X1) | 0 = sK0 | sK0 = X1 | alength = sK0 | $less(sK0,alength) | $less(X1,alength) (4:18) [resolution 330,205]
[SA] new: 836. $less(0,X2) | 0 = sK0 | sK0 = X2 | alength = X3 | $less(X3,alength) | $less(X2,X3) (4:18) [resolution 330,99]
[SA] new: 837. $less(0,X4) | 0 = sK0 | sK0 = X4 | ~$less(X5,X4) | $less(X5,alength) (4:15) [resolution 330,14]
[SA] new: 838. $less(0,i(main_end)) | 0 = sK0 | i(main_end) = sK0 (4:11) [resolution 330,82]
[SA] new: 839. $less(0,alength) | 0 = sK0 | $less(0,sK0) (4:9) [duplicate literal removal 829]
[SA] new: 840. $less(0,alength) | 0 = sK0 (4:6) [duplicate literal removal 828]
[SA] new: 841. $less(0,alength) | 0 = sK0 | alength = 0 (5:9) [duplicate literal removal 827]
[SA] new: 842. $less(i(main_end),alength) | 0 = sK0 | i(main_end) = sK0 | 'Sub'(sK2(0),nl8) (4:15) [evaluation 825]
[SA] new: 843. $less(i(main_end),alength) | 0 = sK0 | i(main_end) = sK0 | i(l8(sK2(0))) = 0 (4:17) [evaluation 824]
[SA] new: 844. $less(X1,alength) | 0 = sK0 | sK0 = X1 | $less(sK0,X1) | $less(0,alength) (4:15) [duplicate literal removal 820]
[SA] new: 845. $less(0,alength) | 0 = sK0 (4:6) [duplicate literal removal 818]
[SA] passive: 821. $less(X2,alength) | 0 = sK0 | sK0 = X2 | X2 = X3 | $less(X3,X2) | $less(0,X3) (4:18) [resolution 330,99]
[SA] passive: 822. $less(X4,alength) | 0 = sK0 | sK0 = X4 | ~$less(X5,0) | $less(X5,X4) (4:15) [resolution 330,14]
[SA] new: 846. 0 = sK0 | i(main_end) = sK0 | $less(X6,0) | i(main_end) = X6 | i(l8(sK2(X6))) = X6 | $less(0,X6) (4:23) [subsumption resolution 823,82]
[SA] passive: 846. 0 = sK0 | i(main_end) = sK0 | $less(X6,0) | i(main_end) = X6 | i(l8(sK2(X6))) = X6 | $less(0,X6) (4:23) [subsumption resolution 823,82]
[SA] new: 847. 0 = sK0 | i(main_end) = sK0 | i(l8(sK2(0))) = 0 (4:13) [subsumption resolution 843,82]
[SA] passive: 847. 0 = sK0 | i(main_end) = sK0 | i(l8(sK2(0))) = 0 (4:13) [subsumption resolution 843,82]
[SA] new: 848. 0 = sK0 | i(main_end) = sK0 | 'Sub'(sK2(0),nl8) (4:11) [subsumption resolution 842,82]
[SA] passive: 848. 0 = sK0 | i(main_end) = sK0 | 'Sub'(sK2(0),nl8) (4:11) [subsumption resolution 842,82]
[SA] passive: 837. $less(0,X4) | 0 = sK0 | sK0 = X4 | ~$less(X5,X4) | $less(X5,alength) (4:15) [resolution 330,14]
[SA] active: 812. ~$less(i(main_end),sK0) | alength = i(main_end) (5:8:1) [resolution 268,13]
[SA] new: 849. alength = i(main_end) | i(l8(sK1(X0,sK2(sK0)))) = X0 | $less(sK0,0) | a(main_end,X0) = a(l11(sK2(sK0)),X0) | i(main_end) = sK0 (6:28) [resolution 812,225]
[SA] new: 850. alength = i(main_end) | $less(sK0,X1) | $less(X1,0) | sK0 = X1 | i(main_end) = X1 | i(l8(sK2(X1))) = X1 (6:23) [resolution 812,214]
[SA] new: 851. alength = i(main_end) | i(l8(sK2(sK0))) = sK0 | $less(sK0,0) | i(main_end) = sK0 (6:17) [resolution 812,127]
[SA] new: 852. alength = i(main_end) | $less(0,i(main_end)) | 0 = sK0 | i(main_end) = 0 (6:15) [resolution 812,229]
[SA] new: 853. alength = i(main_end) | $less(0,i(main_end)) | i(main_end) = sK0 | 0 = sK0 (6:15) [resolution 812,177]
[SA] new: 854. alength = i(main_end) | $less(i(main_end),i(main_end)) | $less(sK0,0) | i(main_end) = sK0 | i(main_end) = sK0 | i(l8(sK2(sK0))) = sK0 (6:26) [resolution 812,214]
[SA] new: 855. alength = i(main_end) | $less(X2,i(main_end)) | $less(sK0,X2) | i(main_end) = sK0 | sK0 = X2 (6:18) [resolution 812,174]
[SA] new: 856. alength = i(main_end) | $less(sK0,X3) | $less(X3,i(main_end)) | i(main_end) = X3 | i(main_end) = sK0 (6:19) [resolution 812,174]
[SA] new: 857. alength = i(main_end) | $less(X4,i(main_end)) | $less(sK0,X4) | sK0 = X4 | i(main_end) = X4 (6:18) [resolution 812,174]
[SA] new: 858. alength = i(main_end) | $less(sK0,i(main_end)) | i(main_end) = sK0 (6:12) [resolution 812,15]
[SA] new: 859. alength = i(main_end) | $less(sK0,i(main_end)) | i(main_end) = sK0 (6:12) [resolution 812,15]
[SA] new: 860. alength = i(main_end) | $less(i(main_end),i(main_end)) | $less(sK0,0) | i(main_end) = sK0 | i(l8(sK2(sK0))) = sK0 (6:22) [duplicate literal removal 854]
[SA] active: 380. $less(0,X0) | $less(X0,sK0) | 0 = X0 | sK0 = X0 (3:12:2) [resolution 174,46]
[SA] new: 861. $less(0,sK0) | 0 = 0 | 0 = sK0 (4:9) [resolution 380,13]
[SA] new: 862. $less(X0,sK0) | 0 = X0 | sK0 = X0 | alength = X0 | $less(sK0,X0) | $less(0,alength) (4:18) [resolution 380,254]
[SA] new: 863. $less(X1,sK0) | 0 = X1 | sK0 = X1 | sK0 = X1 | $less(sK0,X1) | $less(0,alength) (4:18) [resolution 380,205]
[SA] new: 864. $less(X2,sK0) | 0 = X2 | sK0 = X2 | X2 = X3 | $less(X3,X2) | $less(0,X3) (4:18) [resolution 380,99]
[SA] new: 865. $less(X4,sK0) | 0 = X4 | sK0 = X4 | ~$less(X5,0) | $less(X5,X4) (4:15) [resolution 380,14]
[SA] new: 866. $less(i(main_end),sK0) | i(main_end) = 0 | i(main_end) = sK0 | $less(X6,0) | i(main_end) = X6 | i(l8(sK2(X6))) = X6 | $less(0,X6) (4:28) [resolution 380,215]
[SA] new: 867. $less(i(main_end),sK0) | i(main_end) = 0 | i(main_end) = sK0 | $less(0,0) | i(l8(sK2(0))) = 0 (4:21) [resolution 380,84]
[SA] new: 868. $less(i(main_end),sK0) | i(main_end) = 0 | i(main_end) = sK0 | $less(0,0) | 'Sub'(sK2(0),nl8) (4:19) [resolution 380,83]
[SA] new: 869. $less(i(l8(p(nl8))),sK0) | i(l8(p(nl8))) = 0 | i(l8(p(nl8))) = sK0 | zero = nl8 | $less(0,alength) (4:24) [resolution 380,194]
[SA] new: 870. $less(0,sK0) | 0 = 0 | 0 = sK0 | alength = 0 | $less(0,alength) (5:15) [resolution 380,262]
[SA] new: 871. $less(0,sK0) | 0 = 0 | 0 = sK0 | 0 = sK0 | $less(0,alength) (4:15) [resolution 380,116]
[SA] new: 872. $less(0,sK0) | 0 = 0 | 0 = sK0 | 0 = sK0 | $less(0,sK0) (4:15) [resolution 380,114]
[SA] new: 873. $less(sK0,sK0) | 0 = sK0 | sK0 = sK0 | alength = X7 | $less(X7,alength) | $less(0,X7) (4:18) [resolution 380,247]
[SA] new: 874. $less(sK0,sK0) | 0 = sK0 | sK0 = sK0 | sK0 = X8 | $less(X8,alength) | $less(0,X8) (4:18) [resolution 380,199]
[SA] new: 875. $less(sK0,sK0) | 0 = sK0 | sK0 = sK0 | $less(0,alength) (4:12) [resolution 380,96]
[SA] new: 876. $less(0,X0) | 0 = X0 | sK0 = X0 | alength = X1 | $less(X1,alength) | $less(X0,X1) (4:18) [resolution 380,247]
[SA] new: 877. $less(0,X2) | 0 = X2 | sK0 = X2 | sK0 = X3 | $less(X3,alength) | $less(X2,X3) (4:18) [resolution 380,199]
[SA] new: 878. $less(0,X4) | 0 = X4 | sK0 = X4 | $less(X4,alength) (4:12) [resolution 380,96]
[SA] new: 879. $less(0,sK0) | 0 = sK0 | sK0 = sK0 (4:9) [resolution 380,13]
[SA] new: 880. $less(0,X5) | 0 = X5 | sK0 = X5 | alength = sK0 | $less(sK0,sK0) | $less(X5,alength) (4:18) [resolution 380,254]
[SA] new: 881. $less(0,X6) | 0 = X6 | sK0 = X6 | sK0 = sK0 | $less(sK0,sK0) | $less(X6,alength) (4:18) [resolution 380,205]
[SA] new: 882. $less(0,X7) | 0 = X7 | sK0 = X7 | sK0 = X8 | $less(X8,sK0) | $less(X7,X8) (4:18) [resolution 380,99]
[SA] new: 883. $less(0,X9) | 0 = X9 | sK0 = X9 | ~$less(X10,X9) | $less(X10,sK0) (4:15) [resolution 380,14]
[SA] new: 884. $less(0,i(main_end)) | i(main_end) = 0 | i(main_end) = sK0 | alength = i(main_end) (6:16) [resolution 380,812]
[SA] new: 885. $less(0,i(main_end)) | i(main_end) = 0 | i(main_end) = sK0 | i(main_end) = sK0 (6:16) [resolution 380,722]
[SA] new: 886. $less(0,i(main_end)) | i(main_end) = 0 | i(main_end) = sK0 (6:12) [duplicate literal removal 885]
[SA] new: 887. $less(0,sK0) | 0 = 0 | 0 = sK0 (4:9) [duplicate literal removal 872]
[SA] new: 888. $less(0,sK0) | 0 = 0 | 0 = sK0 | $less(0,alength) (4:12) [duplicate literal removal 871]
[SA] new: 889. $less(i(main_end),sK0) | i(main_end) = 0 | i(main_end) = sK0 | 'Sub'(sK2(0),nl8) (4:16) [evaluation 868]
[SA] new: 890. $less(i(main_end),sK0) | i(main_end) = 0 | i(main_end) = sK0 | i(l8(sK2(0))) = 0 (4:18) [evaluation 867]
[SA] new: 891. $less(X1,sK0) | 0 = X1 | sK0 = X1 | $less(sK0,X1) | $less(0,alength) (4:15) [duplicate literal removal 863]
[SA] passive: 864. $less(X2,sK0) | 0 = X2 | sK0 = X2 | X2 = X3 | $less(X3,X2) | $less(0,X3) (4:18) [resolution 380,99]
[SA] passive: 865. $less(X4,sK0) | 0 = X4 | sK0 = X4 | ~$less(X5,0) | $less(X5,X4) (4:15) [resolution 380,14]
[SA] new: 892. i(main_end) = 0 | i(main_end) = sK0 | $less(X6,0) | i(main_end) = X6 | i(l8(sK2(X6))) = X6 | $less(0,X6) (4:24) [subsumption resolution 866,722]
[SA] passive: 892. i(main_end) = 0 | i(main_end) = sK0 | $less(X6,0) | i(main_end) = X6 | i(l8(sK2(X6))) = X6 | $less(0,X6) (4:24) [subsumption resolution 866,722]
[SA] new: 893. i(main_end) = 0 | i(main_end) = sK0 | i(l8(sK2(0))) = 0 (4:14) [subsumption resolution 890,722]
[SA] passive: 893. i(main_end) = 0 | i(main_end) = sK0 | i(l8(sK2(0))) = 0 (4:14) [subsumption resolution 890,722]
[SA] new: 894. i(main_end) = 0 | i(main_end) = sK0 | 'Sub'(sK2(0),nl8) (4:12) [subsumption resolution 889,722]
[SA] passive: 894. i(main_end) = 0 | i(main_end) = sK0 | 'Sub'(sK2(0),nl8) (4:12) [subsumption resolution 889,722]
[SA] passive: 869. $less(i(l8(p(nl8))),sK0) | i(l8(p(nl8))) = 0 | i(l8(p(nl8))) = sK0 | zero = nl8 | $less(0,alength) (4:24) [resolution 380,194]
[SA] passive: 876. $less(0,X0) | 0 = X0 | sK0 = X0 | alength = X1 | $less(X1,alength) | $less(X0,X1) (4:18) [resolution 380,247]
[SA] passive: 877. $less(0,X2) | 0 = X2 | sK0 = X2 | sK0 = X3 | $less(X3,alength) | $less(X2,X3) (4:18) [resolution 380,199]
[SA] passive: 878. $less(0,X4) | 0 = X4 | sK0 = X4 | $less(X4,alength) (4:12) [resolution 380,96]
[SA] passive: 882. $less(0,X7) | 0 = X7 | sK0 = X7 | sK0 = X8 | $less(X8,sK0) | $less(X7,X8) (4:18) [resolution 380,99]
[SA] passive: 883. $less(0,X9) | 0 = X9 | sK0 = X9 | ~$less(X10,X9) | $less(X10,sK0) (4:15) [resolution 380,14]
[SA] passive: 884. $less(0,i(main_end)) | i(main_end) = 0 | i(main_end) = sK0 | alength = i(main_end) (6:16) [resolution 380,812]
[SA] passive: 886. $less(0,i(main_end)) | i(main_end) = 0 | i(main_end) = sK0 (6:12) [duplicate literal removal 885]
[SA] active: 359. v = a(l8(s(s(nl8))),$sum(i(main_end),1)) (4:11:1) [superposition 349,118]
[SA] new: 895. v = a(l11(s(nl8)),$sum(i(main_end),1)) (5:10) [superposition 359,57]
[SA] new: 896. v = a(l11(s(nl8)),$sum(i(main_end),1)) (5:10) [superposition 57,359]
[SA] active: 270. $less(i(l8(sK2(alength))),alength) | alength = i(main_end) | $less(alength,0) (3:13:1) [resolution 143,56]
[SA] new: 897. alength = i(main_end) | $less(alength,0) | alength = alength | $less(sK0,alength) | $less(i(l8(sK2(alength))),alength) (4:19) [resolution 270,254]
[SA] new: 898. alength = i(main_end) | $less(alength,0) | alength = sK0 | $less(sK0,alength) | $less(i(l8(sK2(alength))),alength) (4:19) [resolution 270,205]
[SA] new: 899. alength = i(main_end) | $less(alength,0) | alength = X0 | $less(X0,alength) | $less(i(l8(sK2(alength))),X0) (4:19) [resolution 270,99]
[SA] new: 900. alength = i(main_end) | $less(alength,0) | ~$less(X1,i(l8(sK2(alength)))) | $less(X1,alength) (4:16) [resolution 270,14]
[SA] passive: 899. alength = i(main_end) | $less(alength,0) | alength = X0 | $less(X0,alength) | $less(i(l8(sK2(alength))),X0) (4:19) [resolution 270,99]
[SA] passive: 900. alength = i(main_end) | $less(alength,0) | ~$less(X1,i(l8(sK2(alength)))) | $less(X1,alength) (4:16) [resolution 270,14]
[SA] active: 468. v = a(l11(sK2(sK0)),sK0) | i(main_end) = sK0 (4:11:1) [superposition 70,209]
[SA] new: 901. v = a(l8(sK2(sK0)),sK0) | i(main_end) = sK0 | i(l8(sK2(sK0))) = sK0 (5:17) [superposition 468,72]
[SA] new: 902. v = a(l8(sK2(sK0)),sK0) | i(l8(sK2(sK0))) = sK0 | i(main_end) = sK0 (5:17) [superposition 72,468]
[SA] active: 274. 'Sub'(sK2(X1),nl8) | alength = i(main_end) | $less(X1,0) | ~$less(X1,alength) (3:14:1) [resolution 144,83]
[SA] new: 903. alength = i(main_end) | $less(X0,0) | ~$less(X0,alength) | a(main_end,X1) = a(l11(sK2(X0)),X1) | i(l8(sK1(X1,sK2(X0)))) = X1 (4:27) [resolution 274,150]
[SA] new: 904. alength = i(main_end) | $less(X2,0) | ~$less(X2,alength) | $less(i(l8(sK2(X2))),alength) (4:16) [resolution 274,56]
[SA] passive: 903. alength = i(main_end) | $less(X0,0) | ~$less(X0,alength) | a(main_end,X1) = a(l11(sK2(X0)),X1) | i(l8(sK1(X1,sK2(X0)))) = X1 (4:27) [resolution 274,150]
[SA] passive: 904. alength = i(main_end) | $less(X2,0) | ~$less(X2,alength) | $less(i(l8(sK2(X2))),alength) (4:16) [resolution 274,56]
[SA] active: 470. 0 != sK0 | zero = sK2(sK0) | i(main_end) = sK0 (4:11:1) [superposition 108,209]
[SA] active: 277. ~$less(X5,X4) | alength = i(main_end) | ~$less(X4,alength) | $less(X5,i(main_end)) (3:14:1) [resolution 144,14]
[SA] new: 905. alength = i(main_end) | ~$less(X0,alength) | $less(X1,i(main_end)) | $less(i(main_end),X1) | $less(X0,0) | X0 = X1 | i(main_end) = X0 | i(l8(sK2(X0))) = X0 (4:31) [resolution 277,214]
[SA] new: 906. alength = i(main_end) | ~$less(X2,alength) | $less(X3,i(main_end)) | $less(X4,X3) | $less(X2,X4) | X2 = X3 | X2 = X4 (4:23) [resolution 277,174]
[SA] new: 907. alength = i(main_end) | ~$less(X5,alength) | $less(X6,i(main_end)) | $less(X5,X7) | $less(X7,X6) | X6 = X7 | X5 = X6 (4:23) [resolution 277,174]
[SA] new: 908. alength = i(main_end) | ~$less(X8,alength) | $less(X9,i(main_end)) | $less(X10,X9) | $less(X8,X10) | X8 = X10 | X9 = X10 (4:23) [resolution 277,174]
[SA] new: 909. alength = i(main_end) | ~$less(X11,alength) | $less(X12,i(main_end)) | $less(X11,X12) | X11 = X12 (4:17) [resolution 277,15]
[SA] new: 910. alength = i(main_end) | ~$less(X13,alength) | $less(X14,i(main_end)) | $less(X13,X14) | X13 = X14 (4:17) [resolution 277,15]
[SA] new: 911. alength = i(main_end) | ~$less(alength,alength) | $less(X15,i(main_end)) | $less(0,X15) | 0 = sK0 | sK0 = X15 (4:20) [resolution 277,330]
[SA] new: 912. alength = i(main_end) | ~$less(alength,alength) | $less(X16,i(main_end)) | $less(0,X16) | alength = X16 | 0 = sK0 (4:20) [resolution 277,178]
[SA] new: 913. alength = i(main_end) | ~$less(alength,alength) | $less(X17,i(main_end)) | $less(sK0,X17) | alength = X17 (4:17) [resolution 277,179]
[SA] new: 914. alength = i(main_end) | ~$less(alength,alength) | $less(X18,i(main_end)) | $less(sK0,X18) | sK0 = X18 (4:17) [resolution 277,106]
[SA] new: 915. alength = i(main_end) | ~$less(i(main_end),alength) | $less(X19,i(main_end)) | ~$less(X19,sK0) | alength = i(main_end) (5:19) [resolution 277,268]
[SA] new: 916. alength = i(main_end) | ~$less(i(main_end),alength) | $less(X20,i(main_end)) | ~$less(X20,sK0) | i(main_end) = sK0 (5:19) [resolution 277,224]
[SA] new: 917. alength = i(main_end) | ~$less(i(main_end),alength) | $less(X21,i(main_end)) | i(main_end) = X21 | $less(alength,X21) | alength = i(main_end) (4:23) [resolution 277,176]
[SA] new: 918. alength = i(main_end) | ~$less(i(main_end),alength) | $less(X22,i(main_end)) | ~$less(X22,alength) | alength = i(main_end) (4:19) [resolution 277,144]
[SA] new: 919. alength = i(main_end) | ~$less(0,alength) | $less(X23,i(main_end)) | $less(i(main_end),X24) | $less(X24,X23) | X23 = X24 | i(main_end) = X23 | i(l8(sK2(X23))) = X23 (4:31) [resolution 277,214]
[SA] new: 920. alength = i(main_end) | ~$less(sK0,alength) | $less(X25,i(main_end)) | $less(0,X25) | 0 = X25 | sK0 = X25 (4:20) [resolution 277,380]
[SA] new: 921. alength = i(main_end) | ~$less(sK0,alength) | $less(X26,i(main_end)) | $less(0,X26) | 0 = sK0 | 0 = X26 (4:20) [resolution 277,229]
[SA] new: 922. alength = i(main_end) | ~$less(sK0,alength) | $less(X27,i(main_end)) | $less(0,X27) | sK0 = X27 | 0 = sK0 (4:20) [resolution 277,177]
[SA] new: 923. alength = i(main_end) | ~$less(i(main_end),alength) | $less(alength,i(main_end)) | alength = i(main_end) (4:16) [resolution 277,100]
[SA] new: 924. alength = i(main_end) | ~$less(0,alength) | $less(alength,i(main_end)) | alength = i(main_end) | alength = i(l8(sK2(alength))) (4:21) [resolution 277,142]
[SA] new: 925. alength = i(main_end) | ~$less(alength,alength) | $less(i(l8(p(nl8))),i(main_end)) | zero = nl8 (4:17) [resolution 277,95]
[SA] new: 926. alength = i(main_end) | ~$less(alength,alength) | $less(i(l8(sK2(X28))),i(main_end)) | $less(i(main_end),X28) | i(main_end) = X28 | $less(X28,0) (4:25) [resolution 277,195]
[SA] new: 927. alength = i(main_end) | ~$less(alength,alength) | $less(i(l8(sK2(sK0))),i(main_end)) | alength = i(main_end) (5:18) [resolution 277,272]
[SA] new: 928. alength = i(main_end) | ~$less(alength,alength) | $less(i(l8(sK2(sK0))),i(main_end)) | i(main_end) = sK0 (5:18) [resolution 277,228]
[SA] new: 929. alength = i(main_end) | ~$less(alength,alength) | $less(i(l8(sK2(alength))),i(main_end)) | alength = i(main_end) | $less(alength,0) (4:21) [resolution 277,270]
[SA] new: 930. alength = i(main_end) | ~$less(X29,alength) | $less(i(main_end),i(main_end)) | i(l8(sK1(X30,sK2(X29)))) = X30 | $less(X29,0) | a(main_end,X30) = a(l11(sK2(X29)),X30) | i(main_end) = X29 (4:36) [resolution 277,225]
[SA] new: 931. alength = i(main_end) | ~$less(X31,alength) | $less(i(main_end),i(main_end)) | $less(X31,X32) | $less(X32,0) | X31 = X32 | i(main_end) = X32 | i(l8(sK2(X32))) = X32 (4:31) [resolution 277,214]
[SA] new: 932. alength = i(main_end) | ~$less(X33,alength) | $less(i(main_end),i(main_end)) | i(l8(sK2(X33))) = X33 | $less(X33,0) | i(main_end) = X33 (4:25) [resolution 277,127]
[SA] new: 933. alength = i(main_end) | ~$less(sK0,alength) | $less(0,i(main_end)) | 0 = sK0 (4:14) [resolution 277,97]
[SA] new: 934. alength = i(main_end) | ~$less(alength,alength) | $less(0,i(main_end)) | alength = 0 (4:14) [resolution 277,244]
[SA] new: 935. alength = i(main_end) | ~$less(alength,alength) | $less(0,i(main_end)) | 0 = sK0 (4:14) [resolution 277,113]
[SA] new: 936. alength = i(main_end) | ~$less(X34,alength) | $less(0,i(main_end)) | $less(X34,sK0) | 0 = X34 | sK0 = X34 (4:20) [resolution 277,380]
[SA] new: 937. alength = i(main_end) | ~$less(X35,alength) | $less(0,i(main_end)) | $less(X35,alength) | 0 = sK0 | sK0 = X35 (4:20) [resolution 277,330]
[SA] new: 938. alength = i(main_end) | ~$less(X36,alength) | $less(0,i(main_end)) | $less(X36,sK0) | 0 = sK0 | 0 = X36 (4:20) [resolution 277,229]
[SA] new: 939. alength = i(main_end) | ~$less(X37,alength) | $less(0,i(main_end)) | $less(X37,alength) | alength = X37 | 0 = sK0 (4:20) [resolution 277,178]
[SA] new: 940. alength = i(main_end) | ~$less(X38,alength) | $less(0,i(main_end)) | $less(X38,sK0) | sK0 = X38 | 0 = sK0 (4:20) [resolution 277,177]
[SA] new: 941. alength = i(main_end) | ~$less(alength,alength) | $less(sK0,i(main_end)) (4:11) [resolution 277,45]
[SA] new: 942. alength = i(main_end) | ~$less(X39,alength) | $less(sK0,i(main_end)) | $less(X39,alength) | alength = X39 (4:17) [resolution 277,179]
[SA] new: 943. alength = i(main_end) | ~$less(X40,alength) | $less(sK0,i(main_end)) | $less(X40,alength) | sK0 = X40 (4:17) [resolution 277,106]
[SA] new: 944. alength = i(main_end) | ~$less(i(main_end),alength) | $less(sK0,i(main_end)) | alength = i(main_end) (4:16) [resolution 277,255]
[SA] new: 945. alength = i(main_end) | ~$less(i(main_end),alength) | $less(sK0,i(main_end)) | i(main_end) = sK0 (4:16) [resolution 277,206]
[SA] new: 946. alength = i(main_end) | ~$less(i(main_end),alength) | $less(sK0,i(main_end)) (4:12) [duplicate literal removal 944]
[SA] new: 947. alength = i(main_end) | ~$less(alength,alength) | $less(i(l8(sK2(alength))),i(main_end)) | $less(alength,0) (4:17) [duplicate literal removal 929]
[SA] new: 948. alength = i(main_end) | ~$less(alength,alength) | $less(i(l8(sK2(sK0))),i(main_end)) (5:14) [duplicate literal removal 927]
[SA] new: 949. alength = i(main_end) | ~$less(0,alength) | $less(alength,i(main_end)) | alength = i(l8(sK2(alength))) (4:17) [duplicate literal removal 924]
[SA] new: 950. alength = i(main_end) | ~$less(i(main_end),alength) | $less(alength,i(main_end)) (4:12) [duplicate literal removal 923]
[SA] new: 951. alength = i(main_end) | ~$less(i(main_end),alength) | $less(X22,i(main_end)) | ~$less(X22,alength) (4:15) [duplicate literal removal 918]
[SA] new: 952. alength = i(main_end) | ~$less(i(main_end),alength) | $less(X21,i(main_end)) | i(main_end) = X21 | $less(alength,X21) (4:19) [duplicate literal removal 917]
[SA] new: 953. alength = i(main_end) | ~$less(i(main_end),alength) | $less(X19,i(main_end)) | ~$less(X19,sK0) (5:15) [duplicate literal removal 915]
[SA] passive: 906. alength = i(main_end) | ~$less(X2,alength) | $less(X3,i(main_end)) | $less(X4,X3) | $less(X2,X4) | X2 = X3 | X2 = X4 (4:23) [resolution 277,174]
[SA] passive: 907. alength = i(main_end) | ~$less(X5,alength) | $less(X6,i(main_end)) | $less(X5,X7) | $less(X7,X6) | X6 = X7 | X5 = X6 (4:23) [resolution 277,174]
[SA] passive: 908. alength = i(main_end) | ~$less(X8,alength) | $less(X9,i(main_end)) | $less(X10,X9) | $less(X8,X10) | X8 = X10 | X9 = X10 (4:23) [resolution 277,174]
[SA] passive: 909. alength = i(main_end) | ~$less(X11,alength) | $less(X12,i(main_end)) | $less(X11,X12) | X11 = X12 (4:17) [resolution 277,15]
[SA] new: 954. alength = i(main_end) | $less(X25,i(main_end)) | $less(0,X25) | 0 = X25 | sK0 = X25 (4:17) [subsumption resolution 920,45]
[SA] passive: 954. alength = i(main_end) | $less(X25,i(main_end)) | $less(0,X25) | 0 = X25 | sK0 = X25 (4:17) [subsumption resolution 920,45]
[SA] new: 955. alength = i(main_end) | $less(X26,i(main_end)) | $less(0,X26) | 0 = sK0 | 0 = X26 (4:17) [subsumption resolution 921,45]
[SA] passive: 955. alength = i(main_end) | $less(X26,i(main_end)) | $less(0,X26) | 0 = sK0 | 0 = X26 (4:17) [subsumption resolution 921,45]
[SA] new: 956. alength = i(main_end) | $less(X27,i(main_end)) | $less(0,X27) | sK0 = X27 | 0 = sK0 (4:17) [subsumption resolution 922,45]
[SA] passive: 956. alength = i(main_end) | $less(X27,i(main_end)) | $less(0,X27) | sK0 = X27 | 0 = sK0 (4:17) [subsumption resolution 922,45]
[SA] passive: 936. alength = i(main_end) | ~$less(X34,alength) | $less(0,i(main_end)) | $less(X34,sK0) | 0 = X34 | sK0 = X34 (4:20) [resolution 277,380]
[SA] active: 490. v = a(l11(sK2(sK0)),sK0) | alength = i(main_end) (4:11:1) [superposition 70,258]
[SA] new: 957. v = a(l8(sK2(sK0)),sK0) | alength = i(main_end) | i(l8(sK2(sK0))) = sK0 (5:17) [superposition 490,72]
[SA] new: 958. v = a(l8(sK2(sK0)),sK0) | i(l8(sK2(sK0))) = sK0 | alength = i(main_end) (5:17) [superposition 72,490]
[SA] active: 314. $less(i(l8(p(p(nl8)))),alength) | zero = p(nl8) | zero = nl8 (3:14:1) [resolution 164,56]
[SA] new: 959. zero = p(nl8) | zero = nl8 | alength = i(main_end) | ~$less(alength,alength) | $less(i(l8(p(p(nl8)))),i(main_end)) (4:22) [resolution 314,277]
[SA] new: 960. zero = p(nl8) | zero = nl8 | alength = alength | $less(sK0,alength) | $less(i(l8(p(p(nl8)))),alength) (4:20) [resolution 314,254]
[SA] new: 961. zero = p(nl8) | zero = nl8 | alength = sK0 | $less(sK0,alength) | $less(i(l8(p(p(nl8)))),alength) (4:20) [resolution 314,205]
[SA] new: 962. zero = p(nl8) | zero = nl8 | alength = X0 | $less(X0,alength) | $less(i(l8(p(p(nl8)))),X0) (4:20) [resolution 314,99]
[SA] new: 963. zero = p(nl8) | zero = nl8 | ~$less(X1,i(l8(p(p(nl8))))) | $less(X1,alength) (4:17) [resolution 314,14]
[SA] passive: 962. zero = p(nl8) | zero = nl8 | alength = X0 | $less(X0,alength) | $less(i(l8(p(p(nl8)))),X0) (4:20) [resolution 314,99]
[SA] passive: 963. zero = p(nl8) | zero = nl8 | ~$less(X1,i(l8(p(p(nl8))))) | $less(X1,alength) (4:17) [resolution 314,14]
[SA] active: 492. 0 != sK0 | zero = sK2(sK0) | alength = i(main_end) (4:11:1) [superposition 108,258]
[SA] active: 389. $less(X21,i(main_end)) | $less(alength,X21) | alength = X21 | i(main_end) = X21 (3:14:1) [resolution 174,82]
[SA] new: 964. $less(alength,X0) | alength = X0 | i(main_end) = X0 | $less(X1,0) | i(main_end) = X1 | i(l8(sK2(X1))) = X1 | $less(X0,X1) (4:26) [resolution 389,215]
[SA] new: 965. $less(alength,X2) | alength = X2 | i(main_end) = X2 | $less(X2,0) | i(l8(sK2(X2))) = X2 (4:19) [resolution 389,84]
[SA] new: 966. $less(alength,X3) | alength = X3 | i(main_end) = X3 | $less(X3,0) | 'Sub'(sK2(X3),nl8) (4:17) [resolution 389,83]
[SA] new: 967. $less(alength,i(main_end)) | alength = i(main_end) | i(main_end) = i(main_end) (4:13) [resolution 389,13]
[SA] new: 968. $less(alength,X4) | alength = X4 | i(main_end) = X4 | alength = i(main_end) | ~$less(i(main_end),alength) | $less(X4,i(main_end)) (4:22) [resolution 389,277]
[SA] new: 969. $less(alength,X5) | alength = X5 | i(main_end) = X5 | alength = i(main_end) | $less(sK0,i(main_end)) | $less(X5,alength) (4:21) [resolution 389,254]
[SA] new: 970. $less(alength,X6) | alength = X6 | i(main_end) = X6 | i(main_end) = sK0 | $less(sK0,i(main_end)) | $less(X6,alength) (4:21) [resolution 389,205]
[SA] new: 971. $less(alength,X7) | alength = X7 | i(main_end) = X7 | i(main_end) = X8 | $less(X8,i(main_end)) | $less(X7,X8) (4:21) [resolution 389,99]
[SA] new: 972. $less(alength,X9) | alength = X9 | i(main_end) = X9 | ~$less(X10,X9) | $less(X10,i(main_end)) (4:17) [resolution 389,14]
[SA] passive: 964. $less(alength,X0) | alength = X0 | i(main_end) = X0 | $less(X1,0) | i(main_end) = X1 | i(l8(sK2(X1))) = X1 | $less(X0,X1) (4:26) [resolution 389,215]
[SA] passive: 966. $less(alength,X3) | alength = X3 | i(main_end) = X3 | $less(X3,0) | 'Sub'(sK2(X3),nl8) (4:17) [resolution 389,83]
[SA] passive: 971. $less(alength,X7) | alength = X7 | i(main_end) = X7 | i(main_end) = X8 | $less(X8,i(main_end)) | $less(X7,X8) (4:21) [resolution 389,99]
[SA] passive: 972. $less(alength,X9) | alength = X9 | i(main_end) = X9 | ~$less(X10,X9) | $less(X10,i(main_end)) (4:17) [resolution 389,14]
[SA] active: 583. $less(0,i(main_end)) | alength = i(main_end) | 0 = sK0 (4:11:1) [resolution 178,82]
[SA] new: 973. alength = i(main_end) | 0 = sK0 | $less(X0,0) | i(main_end) = X0 | i(l8(sK2(X0))) = X0 | $less(0,X0) (5:23) [resolution 583,215]
[SA] new: 974. alength = i(main_end) | 0 = sK0 | $less(0,0) | i(l8(sK2(0))) = 0 (5:16) [resolution 583,84]
[SA] new: 975. alength = i(main_end) | 0 = sK0 | $less(0,0) | 'Sub'(sK2(0),nl8) (5:14) [resolution 583,83]
[SA] new: 976. alength = i(main_end) | 0 = sK0 | alength = i(main_end) | ~$less(i(main_end),alength) | $less(0,i(main_end)) (5:19) [resolution 583,277]
[SA] new: 977. alength = i(main_end) | 0 = sK0 | alength = i(main_end) | $less(sK0,i(main_end)) | $less(0,alength) (5:18) [resolution 583,254]
[SA] new: 978. alength = i(main_end) | 0 = sK0 | i(main_end) = sK0 | $less(sK0,i(main_end)) | $less(0,alength) (5:18) [resolution 583,205]
[SA] new: 979. alength = i(main_end) | 0 = sK0 | i(main_end) = X1 | $less(X1,i(main_end)) | $less(0,X1) (5:18) [resolution 583,99]
[SA] new: 980. alength = i(main_end) | 0 = sK0 | ~$less(X2,0) | $less(X2,i(main_end)) (5:14) [resolution 583,14]
[SA] new: 981. alength = i(main_end) | 0 = sK0 | $less(sK0,i(main_end)) | $less(0,alength) (5:14) [duplicate literal removal 977]
[SA] new: 982. alength = i(main_end) | 0 = sK0 | ~$less(i(main_end),alength) | $less(0,i(main_end)) (5:15) [duplicate literal removal 976]
[SA] new: 983. alength = i(main_end) | 0 = sK0 | 'Sub'(sK2(0),nl8) (5:11) [evaluation 975]
[SA] new: 984. alength = i(main_end) | 0 = sK0 | i(l8(sK2(0))) = 0 (5:13) [evaluation 974]
[SA] passive: 979. alength = i(main_end) | 0 = sK0 | i(main_end) = X1 | $less(X1,i(main_end)) | $less(0,X1) (5:18) [resolution 583,99]
[SA] passive: 980. alength = i(main_end) | 0 = sK0 | ~$less(X2,0) | $less(X2,i(main_end)) (5:14) [resolution 583,14]
[SA] active: 409. $less(X21,i(main_end)) | $less(alength,X21) | alength = i(main_end) | alength = X21 (3:14:1) [resolution 174,82]
[SA] new: 985. $less(alength,X0) | alength = i(main_end) | alength = X0 | $less(X1,0) | i(main_end) = X1 | i(l8(sK2(X1))) = X1 | $less(X0,X1) (4:26) [resolution 409,215]
[SA] new: 986. $less(alength,X2) | alength = i(main_end) | alength = X2 | $less(X2,0) | i(l8(sK2(X2))) = X2 (4:19) [resolution 409,84]
[SA] new: 987. $less(alength,X3) | alength = i(main_end) | alength = X3 | $less(X3,0) | 'Sub'(sK2(X3),nl8) (4:17) [resolution 409,83]
[SA] new: 988. $less(alength,i(main_end)) | alength = i(main_end) | alength = i(main_end) (4:12) [resolution 409,13]
[SA] new: 989. $less(alength,X4) | alength = i(main_end) | alength = X4 | alength = i(main_end) | ~$less(i(main_end),alength) | $less(X4,i(main_end)) (4:22) [resolution 409,277]
[SA] new: 990. $less(alength,X5) | alength = i(main_end) | alength = X5 | alength = i(main_end) | $less(sK0,i(main_end)) | $less(X5,alength) (4:21) [resolution 409,254]
[SA] new: 991. $less(alength,X6) | alength = i(main_end) | alength = X6 | i(main_end) = sK0 | $less(sK0,i(main_end)) | $less(X6,alength) (4:21) [resolution 409,205]
[SA] new: 992. $less(alength,X7) | alength = i(main_end) | alength = X7 | i(main_end) = X8 | $less(X8,i(main_end)) | $less(X7,X8) (4:21) [resolution 409,99]
[SA] new: 993. $less(alength,X9) | alength = i(main_end) | alength = X9 | ~$less(X10,X9) | $less(X10,i(main_end)) (4:17) [resolution 409,14]
[SA] new: 994. $less(alength,X5) | alength = i(main_end) | alength = X5 | $less(sK0,i(main_end)) | $less(X5,alength) (4:17) [duplicate literal removal 990]
[SA] new: 995. $less(alength,X4) | alength = i(main_end) | alength = X4 | ~$less(i(main_end),alength) | $less(X4,i(main_end)) (4:18) [duplicate literal removal 989]
[SA] new: 996. $less(alength,i(main_end)) | alength = i(main_end) (4:8) [duplicate literal removal 988]
[SA] passive: 985. $less(alength,X0) | alength = i(main_end) | alength = X0 | $less(X1,0) | i(main_end) = X1 | i(l8(sK2(X1))) = X1 | $less(X0,X1) (4:26) [resolution 409,215]
[SA] passive: 986. $less(alength,X2) | alength = i(main_end) | alength = X2 | $less(X2,0) | i(l8(sK2(X2))) = X2 (4:19) [resolution 409,84]
[SA] passive: 987. $less(alength,X3) | alength = i(main_end) | alength = X3 | $less(X3,0) | 'Sub'(sK2(X3),nl8) (4:17) [resolution 409,83]
[SA] passive: 992. $less(alength,X7) | alength = i(main_end) | alength = X7 | i(main_end) = X8 | $less(X8,i(main_end)) | $less(X7,X8) (4:21) [resolution 409,99]
[SA] passive: 993. $less(alength,X9) | alength = i(main_end) | alength = X9 | ~$less(X10,X9) | $less(X10,i(main_end)) (4:17) [resolution 409,14]
[SA] active: 591. 'Sub'(sK2(0),nl8) | 0 = sK0 | alength = i(main_end) (4:11:1) [subsumption resolution 587,82]
[SA] new: 997. 0 = sK0 | alength = i(main_end) | a(main_end,X0) = a(l11(sK2(0)),X0) | i(l8(sK1(X0,sK2(0)))) = X0 (5:24) [resolution 591,150]
[SA] new: 998. 0 = sK0 | alength = i(main_end) | $less(i(l8(sK2(0))),alength) (5:13) [resolution 591,56]
[SA] passive: 997. 0 = sK0 | alength = i(main_end) | a(main_end,X0) = a(l11(sK2(0)),X0) | i(l8(sK1(X0,sK2(0)))) = X0 (5:24) [resolution 591,150]
[SA] passive: 998. 0 = sK0 | alength = i(main_end) | $less(i(l8(sK2(0))),alength) (5:13) [resolution 591,56]
[SA] active: 193. $less(i(l8(p(nl8))),X0) | alength = X0 | $less(X0,alength) | zero = nl8 (3:15:1) [resolution 95,99]
[SA] new: 999. alength = i(l8(p(nl8))) | $less(i(l8(p(nl8))),alength) | zero = nl8 (4:15) [resolution 193,13]
[SA] new: 1000. alength = X0 | $less(X0,alength) | zero = nl8 | alength = i(main_end) | ~$less(X0,alength) | $less(i(l8(p(nl8))),i(main_end)) (4:23) [resolution 193,277]
[SA] new: 1001. alength = X1 | $less(X1,alength) | zero = nl8 | alength = X1 | $less(sK0,X1) | $less(i(l8(p(nl8))),alength) (4:21) [resolution 193,254]
[SA] new: 1002. alength = X2 | $less(X2,alength) | zero = nl8 | sK0 = X2 | $less(sK0,X2) | $less(i(l8(p(nl8))),alength) (4:21) [resolution 193,205]
[SA] new: 1003. alength = X3 | $less(X3,alength) | zero = nl8 | X3 = X4 | $less(X4,X3) | $less(i(l8(p(nl8))),X4) (4:21) [resolution 193,99]
[SA] new: 1004. alength = X5 | $less(X5,alength) | zero = nl8 | ~$less(X6,i(l8(p(nl8)))) | $less(X6,X5) (4:18) [resolution 193,14]
[SA] new: 1005. alength = i(main_end) | $less(i(main_end),alength) | zero = nl8 | $less(X7,0) | i(main_end) = X7 | i(l8(sK2(X7))) = X7 | $less(i(l8(p(nl8))),X7) (4:30) [resolution 193,215]
[SA] new: 1006. alength = i(main_end) | $less(i(main_end),alength) | zero = nl8 | $less(i(l8(p(nl8))),0) | i(l8(p(nl8))) = i(l8(sK2(i(l8(p(nl8)))))) (4:29) [resolution 193,84]
[SA] new: 1007. alength = i(main_end) | $less(i(main_end),alength) | zero = nl8 | $less(i(l8(p(nl8))),0) | 'Sub'(sK2(i(l8(p(nl8)))),nl8) (4:24) [resolution 193,83]
[SA] new: 1008. alength = i(l8(p(nl8))) | $less(i(l8(p(nl8))),alength) | zero = nl8 | zero = nl8 | $less(i(l8(p(nl8))),alength) (4:24) [resolution 193,194]
[SA] new: 1009. alength = 0 | $less(0,alength) | zero = nl8 | alength = 0 | $less(i(l8(p(nl8))),alength) (5:18) [resolution 193,262]
[SA] new: 1010. alength = 0 | $less(0,alength) | zero = nl8 | 0 = sK0 | $less(i(l8(p(nl8))),alength) (4:18) [resolution 193,116]
[SA] new: 1011. alength = 0 | $less(0,alength) | zero = nl8 | 0 = sK0 | $less(i(l8(p(nl8))),sK0) (4:18) [resolution 193,114]
[SA] new: 1012. alength = sK0 | $less(sK0,alength) | zero = nl8 | alength = X8 | $less(X8,alength) | $less(i(l8(p(nl8))),X8) (4:21) [resolution 193,247]
[SA] new: 1013. alength = sK0 | $less(sK0,alength) | zero = nl8 | sK0 = X9 | $less(X9,alength) | $less(i(l8(p(nl8))),X9) (4:21) [resolution 193,199]
[SA] new: 1014. alength = sK0 | $less(sK0,alength) | zero = nl8 | $less(i(l8(p(nl8))),alength) (4:15) [resolution 193,96]
[SA] new: 1015. alength = 0 | $less(0,alength) | zero = nl8 | $less(i(l8(p(nl8))),alength) (5:15) [duplicate literal removal 1009]
[SA] new: 1016. alength = i(l8(p(nl8))) | $less(i(l8(p(nl8))),alength) | zero = nl8 (4:15) [duplicate literal removal 1008]
[SA] new: 1017. alength = X1 | $less(X1,alength) | zero = nl8 | $less(sK0,X1) | $less(i(l8(p(nl8))),alength) (4:18) [duplicate literal removal 1001]
[SA] passive: 1003. alength = X3 | $less(X3,alength) | zero = nl8 | X3 = X4 | $less(X4,X3) | $less(i(l8(p(nl8))),X4) (4:21) [resolution 193,99]
[SA] passive: 1004. alength = X5 | $less(X5,alength) | zero = nl8 | ~$less(X6,i(l8(p(nl8)))) | $less(X6,X5) (4:18) [resolution 193,14]
[SA] new: 1018. alength = i(main_end) | zero = nl8 | $less(X7,0) | i(main_end) = X7 | i(l8(sK2(X7))) = X7 | $less(i(l8(p(nl8))),X7) (4:26) [subsumption resolution 1005,82]
[SA] passive: 1018. alength = i(main_end) | zero = nl8 | $less(X7,0) | i(main_end) = X7 | i(l8(sK2(X7))) = X7 | $less(i(l8(p(nl8))),X7) (4:26) [subsumption resolution 1005,82]
[SA] new: 1019. alength = i(main_end) | zero = nl8 | $less(i(l8(p(nl8))),0) | i(l8(p(nl8))) = i(l8(sK2(i(l8(p(nl8)))))) (4:25) [subsumption resolution 1006,82]
[SA] passive: 1019. alength = i(main_end) | zero = nl8 | $less(i(l8(p(nl8))),0) | i(l8(p(nl8))) = i(l8(sK2(i(l8(p(nl8)))))) (4:25) [subsumption resolution 1006,82]
[SA] new: 1020. alength = i(main_end) | zero = nl8 | $less(i(l8(p(nl8))),0) | 'Sub'(sK2(i(l8(p(nl8)))),nl8) (4:20) [subsumption resolution 1007,82]
[SA] passive: 1020. alength = i(main_end) | zero = nl8 | $less(i(l8(p(nl8))),0) | 'Sub'(sK2(i(l8(p(nl8)))),nl8) (4:20) [subsumption resolution 1007,82]
[SA] active: 848. 'Sub'(sK2(0),nl8) | i(main_end) = sK0 | 0 = sK0 (4:11:1) [subsumption resolution 842,82]
[SA] new: 1021. i(main_end) = sK0 | 0 = sK0 | a(main_end,X0) = a(l11(sK2(0)),X0) | i(l8(sK1(X0,sK2(0)))) = X0 (5:24) [resolution 848,150]
[SA] new: 1022. i(main_end) = sK0 | 0 = sK0 | $less(i(l8(sK2(0))),alength) (5:13) [resolution 848,56]
[SA] passive: 1021. i(main_end) = sK0 | 0 = sK0 | a(main_end,X0) = a(l11(sK2(0)),X0) | i(l8(sK1(X0,sK2(0)))) = X0 (5:24) [resolution 848,150]
[SA] passive: 1022. i(main_end) = sK0 | 0 = sK0 | $less(i(l8(sK2(0))),alength) (5:13) [resolution 848,56]
[SA] active: 198. $less(sK0,X1) | $less(X1,X0) | $less(X0,alength) | X0 = X1 | sK0 = X0 (3:15:3) [resolution 106,99]
[SA] new: 1023. $less(0,X0) | $less(X0,alength) | 0 = X0 | sK0 = X0 (4:12) [resolution 198,46]
[SA] new: 1024. $less(sK0,X1) | $less(X1,alength) | sK0 = X1 | sK0 = X1 (4:12) [resolution 198,13]
[SA] new: 1025. $less(X2,X3) | $less(X3,alength) | X2 = X3 | sK0 = X3 | alength = i(main_end) | ~$less(X2,alength) | $less(sK0,i(main_end)) (4:23) [resolution 198,277]
[SA] new: 1026. $less(X4,X5) | $less(X5,alength) | X4 = X5 | sK0 = X5 | alength = X4 | $less(sK0,X4) | $less(sK0,alength) (4:21) [resolution 198,254]
[SA] new: 1027. $less(X6,X7) | $less(X7,alength) | X6 = X7 | sK0 = X7 | sK0 = X6 | $less(sK0,X6) | $less(sK0,alength) (4:21) [resolution 198,205]
[SA] new: 1028. $less(X8,X9) | $less(X9,alength) | X8 = X9 | sK0 = X9 | X8 = X10 | $less(X10,X8) | $less(sK0,X10) (4:21) [resolution 198,99]
[SA] new: 1029. $less(X11,X12) | $less(X12,alength) | X11 = X12 | sK0 = X12 | ~$less(X13,sK0) | $less(X13,X11) (4:18) [resolution 198,14]
[SA] new: 1030. $less(i(main_end),X14) | $less(X14,alength) | i(main_end) = X14 | sK0 = X14 | $less(X15,0) | i(main_end) = X15 | i(l8(sK2(X15))) = X15 | $less(sK0,X15) (4:30) [resolution 198,215]
[SA] new: 1031. $less(i(main_end),X16) | $less(X16,alength) | i(main_end) = X16 | sK0 = X16 | $less(sK0,0) | i(l8(sK2(sK0))) = sK0 (4:23) [resolution 198,84]
[SA] new: 1032. $less(i(main_end),X17) | $less(X17,alength) | i(main_end) = X17 | sK0 = X17 | $less(sK0,0) | 'Sub'(sK2(sK0),nl8) (4:21) [resolution 198,83]
[SA] new: 1033. $less(i(l8(p(nl8))),X18) | $less(X18,alength) | i(l8(p(nl8))) = X18 | sK0 = X18 | zero = nl8 | $less(sK0,alength) (4:24) [resolution 198,194]
[SA] new: 1034. $less(0,X19) | $less(X19,alength) | 0 = X19 | sK0 = X19 | alength = 0 | $less(sK0,alength) (5:18) [resolution 198,262]
[SA] new: 1035. $less(0,X20) | $less(X20,alength) | 0 = X20 | sK0 = X20 | 0 = sK0 | $less(sK0,alength) (4:18) [resolution 198,116]
[SA] new: 1036. $less(0,X21) | $less(X21,alength) | 0 = X21 | sK0 = X21 | 0 = sK0 | $less(sK0,sK0) (4:18) [resolution 198,114]
[SA] new: 1037. $less(sK0,X22) | $less(X22,alength) | sK0 = X22 | sK0 = X22 | alength = X23 | $less(X23,alength) | $less(sK0,X23) (4:21) [resolution 198,247]
[SA] new: 1038. $less(sK0,X24) | $less(X24,alength) | sK0 = X24 | sK0 = X24 | sK0 = X25 | $less(X25,alength) | $less(sK0,X25) (4:21) [resolution 198,199]
[SA] new: 1039. $less(sK0,X26) | $less(X26,alength) | sK0 = X26 | sK0 = X26 | $less(sK0,alength) (4:15) [resolution 198,96]
[SA] new: 1040. $less(sK0,sK0) | $less(0,alength) | 0 = sK0 | 0 = sK0 (4:12) [resolution 198,46]
[SA] new: 1041. $less(sK0,X0) | $less(X0,alength) | X0 = X0 | sK0 = X0 (4:12) [resolution 198,13]
[SA] new: 1042. $less(sK0,X1) | $less(X2,alength) | X1 = X2 | sK0 = X2 | alength = i(main_end) | ~$less(X2,alength) | $less(X1,i(main_end)) (4:23) [resolution 198,277]
[SA] new: 1043. $less(sK0,X3) | $less(X4,alength) | X3 = X4 | sK0 = X4 | alength = X4 | $less(sK0,X4) | $less(X3,alength) (4:21) [resolution 198,254]
[SA] new: 1044. $less(sK0,X5) | $less(X6,alength) | X5 = X6 | sK0 = X6 | sK0 = X6 | $less(sK0,X6) | $less(X5,alength) (4:21) [resolution 198,205]
[SA] new: 1045. $less(sK0,X7) | $less(X8,alength) | X7 = X8 | sK0 = X8 | X8 = X9 | $less(X9,X8) | $less(X7,X9) (4:21) [resolution 198,99]
[SA] new: 1046. $less(sK0,X10) | $less(X11,alength) | X10 = X11 | sK0 = X11 | ~$less(X12,X10) | $less(X12,X11) (4:18) [resolution 198,14]
[SA] new: 1047. $less(sK0,X13) | $less(i(main_end),alength) | i(main_end) = X13 | i(main_end) = sK0 | $less(X14,0) | i(main_end) = X14 | i(l8(sK2(X14))) = X14 | $less(X13,X14) (4:31) [resolution 198,215]
[SA] new: 1048. $less(sK0,X15) | $less(i(main_end),alength) | i(main_end) = X15 | i(main_end) = sK0 | $less(X15,0) | i(l8(sK2(X15))) = X15 (4:24) [resolution 198,84]
[SA] new: 1049. $less(sK0,X16) | $less(i(main_end),alength) | i(main_end) = X16 | i(main_end) = sK0 | $less(X16,0) | 'Sub'(sK2(X16),nl8) (4:22) [resolution 198,83]
[SA] new: 1050. $less(sK0,X17) | $less(i(l8(p(nl8))),alength) | i(l8(p(nl8))) = X17 | i(l8(p(nl8))) = sK0 | zero = nl8 | $less(X17,alength) (4:27) [resolution 198,194]
[SA] new: 1051. $less(sK0,X18) | $less(0,alength) | 0 = X18 | 0 = sK0 | alength = 0 | $less(X18,alength) (5:18) [resolution 198,262]
[SA] new: 1052. $less(sK0,X19) | $less(0,alength) | 0 = X19 | 0 = sK0 | 0 = sK0 | $less(X19,alength) (4:18) [resolution 198,116]
[SA] new: 1053. $less(sK0,X20) | $less(0,alength) | 0 = X20 | 0 = sK0 | 0 = sK0 | $less(X20,sK0) (4:18) [resolution 198,114]
[SA] new: 1054. $less(sK0,X21) | $less(sK0,alength) | sK0 = X21 | sK0 = sK0 | alength = X22 | $less(X22,alength) | $less(X21,X22) (4:21) [resolution 198,247]
[SA] new: 1055. $less(sK0,X23) | $less(sK0,alength) | sK0 = X23 | sK0 = sK0 | sK0 = X24 | $less(X24,alength) | $less(X23,X24) (4:21) [resolution 198,199]
[SA] new: 1056. $less(sK0,X25) | $less(sK0,alength) | sK0 = X25 | sK0 = sK0 | $less(X25,alength) (4:15) [resolution 198,96]
[SA] new: 1057. $less(sK0,i(main_end)) | $less(alength,alength) | alength = i(main_end) | alength = sK0 (4:14) [resolution 198,82]
[SA] new: 1058. $less(sK0,i(main_end)) | $less(sK0,alength) | i(main_end) = sK0 | sK0 = sK0 | alength = i(main_end) (6:18) [resolution 198,812]
[SA] new: 1059. $less(sK0,i(main_end)) | $less(sK0,alength) | i(main_end) = sK0 | sK0 = sK0 | i(main_end) = sK0 (6:18) [resolution 198,722]
[SA] new: 1060. $less(sK0,X0) | $less(X0,alength) | alength = X0 | alength = sK0 (4:12) [resolution 198,13]
[SA] new: 1061. $less(sK0,X1) | $less(X1,X2) | X1 = X2 | sK0 = X2 | alength = i(main_end) | ~$less(alength,alength) | $less(X2,i(main_end)) (4:23) [resolution 198,277]
[SA] new: 1062. $less(sK0,X3) | $less(X3,X4) | X3 = X4 | sK0 = X4 | alength = alength | $less(sK0,alength) | $less(X4,alength) (4:21) [resolution 198,254]
[SA] new: 1063. $less(sK0,X5) | $less(X5,X6) | X5 = X6 | sK0 = X6 | alength = sK0 | $less(sK0,alength) | $less(X6,alength) (4:21) [resolution 198,205]
[SA] new: 1064. $less(sK0,X7) | $less(X7,X8) | X7 = X8 | sK0 = X8 | alength = X9 | $less(X9,alength) | $less(X8,X9) (4:21) [resolution 198,99]
[SA] new: 1065. $less(sK0,X10) | $less(X10,X11) | X10 = X11 | sK0 = X11 | ~$less(X12,X11) | $less(X12,alength) (4:18) [resolution 198,14]
[SA] new: 1066. $less(sK0,X13) | $less(X13,i(main_end)) | i(main_end) = X13 | i(main_end) = sK0 (4:15) [resolution 198,82]
[SA] new: 1067. $less(sK0,sK0) | $less(sK0,alength) | sK0 = sK0 | sK0 = sK0 (4:12) [factoring 198]
[SA] new: 1068. $less(sK0,alength) | $less(alength,sK0) | alength = sK0 | sK0 = sK0 (4:12) [factoring 198]
[SA] new: 1069. $less(sK0,alength) | $less(alength,alength) | alength = alength | alength = sK0 (4:12) [factoring 198]
[SA] new: 1070. $less(sK0,sK0) | $less(sK0,alength) | sK0 = sK0 (4:9) [duplicate literal removal 1067]
[SA] new: 1071. $less(sK0,i(main_end)) | $less(sK0,alength) | i(main_end) = sK0 | sK0 = sK0 (6:14) [duplicate literal removal 1059]
[SA] new: 1072. $less(sK0,X20) | $less(0,alength) | 0 = X20 | 0 = sK0 | $less(X20,sK0) (4:15) [duplicate literal removal 1053]
[SA] new: 1073. $less(sK0,X19) | $less(0,alength) | 0 = X19 | 0 = sK0 | $less(X19,alength) (4:15) [duplicate literal removal 1052]
[SA] new: 1074. $less(sK0,X5) | $less(X6,alength) | X5 = X6 | sK0 = X6 | $less(sK0,X6) | $less(X5,alength) (4:18) [duplicate literal removal 1044]
[SA] new: 1075. $less(sK0,sK0) | $less(0,alength) | 0 = sK0 (4:9) [duplicate literal removal 1040]
[SA] new: 1076. $less(sK0,X26) | $less(X26,alength) | sK0 = X26 | $less(sK0,alength) (4:12) [duplicate literal removal 1039]
[SA] new: 1077. $less(sK0,X24) | $less(X24,alength) | sK0 = X24 | sK0 = X25 | $less(X25,alength) | $less(sK0,X25) (4:18) [duplicate literal removal 1038]
[SA] new: 1078. $less(sK0,X22) | $less(X22,alength) | sK0 = X22 | alength = X23 | $less(X23,alength) | $less(sK0,X23) (4:18) [duplicate literal removal 1037]
[SA] new: 1079. $less(sK0,X1) | $less(X1,alength) | sK0 = X1 (4:9) [duplicate literal removal 1024]
[SA] passive: 1028. $less(X8,X9) | $less(X9,alength) | X8 = X9 | sK0 = X9 | X8 = X10 | $less(X10,X8) | $less(sK0,X10) (4:21) [resolution 198,99]
[SA] passive: 1029. $less(X11,X12) | $less(X12,alength) | X11 = X12 | sK0 = X12 | ~$less(X13,sK0) | $less(X13,X11) (4:18) [resolution 198,14]
[SA] passive: 1030. $less(i(main_end),X14) | $less(X14,alength) | i(main_end) = X14 | sK0 = X14 | $less(X15,0) | i(main_end) = X15 | i(l8(sK2(X15))) = X15 | $less(sK0,X15) (4:30) [resolution 198,215]
[SA] new: 1080. $less(i(main_end),X16) | $less(X16,alength) | i(main_end) = X16 | sK0 = X16 | i(l8(sK2(sK0))) = sK0 (4:20) [subsumption resolution 1031,46]
[SA] passive: 1080. $less(i(main_end),X16) | $less(X16,alength) | i(main_end) = X16 | sK0 = X16 | i(l8(sK2(sK0))) = sK0 (4:20) [subsumption resolution 1031,46]
[SA] new: 1081. $less(i(main_end),X17) | $less(X17,alength) | i(main_end) = X17 | sK0 = X17 | 'Sub'(sK2(sK0),nl8) (4:18) [subsumption resolution 1032,46]
[SA] passive: 1081. $less(i(main_end),X17) | $less(X17,alength) | i(main_end) = X17 | sK0 = X17 | 'Sub'(sK2(sK0),nl8) (4:18) [subsumption resolution 1032,46]
[SA] passive: 1045. $less(sK0,X7) | $less(X8,alength) | X7 = X8 | sK0 = X8 | X8 = X9 | $less(X9,X8) | $less(X7,X9) (4:21) [resolution 198,99]
[SA] passive: 1046. $less(sK0,X10) | $less(X11,alength) | X10 = X11 | sK0 = X11 | ~$less(X12,X10) | $less(X12,X11) (4:18) [resolution 198,14]
[SA] new: 1082. $less(sK0,X13) | i(main_end) = X13 | i(main_end) = sK0 | $less(X14,0) | i(main_end) = X14 | i(l8(sK2(X14))) = X14 | $less(X13,X14) (4:27) [subsumption resolution 1047,82]
[SA] passive: 1082. $less(sK0,X13) | i(main_end) = X13 | i(main_end) = sK0 | $less(X14,0) | i(main_end) = X14 | i(l8(sK2(X14))) = X14 | $less(X13,X14) (4:27) [subsumption resolution 1047,82]
[SA] new: 1083. $less(sK0,X16) | i(main_end) = X16 | i(main_end) = sK0 | $less(X16,0) | 'Sub'(sK2(X16),nl8) (4:18) [subsumption resolution 1049,82]
[SA] passive: 1083. $less(sK0,X16) | i(main_end) = X16 | i(main_end) = sK0 | $less(X16,0) | 'Sub'(sK2(X16),nl8) (4:18) [subsumption resolution 1049,82]
[SA] passive: 1065. $less(sK0,X10) | $less(X10,X11) | X10 = X11 | sK0 = X11 | ~$less(X12,X11) | $less(X12,alength) (4:18) [resolution 198,14]
[SA] active: 306. i(main_end) != 4 | s(s(s(s(zero)))) = nl8 (5:11:1) [superposition 109,186]
[SA] active: 204. $less(sK0,X0) | $less(X1,alength) | $less(X0,X1) | alength = X1 | sK0 = X0 (3:15:3) [resolution 106,99]
[SA] new: 1084. $less(X0,alength) | $less(0,X0) | alength = X0 | 0 = sK0 (4:12) [resolution 204,46]
[SA] new: 1085. $less(X1,alength) | $less(sK0,X1) | alength = X1 | sK0 = sK0 (4:12) [resolution 204,13]
[SA] new: 1086. $less(X2,alength) | $less(X3,X2) | alength = X2 | sK0 = X3 | alength = i(main_end) | ~$less(X3,alength) | $less(sK0,i(main_end)) (4:23) [resolution 204,277]
[SA] new: 1087. $less(X4,alength) | $less(X5,X4) | alength = X4 | sK0 = X5 | alength = X5 | $less(sK0,X5) | $less(sK0,alength) (4:21) [resolution 204,254]
[SA] new: 1088. $less(X6,alength) | $less(X7,X6) | alength = X6 | sK0 = X7 | sK0 = X7 | $less(sK0,X7) | $less(sK0,alength) (4:21) [resolution 204,205]
[SA] new: 1089. $less(X8,alength) | $less(X9,X8) | alength = X8 | sK0 = X9 | X9 = X10 | $less(X10,X9) | $less(sK0,X10) (4:21) [resolution 204,99]
[SA] new: 1090. $less(X11,alength) | $less(X12,X11) | alength = X11 | sK0 = X12 | ~$less(X13,sK0) | $less(X13,X12) (4:18) [resolution 204,14]
[SA] new: 1091. $less(X14,alength) | $less(i(main_end),X14) | alength = X14 | i(main_end) = sK0 | $less(X15,0) | i(main_end) = X15 | i(l8(sK2(X15))) = X15 | $less(sK0,X15) (4:30) [resolution 204,215]
[SA] new: 1092. $less(X16,alength) | $less(i(main_end),X16) | alength = X16 | i(main_end) = sK0 | $less(sK0,0) | i(l8(sK2(sK0))) = sK0 (4:23) [resolution 204,84]
[SA] new: 1093. $less(X17,alength) | $less(i(main_end),X17) | alength = X17 | i(main_end) = sK0 | $less(sK0,0) | 'Sub'(sK2(sK0),nl8) (4:21) [resolution 204,83]
[SA] new: 1094. $less(X18,alength) | $less(i(l8(p(nl8))),X18) | alength = X18 | i(l8(p(nl8))) = sK0 | zero = nl8 | $less(sK0,alength) (4:24) [resolution 204,194]
[SA] new: 1095. $less(X19,alength) | $less(0,X19) | alength = X19 | 0 = sK0 | alength = 0 | $less(sK0,alength) (5:18) [resolution 204,262]
[SA] new: 1096. $less(X20,alength) | $less(0,X20) | alength = X20 | 0 = sK0 | 0 = sK0 | $less(sK0,alength) (4:18) [resolution 204,116]
[SA] new: 1097. $less(X21,alength) | $less(0,X21) | alength = X21 | 0 = sK0 | 0 = sK0 | $less(sK0,sK0) (4:18) [resolution 204,114]
[SA] new: 1098. $less(X22,alength) | $less(sK0,X22) | alength = X22 | sK0 = sK0 | alength = X23 | $less(X23,alength) | $less(sK0,X23) (4:21) [resolution 204,247]
[SA] new: 1099. $less(X24,alength) | $less(sK0,X24) | alength = X24 | sK0 = sK0 | sK0 = X25 | $less(X25,alength) | $less(sK0,X25) (4:21) [resolution 204,199]
[SA] new: 1100. $less(X26,alength) | $less(sK0,X26) | alength = X26 | sK0 = sK0 | $less(sK0,alength) (4:15) [resolution 204,96]
[SA] new: 1101. $less(sK0,X0) | $less(X0,alength) | alength = alength | sK0 = X0 (4:12) [resolution 204,13]
[SA] new: 1102. $less(sK0,X1) | $less(X1,X2) | alength = X2 | sK0 = X1 | alength = i(main_end) | ~$less(alength,alength) | $less(X2,i(main_end)) (4:23) [resolution 204,277]
[SA] new: 1103. $less(sK0,X3) | $less(X3,X4) | alength = X4 | sK0 = X3 | alength = alength | $less(sK0,alength) | $less(X4,alength) (4:21) [resolution 204,254]
[SA] new: 1104. $less(sK0,X5) | $less(X5,X6) | alength = X6 | sK0 = X5 | alength = sK0 | $less(sK0,alength) | $less(X6,alength) (4:21) [resolution 204,205]
[SA] new: 1105. $less(sK0,X7) | $less(X7,X8) | alength = X8 | sK0 = X7 | alength = X9 | $less(X9,alength) | $less(X8,X9) (4:21) [resolution 204,99]
[SA] new: 1106. $less(sK0,X10) | $less(X10,X11) | alength = X11 | sK0 = X10 | ~$less(X12,X11) | $less(X12,alength) (4:18) [resolution 204,14]
[SA] new: 1107. $less(sK0,X13) | $less(X13,i(main_end)) | alength = i(main_end) | sK0 = X13 (4:14) [resolution 204,82]
[SA] new: 1108. $less(sK0,sK0) | $less(0,alength) | alength = 0 | sK0 = sK0 (4:12) [resolution 204,46]
[SA] new: 1109. $less(sK0,X0) | $less(X0,alength) | alength = X0 | sK0 = X0 (4:12) [resolution 204,13]
[SA] new: 1110. $less(sK0,X1) | $less(X2,alength) | alength = X2 | sK0 = X1 | alength = i(main_end) | ~$less(X2,alength) | $less(X1,i(main_end)) (4:23) [resolution 204,277]
[SA] new: 1111. $less(sK0,X3) | $less(X4,alength) | alength = X4 | sK0 = X3 | alength = X4 | $less(sK0,X4) | $less(X3,alength) (4:21) [resolution 204,254]
[SA] new: 1112. $less(sK0,X5) | $less(X6,alength) | alength = X6 | sK0 = X5 | sK0 = X6 | $less(sK0,X6) | $less(X5,alength) (4:21) [resolution 204,205]
[SA] new: 1113. $less(sK0,X7) | $less(X8,alength) | alength = X8 | sK0 = X7 | X8 = X9 | $less(X9,X8) | $less(X7,X9) (4:21) [resolution 204,99]
[SA] new: 1114. $less(sK0,X10) | $less(X11,alength) | alength = X11 | sK0 = X10 | ~$less(X12,X10) | $less(X12,X11) (4:18) [resolution 204,14]
[SA] new: 1115. $less(sK0,X13) | $less(i(main_end),alength) | alength = i(main_end) | sK0 = X13 | $less(X14,0) | i(main_end) = X14 | i(l8(sK2(X14))) = X14 | $less(X13,X14) (4:30) [resolution 204,215]
[SA] new: 1116. $less(sK0,X15) | $less(i(main_end),alength) | alength = i(main_end) | sK0 = X15 | $less(X15,0) | i(l8(sK2(X15))) = X15 (4:23) [resolution 204,84]
[SA] new: 1117. $less(sK0,X16) | $less(i(main_end),alength) | alength = i(main_end) | sK0 = X16 | $less(X16,0) | 'Sub'(sK2(X16),nl8) (4:21) [resolution 204,83]
[SA] new: 1118. $less(sK0,X17) | $less(i(l8(p(nl8))),alength) | alength = i(l8(p(nl8))) | sK0 = X17 | zero = nl8 | $less(X17,alength) (4:24) [resolution 204,194]
[SA] new: 1119. $less(sK0,X18) | $less(0,alength) | alength = 0 | sK0 = X18 | alength = 0 | $less(X18,alength) (5:18) [resolution 204,262]
[SA] new: 1120. $less(sK0,X19) | $less(0,alength) | alength = 0 | sK0 = X19 | 0 = sK0 | $less(X19,alength) (4:18) [resolution 204,116]
[SA] new: 1121. $less(sK0,X20) | $less(0,alength) | alength = 0 | sK0 = X20 | 0 = sK0 | $less(X20,sK0) (4:18) [resolution 204,114]
[SA] new: 1122. $less(sK0,X21) | $less(sK0,alength) | alength = sK0 | sK0 = X21 | alength = X22 | $less(X22,alength) | $less(X21,X22) (4:21) [resolution 204,247]
[SA] new: 1123. $less(sK0,X23) | $less(sK0,alength) | alength = sK0 | sK0 = X23 | sK0 = X24 | $less(X24,alength) | $less(X23,X24) (4:21) [resolution 204,199]
[SA] new: 1124. $less(sK0,X25) | $less(sK0,alength) | alength = sK0 | sK0 = X25 | $less(X25,alength) (4:15) [resolution 204,96]
[SA] new: 1125. $less(sK0,i(main_end)) | $less(alength,alength) | alength = alength | i(main_end) = sK0 (4:14) [resolution 204,82]
[SA] new: 1126. $less(sK0,i(main_end)) | $less(sK0,alength) | alength = sK0 | i(main_end) = sK0 | alength = i(main_end) (6:18) [resolution 204,812]
[SA] new: 1127. $less(sK0,i(main_end)) | $less(sK0,alength) | alength = sK0 | i(main_end) = sK0 | i(main_end) = sK0 (6:18) [resolution 204,722]
[SA] new: 1128. $less(sK0,alength) | $less(alength,sK0) | alength = sK0 | alength = sK0 (4:12) [factoring 204]
[SA] new: 1129. $less(sK0,sK0) | $less(sK0,alength) | alength = sK0 | sK0 = sK0 (4:12) [factoring 204]
[SA] new: 1130. $less(sK0,alength) | $less(alength,alength) | alength = alength | alength = sK0 (4:12) [factoring 204]
[SA] new: 1131. $less(sK0,alength) | $less(alength,sK0) | alength = sK0 (4:9) [duplicate literal removal 1128]
[SA] new: 1132. $less(sK0,i(main_end)) | $less(sK0,alength) | alength = sK0 | i(main_end) = sK0 (6:14) [duplicate literal removal 1127]
[SA] new: 1133. $less(sK0,X18) | $less(0,alength) | alength = 0 | sK0 = X18 | $less(X18,alength) (5:15) [duplicate literal removal 1119]
[SA] new: 1134. $less(sK0,X3) | $less(X4,alength) | alength = X4 | sK0 = X3 | $less(sK0,X4) | $less(X3,alength) (4:18) [duplicate literal removal 1111]
[SA] new: 1135. $less(X21,alength) | $less(0,X21) | alength = X21 | 0 = sK0 | $less(sK0,sK0) (4:15) [duplicate literal removal 1097]
[SA] new: 1136. $less(X20,alength) | $less(0,X20) | alength = X20 | 0 = sK0 | $less(sK0,alength) (4:15) [duplicate literal removal 1096]
[SA] new: 1137. $less(X6,alength) | $less(X7,X6) | alength = X6 | sK0 = X7 | $less(sK0,X7) | $less(sK0,alength) (4:18) [duplicate literal removal 1088]
[SA] passive: 1090. $less(X11,alength) | $less(X12,X11) | alength = X11 | sK0 = X12 | ~$less(X13,sK0) | $less(X13,X12) (4:18) [resolution 204,14]
[SA] passive: 1105. $less(sK0,X7) | $less(X7,X8) | alength = X8 | sK0 = X7 | alength = X9 | $less(X9,alength) | $less(X8,X9) (4:21) [resolution 204,99]
[SA] passive: 1106. $less(sK0,X10) | $less(X10,X11) | alength = X11 | sK0 = X10 | ~$less(X12,X11) | $less(X12,alength) (4:18) [resolution 204,14]
[SA] passive: 1107. $less(sK0,X13) | $less(X13,i(main_end)) | alength = i(main_end) | sK0 = X13 (4:14) [resolution 204,82]
[SA] passive: 1113. $less(sK0,X7) | $less(X8,alength) | alength = X8 | sK0 = X7 | X8 = X9 | $less(X9,X8) | $less(X7,X9) (4:21) [resolution 204,99]
[SA] passive: 1114. $less(sK0,X10) | $less(X11,alength) | alength = X11 | sK0 = X10 | ~$less(X12,X10) | $less(X12,X11) (4:18) [resolution 204,14]
[SA] new: 1138. $less(sK0,X13) | alength = i(main_end) | sK0 = X13 | $less(X14,0) | i(main_end) = X14 | i(l8(sK2(X14))) = X14 | $less(X13,X14) (4:26) [subsumption resolution 1115,82]
[SA] passive: 1138. $less(sK0,X13) | alength = i(main_end) | sK0 = X13 | $less(X14,0) | i(main_end) = X14 | i(l8(sK2(X14))) = X14 | $less(X13,X14) (4:26) [subsumption resolution 1115,82]
[SA] new: 1139. $less(sK0,X15) | alength = i(main_end) | sK0 = X15 | $less(X15,0) | i(l8(sK2(X15))) = X15 (4:19) [subsumption resolution 1116,82]
[SA] passive: 1139. $less(sK0,X15) | alength = i(main_end) | sK0 = X15 | $less(X15,0) | i(l8(sK2(X15))) = X15 (4:19) [subsumption resolution 1116,82]
[SA] new: 1140. $less(sK0,X16) | alength = i(main_end) | sK0 = X16 | $less(X16,0) | 'Sub'(sK2(X16),nl8) (4:17) [subsumption resolution 1117,82]
[SA] passive: 1140. $less(sK0,X16) | alength = i(main_end) | sK0 = X16 | $less(X16,0) | 'Sub'(sK2(X16),nl8) (4:17) [subsumption resolution 1117,82]
[SA] active: 358. v = a(l8(s(s(s(s(s(zero)))))),4) (5:11:1) [superposition 349,186]
[SA] new: 1141. v = a(l11(s(s(s(s(zero))))),4) (6:10) [superposition 358,57]
[SA] new: 1142. v = a(l11(s(s(s(s(zero))))),4) (6:10) [superposition 57,358]
[SA] active: 246. $less(sK0,X1) | $less(X1,X0) | $less(X0,alength) | X0 = X1 | alength = X0 (3:15:3) [resolution 179,99]
[SA] new: 1143. $less(0,X0) | $less(X0,alength) | 0 = X0 | alength = X0 (4:12) [resolution 246,46]
[SA] new: 1144. $less(sK0,X1) | $less(X1,alength) | sK0 = X1 | alength = X1 (4:12) [resolution 246,13]
[SA] new: 1145. $less(X2,X3) | $less(X3,alength) | X2 = X3 | alength = X3 | alength = i(main_end) | ~$less(X2,alength) | $less(sK0,i(main_end)) (4:23) [resolution 246,277]
[SA] new: 1146. $less(X4,X5) | $less(X5,alength) | X4 = X5 | alength = X5 | alength = X4 | $less(sK0,X4) | $less(sK0,alength) (4:21) [resolution 246,254]
[SA] new: 1147. $less(X6,X7) | $less(X7,alength) | X6 = X7 | alength = X7 | sK0 = X6 | $less(sK0,X6) | $less(sK0,alength) (4:21) [resolution 246,205]
[SA] new: 1148. $less(X8,X9) | $less(X9,alength) | X8 = X9 | alength = X9 | X8 = X10 | $less(X10,X8) | $less(sK0,X10) (4:21) [resolution 246,99]
[SA] new: 1149. $less(X11,X12) | $less(X12,alength) | X11 = X12 | alength = X12 | ~$less(X13,sK0) | $less(X13,X11) (4:18) [resolution 246,14]
[SA] new: 1150. $less(i(main_end),X14) | $less(X14,alength) | i(main_end) = X14 | alength = X14 | $less(X15,0) | i(main_end) = X15 | i(l8(sK2(X15))) = X15 | $less(sK0,X15) (4:30) [resolution 246,215]
[SA] new: 1151. $less(i(main_end),X16) | $less(X16,alength) | i(main_end) = X16 | alength = X16 | $less(sK0,0) | i(l8(sK2(sK0))) = sK0 (4:23) [resolution 246,84]
[SA] new: 1152. $less(i(main_end),X17) | $less(X17,alength) | i(main_end) = X17 | alength = X17 | $less(sK0,0) | 'Sub'(sK2(sK0),nl8) (4:21) [resolution 246,83]
[SA] new: 1153. $less(i(l8(p(nl8))),X18) | $less(X18,alength) | i(l8(p(nl8))) = X18 | alength = X18 | zero = nl8 | $less(sK0,alength) (4:24) [resolution 246,194]
[SA] new: 1154. $less(0,X19) | $less(X19,alength) | 0 = X19 | alength = X19 | alength = 0 | $less(sK0,alength) (5:18) [resolution 246,262]
[SA] new: 1155. $less(0,X20) | $less(X20,alength) | 0 = X20 | alength = X20 | 0 = sK0 | $less(sK0,alength) (4:18) [resolution 246,116]
[SA] new: 1156. $less(0,X21) | $less(X21,alength) | 0 = X21 | alength = X21 | 0 = sK0 | $less(sK0,sK0) (4:18) [resolution 246,114]
[SA] new: 1157. $less(sK0,X22) | $less(X22,alength) | sK0 = X22 | alength = X22 | alength = X23 | $less(X23,alength) | $less(sK0,X23) (4:21) [resolution 246,247]
[SA] new: 1158. $less(sK0,X24) | $less(X24,alength) | sK0 = X24 | alength = X24 | sK0 = X25 | $less(X25,alength) | $less(sK0,X25) (4:21) [resolution 246,199]
[SA] new: 1159. $less(sK0,X26) | $less(X26,alength) | sK0 = X26 | alength = X26 | $less(sK0,alength) (4:15) [resolution 246,96]
[SA] new: 1160. $less(sK0,sK0) | $less(0,alength) | 0 = sK0 | alength = 0 (4:12) [resolution 246,46]
[SA] new: 1161. $less(sK0,X0) | $less(X0,alength) | X0 = X0 | alength = X0 (4:12) [resolution 246,13]
[SA] new: 1162. $less(sK0,X1) | $less(X2,alength) | X1 = X2 | alength = X2 | alength = i(main_end) | ~$less(X2,alength) | $less(X1,i(main_end)) (4:23) [resolution 246,277]
[SA] new: 1163. $less(sK0,X3) | $less(X4,alength) | X3 = X4 | alength = X4 | alength = X4 | $less(sK0,X4) | $less(X3,alength) (4:21) [resolution 246,254]
[SA] new: 1164. $less(sK0,X5) | $less(X6,alength) | X5 = X6 | alength = X6 | sK0 = X6 | $less(sK0,X6) | $less(X5,alength) (4:21) [resolution 246,205]
[SA] new: 1165. $less(sK0,X7) | $less(X8,alength) | X7 = X8 | alength = X8 | X8 = X9 | $less(X9,X8) | $less(X7,X9) (4:21) [resolution 246,99]
[SA] new: 1166. $less(sK0,X10) | $less(X11,alength) | X10 = X11 | alength = X11 | ~$less(X12,X10) | $less(X12,X11) (4:18) [resolution 246,14]
[SA] new: 1167. $less(sK0,X13) | $less(i(main_end),alength) | i(main_end) = X13 | alength = i(main_end) | $less(X14,0) | i(main_end) = X14 | i(l8(sK2(X14))) = X14 | $less(X13,X14) (4:31) [resolution 246,215]
[SA] new: 1168. $less(sK0,X15) | $less(i(main_end),alength) | i(main_end) = X15 | alength = i(main_end) | $less(X15,0) | i(l8(sK2(X15))) = X15 (4:24) [resolution 246,84]
[SA] new: 1169. $less(sK0,X16) | $less(i(main_end),alength) | i(main_end) = X16 | alength = i(main_end) | $less(X16,0) | 'Sub'(sK2(X16),nl8) (4:22) [resolution 246,83]
[SA] new: 1170. $less(sK0,X17) | $less(i(l8(p(nl8))),alength) | i(l8(p(nl8))) = X17 | alength = i(l8(p(nl8))) | zero = nl8 | $less(X17,alength) (4:27) [resolution 246,194]
[SA] new: 1171. $less(sK0,X18) | $less(0,alength) | 0 = X18 | alength = 0 | alength = 0 | $less(X18,alength) (5:18) [resolution 246,262]
[SA] new: 1172. $less(sK0,X19) | $less(0,alength) | 0 = X19 | alength = 0 | 0 = sK0 | $less(X19,alength) (4:18) [resolution 246,116]
[SA] new: 1173. $less(sK0,X20) | $less(0,alength) | 0 = X20 | alength = 0 | 0 = sK0 | $less(X20,sK0) (4:18) [resolution 246,114]
[SA] new: 1174. $less(sK0,X21) | $less(sK0,alength) | sK0 = X21 | alength = sK0 | alength = X22 | $less(X22,alength) | $less(X21,X22) (4:21) [resolution 246,247]
[SA] new: 1175. $less(sK0,X23) | $less(sK0,alength) | sK0 = X23 | alength = sK0 | sK0 = X24 | $less(X24,alength) | $less(X23,X24) (4:21) [resolution 246,199]
[SA] new: 1176. $less(sK0,X25) | $less(sK0,alength) | sK0 = X25 | alength = sK0 | $less(X25,alength) (4:15) [resolution 246,96]
[SA] new: 1177. $less(sK0,i(main_end)) | $less(alength,alength) | alength = i(main_end) | alength = alength (4:14) [resolution 246,82]
[SA] new: 1178. $less(sK0,i(main_end)) | $less(sK0,alength) | i(main_end) = sK0 | alength = sK0 | alength = i(main_end) (6:18) [resolution 246,812]
[SA] new: 1179. $less(sK0,i(main_end)) | $less(sK0,alength) | i(main_end) = sK0 | alength = sK0 | i(main_end) = sK0 (6:18) [resolution 246,722]
[SA] new: 1180. $less(sK0,X0) | $less(X0,alength) | alength = X0 | alength = alength (4:12) [resolution 246,13]
[SA] new: 1181. $less(sK0,X1) | $less(X1,X2) | X1 = X2 | alength = X2 | alength = i(main_end) | ~$less(alength,alength) | $less(X2,i(main_end)) (4:23) [resolution 246,277]
[SA] new: 1182. $less(sK0,X3) | $less(X3,X4) | X3 = X4 | alength = X4 | alength = alength | $less(sK0,alength) | $less(X4,alength) (4:21) [resolution 246,254]
[SA] new: 1183. $less(sK0,X5) | $less(X5,X6) | X5 = X6 | alength = X6 | alength = sK0 | $less(sK0,alength) | $less(X6,alength) (4:21) [resolution 246,205]
[SA] new: 1184. $less(sK0,X7) | $less(X7,X8) | X7 = X8 | alength = X8 | alength = X9 | $less(X9,alength) | $less(X8,X9) (4:21) [resolution 246,99]
[SA] new: 1185. $less(sK0,X10) | $less(X10,X11) | X10 = X11 | alength = X11 | ~$less(X12,X11) | $less(X12,alength) (4:18) [resolution 246,14]
[SA] new: 1186. $less(sK0,X13) | $less(X13,i(main_end)) | i(main_end) = X13 | alength = i(main_end) (4:15) [resolution 246,82]
[SA] new: 1187. $less(sK0,sK0) | $less(sK0,alength) | sK0 = sK0 | alength = sK0 (4:12) [factoring 246]
[SA] new: 1188. $less(sK0,alength) | $less(alength,sK0) | alength = sK0 | alength = sK0 (4:12) [factoring 246]
[SA] new: 1189. $less(sK0,alength) | $less(alength,alength) | alength = alength | alength = alength (4:12) [factoring 246]
[SA] new: 1190. $less(sK0,alength) | $less(alength,alength) | alength = alength (4:9) [duplicate literal removal 1189]
[SA] new: 1191. $less(sK0,alength) | $less(alength,sK0) | alength = sK0 (4:9) [duplicate literal removal 1188]
[SA] new: 1192. $less(sK0,i(main_end)) | $less(sK0,alength) | i(main_end) = sK0 | alength = sK0 (6:14) [duplicate literal removal 1179]
[SA] new: 1193. $less(sK0,X18) | $less(0,alength) | 0 = X18 | alength = 0 | $less(X18,alength) (5:15) [duplicate literal removal 1171]
[SA] new: 1194. $less(sK0,X3) | $less(X4,alength) | X3 = X4 | alength = X4 | $less(sK0,X4) | $less(X3,alength) (4:18) [duplicate literal removal 1163]
[SA] passive: 1143. $less(0,X0) | $less(X0,alength) | 0 = X0 | alength = X0 (4:12) [resolution 246,46]
[SA] passive: 1148. $less(X8,X9) | $less(X9,alength) | X8 = X9 | alength = X9 | X8 = X10 | $less(X10,X8) | $less(sK0,X10) (4:21) [resolution 246,99]
[SA] passive: 1149. $less(X11,X12) | $less(X12,alength) | X11 = X12 | alength = X12 | ~$less(X13,sK0) | $less(X13,X11) (4:18) [resolution 246,14]
[SA] passive: 1150. $less(i(main_end),X14) | $less(X14,alength) | i(main_end) = X14 | alength = X14 | $less(X15,0) | i(main_end) = X15 | i(l8(sK2(X15))) = X15 | $less(sK0,X15) (4:30) [resolution 246,215]
[SA] new: 1195. $less(i(main_end),X16) | $less(X16,alength) | i(main_end) = X16 | alength = X16 | i(l8(sK2(sK0))) = sK0 (4:20) [subsumption resolution 1151,46]
[SA] passive: 1195. $less(i(main_end),X16) | $less(X16,alength) | i(main_end) = X16 | alength = X16 | i(l8(sK2(sK0))) = sK0 (4:20) [subsumption resolution 1151,46]
[SA] new: 1196. $less(i(main_end),X17) | $less(X17,alength) | i(main_end) = X17 | alength = X17 | 'Sub'(sK2(sK0),nl8) (4:18) [subsumption resolution 1152,46]
[SA] passive: 1196. $less(i(main_end),X17) | $less(X17,alength) | i(main_end) = X17 | alength = X17 | 'Sub'(sK2(sK0),nl8) (4:18) [subsumption resolution 1152,46]
[SA] passive: 1165. $less(sK0,X7) | $less(X8,alength) | X7 = X8 | alength = X8 | X8 = X9 | $less(X9,X8) | $less(X7,X9) (4:21) [resolution 246,99]
[SA] passive: 1166. $less(sK0,X10) | $less(X11,alength) | X10 = X11 | alength = X11 | ~$less(X12,X10) | $less(X12,X11) (4:18) [resolution 246,14]
[SA] new: 1197. $less(sK0,X13) | i(main_end) = X13 | alength = i(main_end) | $less(X14,0) | i(main_end) = X14 | i(l8(sK2(X14))) = X14 | $less(X13,X14) (4:27) [subsumption resolution 1167,82]
[SA] passive: 1197. $less(sK0,X13) | i(main_end) = X13 | alength = i(main_end) | $less(X14,0) | i(main_end) = X14 | i(l8(sK2(X14))) = X14 | $less(X13,X14) (4:27) [subsumption resolution 1167,82]
[SA] new: 1198. $less(sK0,X16) | i(main_end) = X16 | alength = i(main_end) | $less(X16,0) | 'Sub'(sK2(X16),nl8) (4:18) [subsumption resolution 1169,82]
[SA] passive: 1198. $less(sK0,X16) | i(main_end) = X16 | alength = i(main_end) | $less(X16,0) | 'Sub'(sK2(X16),nl8) (4:18) [subsumption resolution 1169,82]
[SA] passive: 1184. $less(sK0,X7) | $less(X7,X8) | X7 = X8 | alength = X8 | alength = X9 | $less(X9,alength) | $less(X8,X9) (4:21) [resolution 246,99]
[SA] passive: 1185. $less(sK0,X10) | $less(X10,X11) | X10 = X11 | alength = X11 | ~$less(X12,X11) | $less(X12,alength) (4:18) [resolution 246,14]
[SA] active: 629. v = a(l11(s(s(s(s(s(zero)))))),5) (6:11:1) [superposition 70,312]
[SA] new: 1199. v = a(l8(s(s(s(s(s(zero)))))),5) | i(l8(s(s(s(s(s(zero))))))) = 5 (7:21) [superposition 629,72]
[SA] new: 1200. v = a(l8(s(s(s(s(s(zero)))))),5) | i(l8(s(s(s(s(s(zero))))))) = 5 (7:21) [superposition 72,629]
[SA] active: 253. $less(sK0,X0) | $less(X1,alength) | $less(X0,X1) | alength = X1 | alength = X0 (3:15:3) [resolution 179,99]
[SA] new: 1201. $less(X0,alength) | $less(0,X0) | alength = X0 | alength = 0 (4:12) [resolution 253,46]
[SA] new: 1202. $less(X1,alength) | $less(sK0,X1) | alength = X1 | alength = sK0 (4:12) [resolution 253,13]
[SA] new: 1203. $less(X2,alength) | $less(X3,X2) | alength = X2 | alength = X3 | alength = i(main_end) | ~$less(X3,alength) | $less(sK0,i(main_end)) (4:23) [resolution 253,277]
[SA] new: 1204. $less(X4,alength) | $less(X5,X4) | alength = X4 | alength = X5 | alength = X5 | $less(sK0,X5) | $less(sK0,alength) (4:21) [resolution 253,254]
[SA] new: 1205. $less(X6,alength) | $less(X7,X6) | alength = X6 | alength = X7 | sK0 = X7 | $less(sK0,X7) | $less(sK0,alength) (4:21) [resolution 253,205]
[SA] new: 1206. $less(X8,alength) | $less(X9,X8) | alength = X8 | alength = X9 | X9 = X10 | $less(X10,X9) | $less(sK0,X10) (4:21) [resolution 253,99]
[SA] new: 1207. $less(X11,alength) | $less(X12,X11) | alength = X11 | alength = X12 | ~$less(X13,sK0) | $less(X13,X12) (4:18) [resolution 253,14]
[SA] new: 1208. $less(X14,alength) | $less(i(main_end),X14) | alength = X14 | alength = i(main_end) | $less(X15,0) | i(main_end) = X15 | i(l8(sK2(X15))) = X15 | $less(sK0,X15) (4:30) [resolution 253,215]
[SA] new: 1209. $less(X16,alength) | $less(i(main_end),X16) | alength = X16 | alength = i(main_end) | $less(sK0,0) | i(l8(sK2(sK0))) = sK0 (4:23) [resolution 253,84]
[SA] new: 1210. $less(X17,alength) | $less(i(main_end),X17) | alength = X17 | alength = i(main_end) | $less(sK0,0) | 'Sub'(sK2(sK0),nl8) (4:21) [resolution 253,83]
[SA] new: 1211. $less(X18,alength) | $less(i(l8(p(nl8))),X18) | alength = X18 | alength = i(l8(p(nl8))) | zero = nl8 | $less(sK0,alength) (4:24) [resolution 253,194]
[SA] new: 1212. $less(X19,alength) | $less(0,X19) | alength = X19 | alength = 0 | alength = 0 | $less(sK0,alength) (5:18) [resolution 253,262]
[SA] new: 1213. $less(X20,alength) | $less(0,X20) | alength = X20 | alength = 0 | 0 = sK0 | $less(sK0,alength) (4:18) [resolution 253,116]
[SA] new: 1214. $less(X21,alength) | $less(0,X21) | alength = X21 | alength = 0 | 0 = sK0 | $less(sK0,sK0) (4:18) [resolution 253,114]
[SA] new: 1215. $less(X22,alength) | $less(sK0,X22) | alength = X22 | alength = sK0 | alength = X23 | $less(X23,alength) | $less(sK0,X23) (4:21) [resolution 253,247]
[SA] new: 1216. $less(X24,alength) | $less(sK0,X24) | alength = X24 | alength = sK0 | sK0 = X25 | $less(X25,alength) | $less(sK0,X25) (4:21) [resolution 253,199]
[SA] new: 1217. $less(X26,alength) | $less(sK0,X26) | alength = X26 | alength = sK0 | $less(sK0,alength) (4:15) [resolution 253,96]
[SA] new: 1218. $less(sK0,X0) | $less(X0,alength) | alength = alength | alength = X0 (4:12) [resolution 253,13]
[SA] new: 1219. $less(sK0,X1) | $less(X1,X2) | alength = X2 | alength = X1 | alength = i(main_end) | ~$less(alength,alength) | $less(X2,i(main_end)) (4:23) [resolution 253,277]
[SA] new: 1220. $less(sK0,X3) | $less(X3,X4) | alength = X4 | alength = X3 | alength = alength | $less(sK0,alength) | $less(X4,alength) (4:21) [resolution 253,254]
[SA] new: 1221. $less(sK0,X5) | $less(X5,X6) | alength = X6 | alength = X5 | alength = sK0 | $less(sK0,alength) | $less(X6,alength) (4:21) [resolution 253,205]
[SA] new: 1222. $less(sK0,X7) | $less(X7,X8) | alength = X8 | alength = X7 | alength = X9 | $less(X9,alength) | $less(X8,X9) (4:21) [resolution 253,99]
[SA] new: 1223. $less(sK0,X10) | $less(X10,X11) | alength = X11 | alength = X10 | ~$less(X12,X11) | $less(X12,alength) (4:18) [resolution 253,14]
[SA] new: 1224. $less(sK0,X13) | $less(X13,i(main_end)) | alength = i(main_end) | alength = X13 (4:14) [resolution 253,82]
[SA] new: 1225. $less(sK0,sK0) | $less(0,alength) | alength = 0 | alength = sK0 (4:12) [resolution 253,46]
[SA] new: 1226. $less(sK0,X0) | $less(X0,alength) | alength = X0 | alength = X0 (4:12) [resolution 253,13]
[SA] new: 1227. $less(sK0,X1) | $less(X2,alength) | alength = X2 | alength = X1 | alength = i(main_end) | ~$less(X2,alength) | $less(X1,i(main_end)) (4:23) [resolution 253,277]
[SA] new: 1228. $less(sK0,X3) | $less(X4,alength) | alength = X4 | alength = X3 | alength = X4 | $less(sK0,X4) | $less(X3,alength) (4:21) [resolution 253,254]
[SA] new: 1229. $less(sK0,X5) | $less(X6,alength) | alength = X6 | alength = X5 | sK0 = X6 | $less(sK0,X6) | $less(X5,alength) (4:21) [resolution 253,205]
[SA] new: 1230. $less(sK0,X7) | $less(X8,alength) | alength = X8 | alength = X7 | X8 = X9 | $less(X9,X8) | $less(X7,X9) (4:21) [resolution 253,99]
[SA] new: 1231. $less(sK0,X10) | $less(X11,alength) | alength = X11 | alength = X10 | ~$less(X12,X10) | $less(X12,X11) (4:18) [resolution 253,14]
[SA] new: 1232. $less(sK0,X13) | $less(i(main_end),alength) | alength = i(main_end) | alength = X13 | $less(X14,0) | i(main_end) = X14 | i(l8(sK2(X14))) = X14 | $less(X13,X14) (4:30) [resolution 253,215]
[SA] new: 1233. $less(sK0,X15) | $less(i(main_end),alength) | alength = i(main_end) | alength = X15 | $less(X15,0) | i(l8(sK2(X15))) = X15 (4:23) [resolution 253,84]
[SA] new: 1234. $less(sK0,X16) | $less(i(main_end),alength) | alength = i(main_end) | alength = X16 | $less(X16,0) | 'Sub'(sK2(X16),nl8) (4:21) [resolution 253,83]
[SA] new: 1235. $less(sK0,X17) | $less(i(l8(p(nl8))),alength) | alength = i(l8(p(nl8))) | alength = X17 | zero = nl8 | $less(X17,alength) (4:24) [resolution 253,194]
[SA] new: 1236. $less(sK0,X18) | $less(0,alength) | alength = 0 | alength = X18 | alength = 0 | $less(X18,alength) (5:18) [resolution 253,262]
[SA] new: 1237. $less(sK0,X19) | $less(0,alength) | alength = 0 | alength = X19 | 0 = sK0 | $less(X19,alength) (4:18) [resolution 253,116]
[SA] new: 1238. $less(sK0,X20) | $less(0,alength) | alength = 0 | alength = X20 | 0 = sK0 | $less(X20,sK0) (4:18) [resolution 253,114]
[SA] new: 1239. $less(sK0,X21) | $less(sK0,alength) | alength = sK0 | alength = X21 | alength = X22 | $less(X22,alength) | $less(X21,X22) (4:21) [resolution 253,247]
[SA] new: 1240. $less(sK0,X23) | $less(sK0,alength) | alength = sK0 | alength = X23 | sK0 = X24 | $less(X24,alength) | $less(X23,X24) (4:21) [resolution 253,199]
[SA] new: 1241. $less(sK0,X25) | $less(sK0,alength) | alength = sK0 | alength = X25 | $less(X25,alength) (4:15) [resolution 253,96]
[SA] new: 1242. $less(sK0,i(main_end)) | $less(alength,alength) | alength = alength | alength = i(main_end) (4:14) [resolution 253,82]
[SA] new: 1243. $less(sK0,i(main_end)) | $less(sK0,alength) | alength = sK0 | alength = i(main_end) | alength = i(main_end) (6:18) [resolution 253,812]
[SA] new: 1244. $less(sK0,i(main_end)) | $less(sK0,alength) | alength = sK0 | alength = i(main_end) | i(main_end) = sK0 (6:18) [resolution 253,722]
[SA] new: 1245. $less(sK0,alength) | $less(alength,sK0) | alength = sK0 | alength = alength (4:12) [factoring 253]
[SA] new: 1246. $less(sK0,sK0) | $less(sK0,alength) | alength = sK0 | alength = sK0 (4:12) [factoring 253]
[SA] new: 1247. $less(sK0,alength) | $less(alength,alength) | alength = alength | alength = alength (4:12) [factoring 253]
[SA] new: 1248. $less(sK0,alength) | $less(alength,alength) | alength = alength (4:9) [duplicate literal removal 1247]
[SA] new: 1249. $less(sK0,sK0) | $less(sK0,alength) | alength = sK0 (4:9) [duplicate literal removal 1246]
[SA] new: 1250. $less(sK0,i(main_end)) | $less(sK0,alength) | alength = sK0 | alength = i(main_end) (6:14) [duplicate literal removal 1243]
[SA] new: 1251. $less(sK0,X18) | $less(0,alength) | alength = 0 | alength = X18 | $less(X18,alength) (5:15) [duplicate literal removal 1236]
[SA] new: 1252. $less(sK0,X3) | $less(X4,alength) | alength = X4 | alength = X3 | $less(sK0,X4) | $less(X3,alength) (4:18) [duplicate literal removal 1228]
[SA] new: 1253. $less(sK0,X0) | $less(X0,alength) | alength = X0 (4:9) [duplicate literal removal 1226]
[SA] new: 1254. $less(X19,alength) | $less(0,X19) | alength = X19 | alength = 0 | $less(sK0,alength) (5:15) [duplicate literal removal 1212]
[SA] new: 1255. $less(X4,alength) | $less(X5,X4) | alength = X4 | alength = X5 | $less(sK0,X5) | $less(sK0,alength) (4:18) [duplicate literal removal 1204]
[SA] passive: 1207. $less(X11,alength) | $less(X12,X11) | alength = X11 | alength = X12 | ~$less(X13,sK0) | $less(X13,X12) (4:18) [resolution 253,14]
[SA] passive: 1222. $less(sK0,X7) | $less(X7,X8) | alength = X8 | alength = X7 | alength = X9 | $less(X9,alength) | $less(X8,X9) (4:21) [resolution 253,99]
[SA] passive: 1223. $less(sK0,X10) | $less(X10,X11) | alength = X11 | alength = X10 | ~$less(X12,X11) | $less(X12,alength) (4:18) [resolution 253,14]
[SA] passive: 1224. $less(sK0,X13) | $less(X13,i(main_end)) | alength = i(main_end) | alength = X13 (4:14) [resolution 253,82]
[SA] passive: 1230. $less(sK0,X7) | $less(X8,alength) | alength = X8 | alength = X7 | X8 = X9 | $less(X9,X8) | $less(X7,X9) (4:21) [resolution 253,99]
[SA] passive: 1231. $less(sK0,X10) | $less(X11,alength) | alength = X11 | alength = X10 | ~$less(X12,X10) | $less(X12,X11) (4:18) [resolution 253,14]
[SA] new: 1256. $less(sK0,X13) | alength = i(main_end) | alength = X13 | $less(X14,0) | i(main_end) = X14 | i(l8(sK2(X14))) = X14 | $less(X13,X14) (4:26) [subsumption resolution 1232,82]
[SA] passive: 1256. $less(sK0,X13) | alength = i(main_end) | alength = X13 | $less(X14,0) | i(main_end) = X14 | i(l8(sK2(X14))) = X14 | $less(X13,X14) (4:26) [subsumption resolution 1232,82]
[SA] new: 1257. $less(sK0,X15) | alength = i(main_end) | alength = X15 | $less(X15,0) | i(l8(sK2(X15))) = X15 (4:19) [subsumption resolution 1233,82]
[SA] passive: 1257. $less(sK0,X15) | alength = i(main_end) | alength = X15 | $less(X15,0) | i(l8(sK2(X15))) = X15 (4:19) [subsumption resolution 1233,82]
[SA] new: 1258. $less(sK0,X16) | alength = i(main_end) | alength = X16 | $less(X16,0) | 'Sub'(sK2(X16),nl8) (4:17) [subsumption resolution 1234,82]
[SA] passive: 1258. $less(sK0,X16) | alength = i(main_end) | alength = X16 | $less(X16,0) | 'Sub'(sK2(X16),nl8) (4:17) [subsumption resolution 1234,82]
[SA] active: 640. i(l8(s(s(s(s(s(s(zero)))))))) = 6 (6:11:1) [evaluation 630]
[SA] new: 1259. i(l8(X0)) != 6 | s(s(s(s(s(s(zero)))))) = X0 (7:14) [superposition 48,640]
[SA] new: 1260. i(l8(X1)) != 6 | s(s(s(s(s(s(zero)))))) = X1 (7:14) [superposition 48,640]
[SA] new: 1261. v = a(l11(s(s(s(s(s(s(zero))))))),6) (7:12) [superposition 70,640]
[SA] new: 1262. i(l8(s(s(s(s(s(s(s(zero))))))))) = $sum(6,1) (7:14) [superposition 74,640]
[SA] new: 1263. 0 != 6 | zero = s(s(s(s(s(s(zero)))))) (7:12) [superposition 108,640]
[SA] new: 1264. i(main_end) != 6 | s(s(s(s(s(s(zero)))))) = nl8 (7:13) [superposition 109,640]
[SA] new: 1265. 1 != 6 | s(zero) = s(s(s(s(s(s(zero)))))) (7:13) [superposition 121,640]
[SA] new: 1266. 2 != 6 | s(s(zero)) = s(s(s(s(s(s(zero)))))) (7:14) [superposition 130,640]
[SA] new: 1267. $sum(i(main_end),1) != 6 | s(nl8) = s(s(s(s(s(s(zero)))))) (7:16) [superposition 158,640]
[SA] new: 1268. 3 != 6 | s(s(s(zero))) = s(s(s(s(s(s(zero)))))) (7:15) [superposition 183,640]
[SA] new: 1269. v = a(l8(s(s(s(s(s(s(s(zero)))))))),6) (7:13) [superposition 349,640]
[SA] new: 1270. i(main_end) != $sum(6,-1) | s(nl8) = s(s(s(s(s(s(zero)))))) (7:16) [evaluation 1267]
[SA] new: 1271. i(main_end) != 5 | s(nl8) = s(s(s(s(s(s(zero)))))) (7:14) [evaluation 1270]
[SA] new: 1272. i(main_end) != 5 | s(s(s(s(s(zero))))) = nl8 (7:12) [term algebras injectivity 1271]
[SA] new: 1273. i(l8(s(s(s(s(s(s(s(zero))))))))) = 7 (7:12) [evaluation 1262]
[SA] passive: 1259. i(l8(X0)) != 6 | s(s(s(s(s(s(zero)))))) = X0 (7:14) [superposition 48,640]
[SA] passive: 1261. v = a(l11(s(s(s(s(s(s(zero))))))),6) (7:12) [superposition 70,640]
[SA] passive: 1273. i(l8(s(s(s(s(s(s(s(zero))))))))) = 7 (7:12) [evaluation 1262]
[SA] passive: 1264. i(main_end) != 6 | s(s(s(s(s(s(zero)))))) = nl8 (7:13) [superposition 109,640]
[SA] passive: 1269. v = a(l8(s(s(s(s(s(s(s(zero)))))))),6) (7:13) [superposition 349,640]
[SA] active: 324. ~$less(X3,0) | sK0 = X2 | 0 = sK0 | $less(X2,sK0) | $less(X3,X2) (3:15:1) [resolution 177,14]
[SA] new: 1274. sK0 = X0 | 0 = sK0 | $less(X0,sK0) | $less(X1,X0) | $less(i(main_end),X2) | $less(X2,X1) | X1 = X2 | i(main_end) = X1 | i(l8(sK2(X1))) = X1 (4:32) [resolution 324,214]
[SA] new: 1275. sK0 = X3 | 0 = sK0 | $less(X3,sK0) | $less(X4,X3) | $less(0,X4) | 0 = X4 (4:18) [resolution 324,15]
[SA] new: 1276. sK0 = X5 | 0 = sK0 | $less(X5,sK0) | $less(X6,X5) | $less(0,X6) | 0 = X6 (4:18) [resolution 324,15]
[SA] new: 1277. sK0 = X7 | 0 = sK0 | $less(X7,sK0) | $less(X8,X7) | $less(X9,X8) | $less(0,X9) | 0 = X8 | 0 = X9 (4:24) [resolution 324,174]
[SA] new: 1278. sK0 = X10 | 0 = sK0 | $less(X10,sK0) | $less(X11,X10) | $less(0,X12) | $less(X12,X11) | X11 = X12 | 0 = X11 (4:24) [resolution 324,174]
[SA] new: 1279. sK0 = X13 | 0 = sK0 | $less(X13,sK0) | $less(X14,X13) | $less(X15,X14) | $less(0,X15) | 0 = X15 | X14 = X15 (4:24) [resolution 324,174]
[SA] new: 1280. sK0 = X16 | 0 = sK0 | $less(X16,sK0) | $less(X17,X16) | $less(sK0,X17) | $less(0,alength) | 0 = X17 | 0 = sK0 (4:24) [resolution 324,198]
[SA] new: 1281. sK0 = X18 | 0 = sK0 | $less(X18,sK0) | $less(X19,X18) | $less(sK0,X19) | $less(0,alength) | alength = 0 | sK0 = X19 (4:24) [resolution 324,204]
[SA] new: 1282. sK0 = X20 | 0 = sK0 | $less(X20,sK0) | $less(X21,X20) | $less(i(main_end),X21) | $less(0,0) | 0 = X21 | i(main_end) = 0 | i(l8(sK2(0))) = 0 (4:32) [resolution 324,214]
[SA] new: 1283. sK0 = X22 | 0 = sK0 | $less(X22,sK0) | $less(X23,X22) | $less(sK0,X23) | $less(0,alength) | 0 = X23 | alength = 0 (4:24) [resolution 324,246]
[SA] new: 1284. sK0 = X24 | 0 = sK0 | $less(X24,sK0) | $less(X25,X24) | $less(sK0,X25) | $less(0,alength) | alength = 0 | alength = X25 (4:24) [resolution 324,253]
[SA] new: 1285. sK0 = X26 | 0 = sK0 | $less(X26,sK0) | $less(alength,X26) | alength = i(main_end) | alength = i(l8(sK2(alength))) (4:22) [resolution 324,142]
[SA] new: 1286. sK0 = X27 | 0 = sK0 | $less(X27,sK0) | $less(i(main_end),X27) | i(l8(sK1(X28,sK2(0)))) = X28 | $less(0,0) | a(main_end,X28) = a(l11(sK2(0)),X28) | i(main_end) = 0 (4:37) [resolution 324,225]
[SA] new: 1287. sK0 = X29 | 0 = sK0 | $less(X29,sK0) | $less(i(main_end),X29) | $less(0,X30) | $less(X30,0) | 0 = X30 | i(main_end) = X30 | i(l8(sK2(X30))) = X30 (4:32) [resolution 324,214]
[SA] new: 1288. sK0 = X31 | 0 = sK0 | $less(X31,sK0) | $less(i(main_end),X31) | i(l8(sK2(0))) = 0 | $less(0,0) | i(main_end) = 0 (4:26) [resolution 324,127]
[SA] new: 1289. sK0 = X32 | 0 = sK0 | $less(X32,sK0) | $less(i(l8(p(nl8))),X32) | alength = 0 | $less(0,alength) | zero = nl8 (4:24) [resolution 324,193]
[SA] new: 1290. sK0 = X33 | 0 = sK0 | $less(X33,sK0) | $less(0,X33) | $less(0,sK0) | 0 = 0 | 0 = sK0 (4:21) [resolution 324,380]
[SA] new: 1291. sK0 = X34 | 0 = sK0 | $less(X34,sK0) | $less(0,X34) | $less(0,alength) | 0 = sK0 | 0 = sK0 (4:21) [resolution 324,330]
[SA] new: 1292. sK0 = X35 | 0 = sK0 | $less(X35,sK0) | $less(0,X35) | $less(0,sK0) | 0 = sK0 | 0 = 0 (4:21) [resolution 324,229]
[SA] new: 1293. sK0 = X36 | 0 = sK0 | $less(X36,sK0) | $less(0,X36) | $less(0,alength) | alength = 0 | 0 = sK0 (4:21) [resolution 324,178]
[SA] new: 1294. sK0 = X37 | 0 = sK0 | $less(X37,sK0) | $less(0,X37) | $less(0,sK0) | 0 = sK0 | 0 = sK0 (4:21) [resolution 324,177]
[SA] new: 1295. sK0 = X38 | 0 = sK0 | $less(X38,sK0) | $less(sK0,X38) | $less(X39,alength) | $less(0,X39) | alength = X39 | alength = 0 (4:24) [resolution 324,253]
[SA] new: 1296. sK0 = X40 | 0 = sK0 | $less(X40,sK0) | $less(sK0,X40) | $less(0,X41) | $less(X41,alength) | 0 = X41 | alength = X41 (4:24) [resolution 324,246]
[SA] new: 1297. sK0 = X42 | 0 = sK0 | $less(X42,sK0) | $less(sK0,X42) | $less(X43,alength) | $less(0,X43) | alength = X43 | 0 = sK0 (4:24) [resolution 324,204]
[SA] new: 1298. sK0 = X44 | 0 = sK0 | $less(X44,sK0) | $less(sK0,X44) | $less(0,X45) | $less(X45,alength) | 0 = X45 | sK0 = X45 (4:24) [resolution 324,198]
[SA] new: 1299. sK0 = X46 | 0 = sK0 | $less(X46,sK0) | $less(sK0,X46) | $less(0,alength) | alength = 0 (4:18) [resolution 324,179]
[SA] new: 1300. sK0 = X47 | 0 = sK0 | $less(X47,sK0) | $less(sK0,X47) | $less(0,alength) | 0 = sK0 (4:18) [resolution 324,106]
[SA] new: 1301. sK0 = X47 | 0 = sK0 | $less(X47,sK0) | $less(sK0,X47) | $less(0,alength) (4:15) [duplicate literal removal 1300]
[SA] new: 1302. sK0 = X42 | 0 = sK0 | $less(X42,sK0) | $less(sK0,X42) | $less(X43,alength) | $less(0,X43) | alength = X43 (4:21) [duplicate literal removal 1297]
[SA] new: 1303. sK0 = X37 | 0 = sK0 | $less(X37,sK0) | $less(0,X37) | $less(0,sK0) (4:15) [duplicate literal removal 1294]
[SA] new: 1304. sK0 = X36 | 0 = sK0 | $less(X36,sK0) | $less(0,X36) | $less(0,alength) | alength = 0 (4:18) [duplicate literal removal 1293]
[SA] new: 1305. sK0 = X35 | 0 = sK0 | $less(X35,sK0) | $less(0,X35) | $less(0,sK0) | 0 = 0 (4:18) [duplicate literal removal 1292]
[SA] new: 1306. sK0 = X34 | 0 = sK0 | $less(X34,sK0) | $less(0,X34) | $less(0,alength) (4:15) [duplicate literal removal 1291]
[SA] new: 1307. sK0 = X33 | 0 = sK0 | $less(X33,sK0) | $less(0,X33) | $less(0,sK0) | 0 = 0 (4:18) [duplicate literal removal 1290]
[SA] new: 1308. sK0 = X31 | 0 = sK0 | $less(X31,sK0) | $less(i(main_end),X31) | i(l8(sK2(0))) = 0 | i(main_end) = 0 (4:23) [evaluation 1288]
[SA] new: 1309. sK0 = X27 | 0 = sK0 | $less(X27,sK0) | $less(i(main_end),X27) | i(l8(sK1(X28,sK2(0)))) = X28 | a(main_end,X28) = a(l11(sK2(0)),X28) | i(main_end) = 0 (4:34) [evaluation 1286]
[SA] new: 1310. sK0 = X20 | 0 = sK0 | $less(X20,sK0) | $less(X21,X20) | $less(i(main_end),X21) | 0 = X21 | i(main_end) = 0 | i(l8(sK2(0))) = 0 (4:29) [evaluation 1282]
[SA] new: 1311. sK0 = X16 | 0 = sK0 | $less(X16,sK0) | $less(X17,X16) | $less(sK0,X17) | $less(0,alength) | 0 = X17 (4:21) [duplicate literal removal 1280]
[SA] passive: 1274. sK0 = X0 | 0 = sK0 | $less(X0,sK0) | $less(X1,X0) | $less(i(main_end),X2) | $less(X2,X1) | X1 = X2 | i(main_end) = X1 | i(l8(sK2(X1))) = X1 (4:32) [resolution 324,214]
[SA] passive: 1277. sK0 = X7 | 0 = sK0 | $less(X7,sK0) | $less(X8,X7) | $less(X9,X8) | $less(0,X9) | 0 = X8 | 0 = X9 (4:24) [resolution 324,174]
[SA] passive: 1278. sK0 = X10 | 0 = sK0 | $less(X10,sK0) | $less(X11,X10) | $less(0,X12) | $less(X12,X11) | X11 = X12 | 0 = X11 (4:24) [resolution 324,174]
[SA] passive: 1279. sK0 = X13 | 0 = sK0 | $less(X13,sK0) | $less(X14,X13) | $less(X15,X14) | $less(0,X15) | 0 = X15 | X14 = X15 (4:24) [resolution 324,174]
[SA] active: 761. $less(0,i(main_end)) | i(main_end) = sK0 | 0 = sK0 (6:11:1) [duplicate literal removal 749]
[SA] new: 1312. i(main_end) = sK0 | 0 = sK0 | $less(X0,0) | i(main_end) = X0 | i(l8(sK2(X0))) = X0 | $less(0,X0) (7:23) [resolution 761,215]
[SA] new: 1313. i(main_end) = sK0 | 0 = sK0 | $less(0,0) | i(l8(sK2(0))) = 0 (7:16) [resolution 761,84]
[SA] new: 1314. i(main_end) = sK0 | 0 = sK0 | $less(0,0) | 'Sub'(sK2(0),nl8) (7:14) [resolution 761,83]
[SA] new: 1315. i(main_end) = sK0 | 0 = sK0 | alength = i(main_end) | ~$less(i(main_end),alength) | $less(0,i(main_end)) (7:19) [resolution 761,277]
[SA] new: 1316. i(main_end) = sK0 | 0 = sK0 | alength = i(main_end) | $less(sK0,i(main_end)) | $less(0,alength) (7:18) [resolution 761,254]
[SA] new: 1317. i(main_end) = sK0 | 0 = sK0 | i(main_end) = sK0 | $less(sK0,i(main_end)) | $less(0,alength) (7:18) [resolution 761,205]
[SA] new: 1318. i(main_end) = sK0 | 0 = sK0 | i(main_end) = X1 | $less(X1,i(main_end)) | $less(0,X1) (7:18) [resolution 761,99]
[SA] new: 1319. i(main_end) = sK0 | 0 = sK0 | ~$less(X2,0) | $less(X2,i(main_end)) (7:14) [resolution 761,14]
[SA] new: 1320. i(main_end) = sK0 | 0 = sK0 | $less(sK0,i(main_end)) | $less(0,alength) (7:14) [duplicate literal removal 1317]
[SA] new: 1321. i(main_end) = sK0 | 0 = sK0 | 'Sub'(sK2(0),nl8) (7:11) [evaluation 1314]
[SA] new: 1322. i(main_end) = sK0 | 0 = sK0 | i(l8(sK2(0))) = 0 (7:13) [evaluation 1313]
[SA] passive: 1318. i(main_end) = sK0 | 0 = sK0 | i(main_end) = X1 | $less(X1,i(main_end)) | $less(0,X1) (7:18) [resolution 761,99]
[SA] passive: 1319. i(main_end) = sK0 | 0 = sK0 | ~$less(X2,0) | $less(X2,i(main_end)) (7:14) [resolution 761,14]
[SA] active: 333. ~$less(X4,X3) | sK0 = X3 | 0 = sK0 | $less(0,X3) | $less(X4,sK0) (3:15:1) [resolution 177,14]
[SA] new: 1323. sK0 = X0 | 0 = sK0 | $less(0,X0) | $less(X1,sK0) | $less(X0,X1) | X0 = X1 (4:18) [resolution 333,15]
[SA] new: 1324. sK0 = X2 | 0 = sK0 | $less(0,X2) | $less(X3,sK0) | $less(X2,X3) | X2 = X3 (4:18) [resolution 333,15]
[SA] new: 1325. sK0 = X4 | 0 = sK0 | $less(0,X4) | $less(X5,sK0) | $less(X6,X5) | $less(X4,X6) | X4 = X5 | X4 = X6 (4:24) [resolution 333,174]
[SA] new: 1326. sK0 = X7 | 0 = sK0 | $less(0,X7) | $less(X8,sK0) | $less(X7,X9) | $less(X9,X8) | X8 = X9 | X7 = X8 (4:24) [resolution 333,174]
[SA] new: 1327. sK0 = X10 | 0 = sK0 | $less(0,X10) | $less(X11,sK0) | $less(X12,X11) | $less(X10,X12) | X10 = X12 | X11 = X12 (4:24) [resolution 333,174]
[SA] new: 1328. sK0 = X13 | 0 = sK0 | $less(0,X13) | $less(X14,sK0) | $less(sK0,X14) | $less(X13,alength) | X13 = X14 | sK0 = X13 (4:24) [resolution 333,198]
[SA] new: 1329. sK0 = X15 | 0 = sK0 | $less(0,X15) | $less(X16,sK0) | $less(sK0,X16) | $less(X15,alength) | alength = X15 | sK0 = X16 (4:24) [resolution 333,204]
[SA] new: 1330. sK0 = X17 | 0 = sK0 | $less(0,X17) | $less(X18,sK0) | $less(i(main_end),X18) | $less(X17,0) | X17 = X18 | i(main_end) = X17 | i(l8(sK2(X17))) = X17 (4:32) [resolution 333,214]
[SA] new: 1331. sK0 = X19 | 0 = sK0 | $less(0,X19) | $less(X20,sK0) | $less(sK0,X20) | $less(X19,alength) | X19 = X20 | alength = X19 (4:24) [resolution 333,246]
[SA] new: 1332. sK0 = X21 | 0 = sK0 | $less(0,X21) | $less(X22,sK0) | $less(sK0,X22) | $less(X21,alength) | alength = X21 | alength = X22 (4:24) [resolution 333,253]
[SA] new: 1333. alength = sK0 | 0 = sK0 | $less(0,alength) | $less(X23,sK0) | $less(sK0,X23) | sK0 = X23 (4:18) [resolution 333,106]
[SA] new: 1334. alength = sK0 | 0 = sK0 | $less(0,alength) | $less(X24,sK0) | $less(0,X24) | alength = X24 | 0 = sK0 (4:21) [resolution 333,178]
[SA] new: 1335. alength = sK0 | 0 = sK0 | $less(0,alength) | $less(X25,sK0) | $less(sK0,X25) | alength = X25 (4:18) [resolution 333,179]
[SA] new: 1336. alength = sK0 | 0 = sK0 | $less(0,alength) | $less(X26,sK0) | $less(sK0,X27) | $less(X27,X26) | X26 = X27 | sK0 = X26 (4:24) [resolution 333,198]
[SA] new: 1337. alength = sK0 | 0 = sK0 | $less(0,alength) | $less(X28,sK0) | $less(sK0,X29) | $less(X29,X28) | alength = X28 | sK0 = X29 (4:24) [resolution 333,204]
[SA] new: 1338. alength = sK0 | 0 = sK0 | $less(0,alength) | $less(X30,sK0) | $less(sK0,X31) | $less(X31,X30) | X30 = X31 | alength = X30 (4:24) [resolution 333,246]
[SA] new: 1339. alength = sK0 | 0 = sK0 | $less(0,alength) | $less(X32,sK0) | $less(sK0,X33) | $less(X33,X32) | alength = X32 | alength = X33 (4:24) [resolution 333,253]
[SA] new: 1340. alength = sK0 | 0 = sK0 | $less(0,alength) | $less(X34,sK0) | $less(0,X34) | 0 = sK0 | sK0 = X34 (4:21) [resolution 333,330]
[SA] new: 1341. i(main_end) = sK0 | 0 = sK0 | $less(0,i(main_end)) | $less(X35,sK0) | $less(alength,X35) | alength = i(main_end) | alength = X35 (4:24) [resolution 333,409]
[SA] new: 1342. i(main_end) = sK0 | 0 = sK0 | $less(0,i(main_end)) | $less(X36,sK0) | $less(alength,X36) | alength = X36 | i(main_end) = X36 (4:24) [resolution 333,389]
[SA] new: 1343. i(main_end) = sK0 | 0 = sK0 | $less(0,i(main_end)) | $less(X37,sK0) | ~$less(X37,sK0) | alength = i(main_end) (5:21) [resolution 333,268]
[SA] new: 1344. i(main_end) = sK0 | 0 = sK0 | $less(0,i(main_end)) | $less(X38,sK0) | ~$less(X38,sK0) | i(main_end) = sK0 (5:21) [resolution 333,224]
[SA] new: 1345. i(main_end) = sK0 | 0 = sK0 | $less(0,i(main_end)) | $less(X39,sK0) | i(main_end) = X39 | $less(alength,X39) | alength = i(main_end) (4:25) [resolution 333,176]
[SA] new: 1346. i(main_end) = sK0 | 0 = sK0 | $less(0,i(main_end)) | $less(X40,sK0) | ~$less(X40,alength) | alength = i(main_end) (4:21) [resolution 333,144]
[SA] new: 1347. 0 = sK0 | 0 = sK0 | $less(0,0) | $less(X41,sK0) | $less(i(main_end),X42) | $less(X42,X41) | X41 = X42 | i(main_end) = X41 | i(l8(sK2(X41))) = X41 (4:32) [resolution 333,214]
[SA] new: 1348. sK0 = sK0 | 0 = sK0 | $less(0,sK0) | $less(X43,sK0) | $less(0,X43) | 0 = X43 | sK0 = X43 (4:21) [resolution 333,380]
[SA] new: 1349. sK0 = sK0 | 0 = sK0 | $less(0,sK0) | $less(X44,sK0) | $less(0,X44) | 0 = sK0 | 0 = X44 (4:21) [resolution 333,229]
[SA] new: 1350. sK0 = sK0 | 0 = sK0 | $less(0,sK0) | $less(X45,sK0) | $less(0,X45) | sK0 = X45 | 0 = sK0 (4:21) [resolution 333,177]
[SA] new: 1351. i(main_end) = sK0 | 0 = sK0 | $less(0,i(main_end)) | $less(alength,sK0) | alength = i(main_end) (4:18) [resolution 333,100]
[SA] new: 1352. 0 = sK0 | 0 = sK0 | $less(0,0) | $less(alength,sK0) | alength = i(main_end) | alength = i(l8(sK2(alength))) (4:22) [resolution 333,142]
[SA] new: 1353. alength = sK0 | 0 = sK0 | $less(0,alength) | $less(i(l8(p(nl8))),sK0) | zero = nl8 (4:18) [resolution 333,95]
[SA] new: 1354. alength = sK0 | 0 = sK0 | $less(0,alength) | $less(i(l8(p(p(nl8)))),sK0) | zero = p(nl8) | zero = nl8 (4:23) [resolution 333,314]
[SA] new: 1355. alength = sK0 | 0 = sK0 | $less(0,alength) | $less(i(l8(sK2(X46))),sK0) | $less(i(main_end),X46) | i(main_end) = X46 | $less(X46,0) (4:26) [resolution 333,195]
[SA] new: 1356. alength = sK0 | 0 = sK0 | $less(0,alength) | $less(i(l8(sK2(sK0))),sK0) | alength = i(main_end) (5:19) [resolution 333,272]
[SA] new: 1357. alength = sK0 | 0 = sK0 | $less(0,alength) | $less(i(l8(sK2(sK0))),sK0) | i(main_end) = sK0 (5:19) [resolution 333,228]
[SA] new: 1358. alength = sK0 | 0 = sK0 | $less(0,alength) | $less(i(l8(sK2(alength))),sK0) | alength = i(main_end) | $less(alength,0) (4:22) [resolution 333,270]
[SA] new: 1359. sK0 = X47 | 0 = sK0 | $less(0,X47) | $less(i(main_end),sK0) | i(l8(sK1(X48,sK2(X47)))) = X48 | $less(X47,0) | a(main_end,X48) = a(l11(sK2(X47)),X48) | i(main_end) = X47 (4:37) [resolution 333,225]
[SA] new: 1360. sK0 = X49 | 0 = sK0 | $less(0,X49) | $less(i(main_end),sK0) | $less(X49,X50) | $less(X50,0) | X49 = X50 | i(main_end) = X50 | i(l8(sK2(X50))) = X50 (4:32) [resolution 333,214]
[SA] new: 1361. sK0 = X51 | 0 = sK0 | $less(0,X51) | $less(i(main_end),sK0) | i(l8(sK2(X51))) = X51 | $less(X51,0) | i(main_end) = X51 (4:26) [resolution 333,127]
[SA] new: 1362. sK0 = X52 | 0 = sK0 | $less(0,X52) | $less(i(l8(p(nl8))),sK0) | alength = X52 | $less(X52,alength) | zero = nl8 (4:24) [resolution 333,193]
[SA] new: 1363. sK0 = X53 | 0 = sK0 | $less(0,X53) | $less(0,sK0) | $less(X53,sK0) | 0 = X53 | sK0 = X53 (4:21) [resolution 333,380]
[SA] new: 1364. sK0 = X54 | 0 = sK0 | $less(0,X54) | $less(0,sK0) | $less(X54,alength) | 0 = sK0 | sK0 = X54 (4:21) [resolution 333,330]
[SA] new: 1365. sK0 = X55 | 0 = sK0 | $less(0,X55) | $less(0,sK0) | $less(X55,sK0) | 0 = sK0 | 0 = X55 (4:21) [resolution 333,229]
[SA] new: 1366. sK0 = X56 | 0 = sK0 | $less(0,X56) | $less(0,sK0) | $less(X56,alength) | alength = X56 | 0 = sK0 (4:21) [resolution 333,178]
[SA] new: 1367. sK0 = X57 | 0 = sK0 | $less(0,X57) | $less(0,sK0) | $less(X57,sK0) | sK0 = X57 | 0 = sK0 (4:21) [resolution 333,177]
[SA] new: 1368. alength = sK0 | 0 = sK0 | $less(0,alength) | $less(0,sK0) | alength = 0 (4:15) [resolution 333,244]
[SA] new: 1369. alength = sK0 | 0 = sK0 | $less(0,alength) | $less(0,sK0) | 0 = sK0 (4:15) [resolution 333,113]
[SA] new: 1370. i(main_end) = sK0 | 0 = sK0 | $less(0,i(main_end)) | $less(0,sK0) | i(main_end) = sK0 | 0 = sK0 (7:21) [resolution 333,761]
[SA] new: 1371. i(main_end) = sK0 | 0 = sK0 | $less(0,i(main_end)) | $less(0,sK0) | alength = i(main_end) | 0 = sK0 (5:21) [resolution 333,583]
[SA] new: 1372. sK0 = sK0 | 0 = sK0 | $less(0,sK0) | $less(0,sK0) | 0 = sK0 (4:15) [resolution 333,97]
[SA] new: 1373. alength = sK0 | 0 = sK0 | $less(0,alength) | $less(sK0,sK0) (4:12) [resolution 333,45]
[SA] new: 1374. sK0 = X58 | 0 = sK0 | $less(0,X58) | $less(sK0,sK0) | $less(X59,alength) | $less(X58,X59) | alength = X59 | alength = X58 (4:24) [resolution 333,253]
[SA] new: 1375. sK0 = X60 | 0 = sK0 | $less(0,X60) | $less(sK0,sK0) | $less(X60,X61) | $less(X61,alength) | X60 = X61 | alength = X61 (4:24) [resolution 333,246]
[SA] new: 1376. sK0 = X62 | 0 = sK0 | $less(0,X62) | $less(sK0,sK0) | $less(X63,alength) | $less(X62,X63) | alength = X63 | sK0 = X62 (4:24) [resolution 333,204]
[SA] new: 1377. sK0 = X64 | 0 = sK0 | $less(0,X64) | $less(sK0,sK0) | $less(X64,X65) | $less(X65,alength) | X64 = X65 | sK0 = X65 (4:24) [resolution 333,198]
[SA] new: 1378. sK0 = X66 | 0 = sK0 | $less(0,X66) | $less(sK0,sK0) | $less(X66,alength) | alength = X66 (4:18) [resolution 333,179]
[SA] new: 1379. sK0 = X67 | 0 = sK0 | $less(0,X67) | $less(sK0,sK0) | $less(X67,alength) | sK0 = X67 (4:18) [resolution 333,106]
[SA] new: 1380. i(main_end) = sK0 | 0 = sK0 | $less(0,i(main_end)) | $less(sK0,sK0) | alength = i(main_end) (4:18) [resolution 333,255]
[SA] new: 1381. i(main_end) = sK0 | 0 = sK0 | $less(0,i(main_end)) | $less(sK0,sK0) | i(main_end) = sK0 (4:18) [resolution 333,206]
[SA] new: 1382. i(main_end) = sK0 | 0 = sK0 | $less(0,i(main_end)) | $less(sK0,sK0) (4:14) [duplicate literal removal 1381]
[SA] new: 1383. sK0 = X67 | 0 = sK0 | $less(0,X67) | $less(sK0,sK0) | $less(X67,alength) (4:15) [duplicate literal removal 1379]
[SA] new: 1384. sK0 = X62 | 0 = sK0 | $less(0,X62) | $less(sK0,sK0) | $less(X63,alength) | $less(X62,X63) | alength = X63 (4:21) [duplicate literal removal 1376]
[SA] new: 1385. sK0 = sK0 | 0 = sK0 | $less(0,sK0) (4:9) [duplicate literal removal 1372]
[SA] new: 1386. i(main_end) = sK0 | 0 = sK0 | $less(0,i(main_end)) | $less(0,sK0) | alength = i(main_end) (5:18) [duplicate literal removal 1371]
[SA] new: 1387. i(main_end) = sK0 | 0 = sK0 | $less(0,i(main_end)) | $less(0,sK0) (7:14) [duplicate literal removal 1370]
[SA] new: 1388. alength = sK0 | 0 = sK0 | $less(0,alength) | $less(0,sK0) (4:12) [duplicate literal removal 1369]
[SA] new: 1389. sK0 = X57 | 0 = sK0 | $less(0,X57) | $less(0,sK0) | $less(X57,sK0) (4:15) [duplicate literal removal 1367]
[SA] new: 1390. sK0 = X56 | 0 = sK0 | $less(0,X56) | $less(0,sK0) | $less(X56,alength) | alength = X56 (4:18) [duplicate literal removal 1366]
[SA] new: 1391. sK0 = X55 | 0 = sK0 | $less(0,X55) | $less(0,sK0) | $less(X55,sK0) | 0 = X55 (4:18) [duplicate literal removal 1365]
[SA] new: 1392. sK0 = X54 | 0 = sK0 | $less(0,X54) | $less(0,sK0) | $less(X54,alength) (4:15) [duplicate literal removal 1364]
[SA] new: 1393. sK0 = X53 | 0 = sK0 | $less(0,X53) | $less(0,sK0) | $less(X53,sK0) | 0 = X53 (4:18) [duplicate literal removal 1363]
[SA] new: 1394. 0 = sK0 | $less(0,0) | $less(alength,sK0) | alength = i(main_end) | alength = i(l8(sK2(alength))) (4:19) [duplicate literal removal 1352]
[SA] new: 1395. 0 = sK0 | $less(alength,sK0) | alength = i(main_end) | alength = i(l8(sK2(alength))) (4:16) [evaluation 1394]
[SA] new: 1396. sK0 = sK0 | 0 = sK0 | $less(0,sK0) | $less(X45,sK0) | $less(0,X45) | sK0 = X45 (4:18) [duplicate literal removal 1350]
[SA] new: 1397. sK0 = sK0 | 0 = sK0 | $less(0,sK0) | $less(X44,sK0) | $less(0,X44) | 0 = X44 (4:18) [duplicate literal removal 1349]
[SA] new: 1398. 0 = sK0 | $less(0,0) | $less(X41,sK0) | $less(i(main_end),X42) | $less(X42,X41) | X41 = X42 | i(main_end) = X41 | i(l8(sK2(X41))) = X41 (4:29) [duplicate literal removal 1347]
[SA] new: 1399. 0 = sK0 | $less(X41,sK0) | $less(i(main_end),X42) | $less(X42,X41) | X41 = X42 | i(main_end) = X41 | i(l8(sK2(X41))) = X41 (4:26) [evaluation 1398]
[SA] new: 1400. i(main_end) = sK0 | 0 = sK0 | $less(0,i(main_end)) | $less(X38,sK0) | ~$less(X38,sK0) (5:17) [duplicate literal removal 1344]
[SA] new: 1401. alength = sK0 | 0 = sK0 | $less(0,alength) | $less(X34,sK0) | $less(0,X34) | sK0 = X34 (4:18) [duplicate literal removal 1340]
[SA] new: 1402. alength = sK0 | 0 = sK0 | $less(0,alength) | $less(X24,sK0) | $less(0,X24) | alength = X24 (4:18) [duplicate literal removal 1334]
[SA] new: 1403. sK0 = X13 | 0 = sK0 | $less(0,X13) | $less(X14,sK0) | $less(sK0,X14) | $less(X13,alength) | X13 = X14 (4:21) [duplicate literal removal 1328]
[SA] passive: 1323. sK0 = X0 | 0 = sK0 | $less(0,X0) | $less(X1,sK0) | $less(X0,X1) | X0 = X1 (4:18) [resolution 333,15]
[SA] passive: 1325. sK0 = X4 | 0 = sK0 | $less(0,X4) | $less(X5,sK0) | $less(X6,X5) | $less(X4,X6) | X4 = X5 | X4 = X6 (4:24) [resolution 333,174]
[SA] passive: 1326. sK0 = X7 | 0 = sK0 | $less(0,X7) | $less(X8,sK0) | $less(X7,X9) | $less(X9,X8) | X8 = X9 | X7 = X8 (4:24) [resolution 333,174]
[SA] passive: 1327. sK0 = X10 | 0 = sK0 | $less(0,X10) | $less(X11,sK0) | $less(X12,X11) | $less(X10,X12) | X10 = X12 | X11 = X12 (4:24) [resolution 333,174]
[SA] passive: 1330. sK0 = X17 | 0 = sK0 | $less(0,X17) | $less(X18,sK0) | $less(i(main_end),X18) | $less(X17,0) | X17 = X18 | i(main_end) = X17 | i(l8(sK2(X17))) = X17 (4:32) [resolution 333,214]
[SA] passive: 1359. sK0 = X47 | 0 = sK0 | $less(0,X47) | $less(i(main_end),sK0) | i(l8(sK1(X48,sK2(X47)))) = X48 | $less(X47,0) | a(main_end,X48) = a(l11(sK2(X47)),X48) | i(main_end) = X47 (4:37) [resolution 333,225]
[SA] passive: 1360. sK0 = X49 | 0 = sK0 | $less(0,X49) | $less(i(main_end),sK0) | $less(X49,X50) | $less(X50,0) | X49 = X50 | i(main_end) = X50 | i(l8(sK2(X50))) = X50 (4:32) [resolution 333,214]
[SA] passive: 1361. sK0 = X51 | 0 = sK0 | $less(0,X51) | $less(i(main_end),sK0) | i(l8(sK2(X51))) = X51 | $less(X51,0) | i(main_end) = X51 (4:26) [resolution 333,127]
[SA] active: 261. $less(0,X0) | $less(X0,alength) | alength = X0 | alength = 0 (4:12:2) [resolution 244,99]
[SA] new: 1404. $less(0,alength) | alength = 0 | alength = 0 (5:9) [resolution 261,13]
[SA] new: 1405. $less(X0,alength) | alength = X0 | alength = 0 | sK0 = X0 | 0 = sK0 | $less(0,X0) | $less(0,sK0) (5:21) [resolution 261,333]
[SA] new: 1406. $less(X1,alength) | alength = X1 | alength = 0 | alength = i(main_end) | ~$less(X1,alength) | $less(0,i(main_end)) (5:20) [resolution 261,277]
[SA] new: 1407. $less(X2,alength) | alength = X2 | alength = 0 | alength = X2 | $less(sK0,X2) | $less(0,alength) (5:18) [resolution 261,254]
[SA] new: 1408. $less(X3,alength) | alength = X3 | alength = 0 | sK0 = X3 | $less(sK0,X3) | $less(0,alength) (5:18) [resolution 261,205]
[SA] new: 1409. $less(X4,alength) | alength = X4 | alength = 0 | X4 = X5 | $less(X5,X4) | $less(0,X5) (5:18) [resolution 261,99]
[SA] new: 1410. $less(X6,alength) | alength = X6 | alength = 0 | ~$less(X7,0) | $less(X7,X6) (5:15) [resolution 261,14]
[SA] new: 1411. $less(i(main_end),alength) | alength = i(main_end) | alength = 0 | $less(X8,0) | i(main_end) = X8 | i(l8(sK2(X8))) = X8 | $less(0,X8) (5:27) [resolution 261,215]
[SA] new: 1412. $less(i(main_end),alength) | alength = i(main_end) | alength = 0 | $less(0,0) | i(l8(sK2(0))) = 0 (5:20) [resolution 261,84]
[SA] new: 1413. $less(i(main_end),alength) | alength = i(main_end) | alength = 0 | $less(0,0) | 'Sub'(sK2(0),nl8) (5:18) [resolution 261,83]
[SA] new: 1414. $less(i(l8(p(nl8))),alength) | alength = i(l8(p(nl8))) | alength = 0 | zero = nl8 | $less(0,alength) (5:21) [resolution 261,194]
[SA] new: 1415. $less(0,alength) | alength = 0 | alength = 0 | sK0 = X9 | 0 = sK0 | $less(X9,sK0) | $less(0,X9) (5:21) [resolution 261,324]
[SA] new: 1416. $less(0,alength) | alength = 0 | alength = 0 | alength = 0 | $less(0,alength) (5:15) [resolution 261,262]
[SA] new: 1417. $less(0,alength) | alength = 0 | alength = 0 | 0 = sK0 | $less(0,alength) (5:15) [resolution 261,116]
[SA] new: 1418. $less(0,alength) | alength = 0 | alength = 0 | 0 = sK0 | $less(0,sK0) (5:15) [resolution 261,114]
[SA] new: 1419. $less(sK0,alength) | alength = sK0 | alength = 0 | alength = X10 | $less(X10,alength) | $less(0,X10) (5:18) [resolution 261,247]
[SA] new: 1420. $less(sK0,alength) | alength = sK0 | alength = 0 | sK0 = X11 | $less(X11,alength) | $less(0,X11) (5:18) [resolution 261,199]
[SA] new: 1421. $less(sK0,alength) | alength = sK0 | alength = 0 | $less(0,alength) (5:12) [resolution 261,96]
[SA] new: 1422. $less(0,alength) | alength = alength | alength = 0 (5:9) [resolution 261,13]
[SA] new: 1423. $less(0,X0) | alength = X0 | alength = 0 | alength = sK0 | 0 = sK0 | $less(0,alength) | $less(X0,sK0) (5:21) [resolution 261,333]
[SA] new: 1424. $less(0,X1) | alength = X1 | alength = 0 | alength = i(main_end) | ~$less(alength,alength) | $less(X1,i(main_end)) (5:20) [resolution 261,277]
[SA] new: 1425. $less(0,X2) | alength = X2 | alength = 0 | alength = alength | $less(sK0,alength) | $less(X2,alength) (5:18) [resolution 261,254]
[SA] new: 1426. $less(0,X3) | alength = X3 | alength = 0 | alength = sK0 | $less(sK0,alength) | $less(X3,alength) (5:18) [resolution 261,205]
[SA] new: 1427. $less(0,X4) | alength = X4 | alength = 0 | alength = X5 | $less(X5,alength) | $less(X4,X5) (5:18) [resolution 261,99]
[SA] new: 1428. $less(0,X6) | alength = X6 | alength = 0 | ~$less(X7,X6) | $less(X7,alength) (5:15) [resolution 261,14]
[SA] new: 1429. $less(0,i(main_end)) | alength = i(main_end) | alength = 0 (5:11) [resolution 261,82]
[SA] new: 1430. $less(0,alength) | alength = 0 | 0 = sK0 | $less(0,sK0) (5:12) [duplicate literal removal 1418]
[SA] new: 1431. $less(0,alength) | alength = 0 | 0 = sK0 (5:9) [duplicate literal removal 1417]
[SA] new: 1432. $less(0,alength) | alength = 0 (5:6) [duplicate literal removal 1416]
[SA] new: 1433. $less(0,alength) | alength = 0 | sK0 = X9 | 0 = sK0 | $less(X9,sK0) | $less(0,X9) (5:18) [duplicate literal removal 1415]
[SA] new: 1434. $less(i(main_end),alength) | alength = i(main_end) | alength = 0 | 'Sub'(sK2(0),nl8) (5:15) [evaluation 1413]
[SA] new: 1435. $less(i(main_end),alength) | alength = i(main_end) | alength = 0 | i(l8(sK2(0))) = 0 (5:17) [evaluation 1412]
[SA] new: 1436. $less(X2,alength) | alength = X2 | alength = 0 | $less(sK0,X2) | $less(0,alength) (5:15) [duplicate literal removal 1407]
[SA] new: 1437. $less(0,alength) | alength = 0 (5:6) [duplicate literal removal 1404]
[SA] passive: 1409. $less(X4,alength) | alength = X4 | alength = 0 | X4 = X5 | $less(X5,X4) | $less(0,X5) (5:18) [resolution 261,99]
[SA] passive: 1410. $less(X6,alength) | alength = X6 | alength = 0 | ~$less(X7,0) | $less(X7,X6) (5:15) [resolution 261,14]
[SA] new: 1438. alength = i(main_end) | alength = 0 | $less(X8,0) | i(main_end) = X8 | i(l8(sK2(X8))) = X8 | $less(0,X8) (5:23) [subsumption resolution 1411,82]
[SA] passive: 1438. alength = i(main_end) | alength = 0 | $less(X8,0) | i(main_end) = X8 | i(l8(sK2(X8))) = X8 | $less(0,X8) (5:23) [subsumption resolution 1411,82]
[SA] new: 1439. alength = i(main_end) | alength = 0 | i(l8(sK2(0))) = 0 (5:13) [subsumption resolution 1435,82]
[SA] passive: 1439. alength = i(main_end) | alength = 0 | i(l8(sK2(0))) = 0 (5:13) [subsumption resolution 1435,82]
[SA] new: 1440. alength = i(main_end) | alength = 0 | 'Sub'(sK2(0),nl8) (5:11) [subsumption resolution 1434,82]
[SA] passive: 1440. alength = i(main_end) | alength = 0 | 'Sub'(sK2(0),nl8) (5:11) [subsumption resolution 1434,82]
[SA] passive: 1427. $less(0,X4) | alength = X4 | alength = 0 | alength = X5 | $less(X5,alength) | $less(X4,X5) (5:18) [resolution 261,99]
[SA] passive: 1428. $less(0,X6) | alength = X6 | alength = 0 | ~$less(X7,X6) | $less(X7,alength) (5:15) [resolution 261,14]
[SA] passive: 1429. $less(0,i(main_end)) | alength = i(main_end) | alength = 0 (5:11) [resolution 261,82]
[SA] active: 336. 'Sub'(sK2(0),nl8) | i(main_end) = sK0 | 0 = sK0 | $less(i(main_end),sK0) (3:15:1) [evaluation 326]
[SA] new: 1441. i(main_end) = sK0 | 0 = sK0 | $less(i(main_end),sK0) | a(main_end,X0) = a(l11(sK2(0)),X0) | i(l8(sK1(X0,sK2(0)))) = X0 (4:28) [resolution 336,150]
[SA] new: 1442. i(main_end) = sK0 | 0 = sK0 | $less(i(main_end),sK0) | $less(i(l8(sK2(0))),alength) (4:17) [resolution 336,56]
[SA] active: 1429. $less(0,i(main_end)) | alength = i(main_end) | alength = 0 (5:11:1) [resolution 261,82]
[SA] new: 1443. alength = i(main_end) | alength = 0 | $less(X0,0) | i(main_end) = X0 | i(l8(sK2(X0))) = X0 | $less(0,X0) (6:23) [resolution 1429,215]
[SA] new: 1444. alength = i(main_end) | alength = 0 | $less(0,0) | i(l8(sK2(0))) = 0 (6:16) [resolution 1429,84]
[SA] new: 1445. alength = i(main_end) | alength = 0 | $less(0,0) | 'Sub'(sK2(0),nl8) (6:14) [resolution 1429,83]
[SA] new: 1446. alength = i(main_end) | alength = 0 | i(main_end) = sK0 | 0 = sK0 | $less(0,i(main_end)) | $less(0,sK0) (6:21) [resolution 1429,333]
[SA] new: 1447. alength = i(main_end) | alength = 0 | alength = i(main_end) | ~$less(i(main_end),alength) | $less(0,i(main_end)) (6:19) [resolution 1429,277]
[SA] new: 1448. alength = i(main_end) | alength = 0 | alength = i(main_end) | $less(sK0,i(main_end)) | $less(0,alength) (6:18) [resolution 1429,254]
[SA] new: 1449. alength = i(main_end) | alength = 0 | i(main_end) = sK0 | $less(sK0,i(main_end)) | $less(0,alength) (6:18) [resolution 1429,205]
[SA] new: 1450. alength = i(main_end) | alength = 0 | i(main_end) = X1 | $less(X1,i(main_end)) | $less(0,X1) (6:18) [resolution 1429,99]
[SA] new: 1451. alength = i(main_end) | alength = 0 | ~$less(X2,0) | $less(X2,i(main_end)) (6:14) [resolution 1429,14]
[SA] new: 1452. alength = i(main_end) | alength = 0 | $less(sK0,i(main_end)) | $less(0,alength) (6:14) [duplicate literal removal 1448]
[SA] new: 1453. alength = i(main_end) | alength = 0 | ~$less(i(main_end),alength) | $less(0,i(main_end)) (6:15) [duplicate literal removal 1447]
[SA] new: 1454. alength = i(main_end) | alength = 0 | 'Sub'(sK2(0),nl8) (6:11) [evaluation 1445]
[SA] new: 1455. alength = i(main_end) | alength = 0 | i(l8(sK2(0))) = 0 (6:13) [evaluation 1444]
[SA] passive: 1450. alength = i(main_end) | alength = 0 | i(main_end) = X1 | $less(X1,i(main_end)) | $less(0,X1) (6:18) [resolution 1429,99]
[SA] passive: 1451. alength = i(main_end) | alength = 0 | ~$less(X2,0) | $less(X2,i(main_end)) (6:14) [resolution 1429,14]
[SA] active: 388. $less(sK0,X19) | $less(X20,alength) | $less(X19,X20) | X19 = X20 | sK0 = X19 (3:15:3) [resolution 174,96]
[SA] new: 1456. $less(X0,alength) | $less(0,X0) | 0 = X0 | 0 = sK0 (4:12) [resolution 388,46]
[SA] new: 1457. $less(X1,alength) | $less(sK0,X1) | sK0 = X1 | sK0 = sK0 (4:12) [resolution 388,13]
[SA] new: 1458. $less(X2,alength) | $less(X3,X2) | X2 = X3 | sK0 = X3 | sK0 = X3 | 0 = sK0 | $less(0,X3) | $less(sK0,sK0) (4:24) [resolution 388,333]
[SA] new: 1459. $less(X4,alength) | $less(X5,X4) | X4 = X5 | sK0 = X5 | alength = i(main_end) | ~$less(X5,alength) | $less(sK0,i(main_end)) (4:23) [resolution 388,277]
[SA] new: 1460. $less(X6,alength) | $less(X7,X6) | X6 = X7 | sK0 = X7 | alength = X7 | $less(sK0,X7) | $less(sK0,alength) (4:21) [resolution 388,254]
[SA] new: 1461. $less(X8,alength) | $less(X9,X8) | X8 = X9 | sK0 = X9 | sK0 = X9 | $less(sK0,X9) | $less(sK0,alength) (4:21) [resolution 388,205]
[SA] new: 1462. $less(X10,alength) | $less(X11,X10) | X10 = X11 | sK0 = X11 | X11 = X12 | $less(X12,X11) | $less(sK0,X12) (4:21) [resolution 388,99]
[SA] new: 1463. $less(X13,alength) | $less(X14,X13) | X13 = X14 | sK0 = X14 | ~$less(X15,sK0) | $less(X15,X14) (4:18) [resolution 388,14]
[SA] new: 1464. $less(X16,alength) | $less(i(main_end),X16) | i(main_end) = X16 | i(main_end) = sK0 | $less(X17,0) | i(main_end) = X17 | i(l8(sK2(X17))) = X17 | $less(sK0,X17) (4:31) [resolution 388,215]
[SA] new: 1465. $less(X18,alength) | $less(i(main_end),X18) | i(main_end) = X18 | i(main_end) = sK0 | $less(sK0,0) | i(l8(sK2(sK0))) = sK0 (4:24) [resolution 388,84]
[SA] new: 1466. $less(X19,alength) | $less(i(main_end),X19) | i(main_end) = X19 | i(main_end) = sK0 | $less(sK0,0) | 'Sub'(sK2(sK0),nl8) (4:22) [resolution 388,83]
[SA] new: 1467. $less(X20,alength) | $less(i(l8(p(nl8))),X20) | i(l8(p(nl8))) = X20 | i(l8(p(nl8))) = sK0 | zero = nl8 | $less(sK0,alength) (4:27) [resolution 388,194]
[SA] new: 1468. $less(X21,alength) | $less(0,X21) | 0 = X21 | 0 = sK0 | sK0 = X22 | 0 = sK0 | $less(X22,sK0) | $less(sK0,X22) (4:24) [resolution 388,324]
[SA] new: 1469. $less(X23,alength) | $less(0,X23) | 0 = X23 | 0 = sK0 | alength = 0 | $less(sK0,alength) (5:18) [resolution 388,262]
[SA] new: 1470. $less(X24,alength) | $less(0,X24) | 0 = X24 | 0 = sK0 | 0 = sK0 | $less(sK0,alength) (4:18) [resolution 388,116]
[SA] new: 1471. $less(X25,alength) | $less(0,X25) | 0 = X25 | 0 = sK0 | 0 = sK0 | $less(sK0,sK0) (4:18) [resolution 388,114]
[SA] new: 1472. $less(X26,alength) | $less(sK0,X26) | sK0 = X26 | sK0 = sK0 | alength = X27 | $less(X27,alength) | $less(sK0,X27) (4:21) [resolution 388,247]
[SA] new: 1473. $less(X28,alength) | $less(sK0,X28) | sK0 = X28 | sK0 = sK0 | sK0 = X29 | $less(X29,alength) | $less(sK0,X29) (4:21) [resolution 388,199]
[SA] new: 1474. $less(X30,alength) | $less(sK0,X30) | sK0 = X30 | sK0 = sK0 | $less(sK0,alength) (4:15) [resolution 388,96]
[SA] new: 1475. $less(sK0,X0) | $less(X0,alength) | alength = X0 | sK0 = X0 (4:12) [resolution 388,13]
[SA] new: 1476. $less(sK0,X1) | $less(X1,X2) | X1 = X2 | sK0 = X1 | alength = sK0 | 0 = sK0 | $less(0,alength) | $less(X2,sK0) (4:24) [resolution 388,333]
[SA] new: 1477. $less(sK0,X3) | $less(X3,X4) | X3 = X4 | sK0 = X3 | alength = i(main_end) | ~$less(alength,alength) | $less(X4,i(main_end)) (4:23) [resolution 388,277]
[SA] new: 1478. $less(sK0,X5) | $less(X5,X6) | X5 = X6 | sK0 = X5 | alength = alength | $less(sK0,alength) | $less(X6,alength) (4:21) [resolution 388,254]
[SA] new: 1479. $less(sK0,X7) | $less(X7,X8) | X7 = X8 | sK0 = X7 | alength = sK0 | $less(sK0,alength) | $less(X8,alength) (4:21) [resolution 388,205]
[SA] new: 1480. $less(sK0,X9) | $less(X9,X10) | X9 = X10 | sK0 = X9 | alength = X11 | $less(X11,alength) | $less(X10,X11) (4:21) [resolution 388,99]
[SA] new: 1481. $less(sK0,X12) | $less(X12,X13) | X12 = X13 | sK0 = X12 | ~$less(X14,X13) | $less(X14,alength) (4:18) [resolution 388,14]
[SA] new: 1482. $less(sK0,X15) | $less(X15,i(main_end)) | i(main_end) = X15 | sK0 = X15 (4:14) [resolution 388,82]
[SA] new: 1483. $less(sK0,sK0) | $less(0,alength) | 0 = sK0 | sK0 = sK0 (4:12) [resolution 388,46]
[SA] new: 1484. $less(sK0,X0) | $less(X0,alength) | X0 = X0 | sK0 = X0 (4:12) [resolution 388,13]
[SA] new: 1485. $less(sK0,X1) | $less(X2,alength) | X1 = X2 | sK0 = X1 | sK0 = X2 | 0 = sK0 | $less(0,X2) | $less(X1,sK0) (4:24) [resolution 388,333]
[SA] new: 1486. $less(sK0,X3) | $less(X4,alength) | X3 = X4 | sK0 = X3 | alength = i(main_end) | ~$less(X4,alength) | $less(X3,i(main_end)) (4:23) [resolution 388,277]
[SA] new: 1487. $less(sK0,X5) | $less(X6,alength) | X5 = X6 | sK0 = X5 | alength = X6 | $less(sK0,X6) | $less(X5,alength) (4:21) [resolution 388,254]
[SA] new: 1488. $less(sK0,X7) | $less(X8,alength) | X7 = X8 | sK0 = X7 | sK0 = X8 | $less(sK0,X8) | $less(X7,alength) (4:21) [resolution 388,205]
[SA] new: 1489. $less(sK0,X9) | $less(X10,alength) | X9 = X10 | sK0 = X9 | X10 = X11 | $less(X11,X10) | $less(X9,X11) (4:21) [resolution 388,99]
[SA] new: 1490. $less(sK0,X12) | $less(X13,alength) | X12 = X13 | sK0 = X12 | ~$less(X14,X12) | $less(X14,X13) (4:18) [resolution 388,14]
[SA] new: 1491. $less(sK0,X15) | $less(i(main_end),alength) | i(main_end) = X15 | sK0 = X15 | $less(X16,0) | i(main_end) = X16 | i(l8(sK2(X16))) = X16 | $less(X15,X16) (4:30) [resolution 388,215]
[SA] new: 1492. $less(sK0,X17) | $less(i(main_end),alength) | i(main_end) = X17 | sK0 = X17 | $less(X17,0) | i(l8(sK2(X17))) = X17 (4:23) [resolution 388,84]
[SA] new: 1493. $less(sK0,X18) | $less(i(main_end),alength) | i(main_end) = X18 | sK0 = X18 | $less(X18,0) | 'Sub'(sK2(X18),nl8) (4:21) [resolution 388,83]
[SA] new: 1494. $less(sK0,X19) | $less(i(l8(p(nl8))),alength) | i(l8(p(nl8))) = X19 | sK0 = X19 | zero = nl8 | $less(X19,alength) (4:24) [resolution 388,194]
[SA] new: 1495. $less(sK0,X20) | $less(0,alength) | 0 = X20 | sK0 = X20 | sK0 = X21 | 0 = sK0 | $less(X21,sK0) | $less(X20,X21) (4:24) [resolution 388,324]
[SA] new: 1496. $less(sK0,X22) | $less(0,alength) | 0 = X22 | sK0 = X22 | alength = 0 | $less(X22,alength) (5:18) [resolution 388,262]
[SA] new: 1497. $less(sK0,X23) | $less(0,alength) | 0 = X23 | sK0 = X23 | 0 = sK0 | $less(X23,alength) (4:18) [resolution 388,116]
[SA] new: 1498. $less(sK0,X24) | $less(0,alength) | 0 = X24 | sK0 = X24 | 0 = sK0 | $less(X24,sK0) (4:18) [resolution 388,114]
[SA] new: 1499. $less(sK0,X25) | $less(sK0,alength) | sK0 = X25 | sK0 = X25 | alength = X26 | $less(X26,alength) | $less(X25,X26) (4:21) [resolution 388,247]
[SA] new: 1500. $less(sK0,X27) | $less(sK0,alength) | sK0 = X27 | sK0 = X27 | sK0 = X28 | $less(X28,alength) | $less(X27,X28) (4:21) [resolution 388,199]
[SA] new: 1501. $less(sK0,X29) | $less(sK0,alength) | sK0 = X29 | sK0 = X29 | $less(X29,alength) (4:15) [resolution 388,96]
[SA] new: 1502. $less(sK0,i(main_end)) | $less(alength,alength) | alength = i(main_end) | i(main_end) = sK0 (4:15) [resolution 388,82]
[SA] new: 1503. $less(sK0,i(main_end)) | $less(sK0,alength) | i(main_end) = sK0 | i(main_end) = sK0 | alength = i(main_end) (6:19) [resolution 388,812]
[SA] new: 1504. $less(sK0,i(main_end)) | $less(sK0,alength) | i(main_end) = sK0 | i(main_end) = sK0 | i(main_end) = sK0 (6:19) [resolution 388,722]
[SA] new: 1505. $less(sK0,alength) | $less(alength,sK0) | alength = sK0 | alength = sK0 (4:12) [factoring 388]
[SA] new: 1506. $less(sK0,sK0) | $less(sK0,alength) | sK0 = sK0 | sK0 = sK0 (4:12) [factoring 388]
[SA] new: 1507. $less(sK0,alength) | $less(alength,alength) | alength = alength | alength = sK0 (4:12) [factoring 388]
[SA] new: 1508. $less(sK0,sK0) | $less(sK0,alength) | sK0 = sK0 (4:9) [duplicate literal removal 1506]
[SA] new: 1509. $less(sK0,alength) | $less(alength,sK0) | alength = sK0 (4:9) [duplicate literal removal 1505]
[SA] new: 1510. $less(sK0,i(main_end)) | $less(sK0,alength) | i(main_end) = sK0 (6:11) [duplicate literal removal 1504]
[SA] new: 1511. $less(sK0,i(main_end)) | $less(sK0,alength) | i(main_end) = sK0 | alength = i(main_end) (6:15) [duplicate literal removal 1503]
[SA] new: 1512. $less(sK0,X29) | $less(sK0,alength) | sK0 = X29 | $less(X29,alength) (4:12) [duplicate literal removal 1501]
[SA] new: 1513. $less(sK0,X27) | $less(sK0,alength) | sK0 = X27 | sK0 = X28 | $less(X28,alength) | $less(X27,X28) (4:18) [duplicate literal removal 1500]
[SA] new: 1514. $less(sK0,X25) | $less(sK0,alength) | sK0 = X25 | alength = X26 | $less(X26,alength) | $less(X25,X26) (4:18) [duplicate literal removal 1499]
[SA] new: 1515. $less(X25,alength) | $less(0,X25) | 0 = X25 | 0 = sK0 | $less(sK0,sK0) (4:15) [duplicate literal removal 1471]
[SA] new: 1516. $less(X24,alength) | $less(0,X24) | 0 = X24 | 0 = sK0 | $less(sK0,alength) (4:15) [duplicate literal removal 1470]
[SA] new: 1517. $less(X21,alength) | $less(0,X21) | 0 = X21 | 0 = sK0 | sK0 = X22 | $less(X22,sK0) | $less(sK0,X22) (4:21) [duplicate literal removal 1468]
[SA] new: 1518. $less(X8,alength) | $less(X9,X8) | X8 = X9 | sK0 = X9 | $less(sK0,X9) | $less(sK0,alength) (4:18) [duplicate literal removal 1461]
[SA] new: 1519. $less(X2,alength) | $less(X3,X2) | X2 = X3 | sK0 = X3 | 0 = sK0 | $less(0,X3) | $less(sK0,sK0) (4:21) [duplicate literal removal 1458]
[SA] new: 1520. $less(X2,alength) | $less(X3,X2) | X2 = X3 | sK0 = X3 | 0 = sK0 | $less(0,X3) (4:18) [subsumption resolution 1519,13]
[SA] passive: 1520. $less(X2,alength) | $less(X3,X2) | X2 = X3 | sK0 = X3 | 0 = sK0 | $less(0,X3) (4:18) [subsumption resolution 1519,13]
[SA] passive: 1462. $less(X10,alength) | $less(X11,X10) | X10 = X11 | sK0 = X11 | X11 = X12 | $less(X12,X11) | $less(sK0,X12) (4:21) [resolution 388,99]
[SA] passive: 1463. $less(X13,alength) | $less(X14,X13) | X13 = X14 | sK0 = X14 | ~$less(X15,sK0) | $less(X15,X14) (4:18) [resolution 388,14]
[SA] passive: 1481. $less(sK0,X12) | $less(X12,X13) | X12 = X13 | sK0 = X12 | ~$less(X14,X13) | $less(X14,alength) (4:18) [resolution 388,14]
[SA] passive: 1482. $less(sK0,X15) | $less(X15,i(main_end)) | i(main_end) = X15 | sK0 = X15 (4:14) [resolution 388,82]
[SA] passive: 1490. $less(sK0,X12) | $less(X13,alength) | X12 = X13 | sK0 = X12 | ~$less(X14,X12) | $less(X14,X13) (4:18) [resolution 388,14]
[SA] new: 1521. $less(sK0,X15) | i(main_end) = X15 | sK0 = X15 | $less(X16,0) | i(main_end) = X16 | i(l8(sK2(X16))) = X16 | $less(X15,X16) (4:26) [subsumption resolution 1491,82]
[SA] passive: 1521. $less(sK0,X15) | i(main_end) = X15 | sK0 = X15 | $less(X16,0) | i(main_end) = X16 | i(l8(sK2(X16))) = X16 | $less(X15,X16) (4:26) [subsumption resolution 1491,82]
[SA] new: 1522. $less(sK0,X18) | i(main_end) = X18 | sK0 = X18 | $less(X18,0) | 'Sub'(sK2(X18),nl8) (4:17) [subsumption resolution 1493,82]
[SA] passive: 1522. $less(sK0,X18) | i(main_end) = X18 | sK0 = X18 | $less(X18,0) | 'Sub'(sK2(X18),nl8) (4:17) [subsumption resolution 1493,82]
[SA] active: 1440. 'Sub'(sK2(0),nl8) | alength = 0 | alength = i(main_end) (5:11:1) [subsumption resolution 1434,82]
[SA] new: 1523. alength = 0 | alength = i(main_end) | a(main_end,X0) = a(l11(sK2(0)),X0) | i(l8(sK1(X0,sK2(0)))) = X0 (6:24) [resolution 1440,150]
[SA] new: 1524. alength = 0 | alength = i(main_end) | $less(i(l8(sK2(0))),alength) (6:13) [resolution 1440,56]
[SA] passive: 1523. alength = 0 | alength = i(main_end) | a(main_end,X0) = a(l11(sK2(0)),X0) | i(l8(sK1(X0,sK2(0)))) = X0 (6:24) [resolution 1440,150]
[SA] passive: 1524. alength = 0 | alength = i(main_end) | $less(i(l8(sK2(0))),alength) (6:13) [resolution 1440,56]
[SA] active: 408. $less(sK0,X19) | $less(X20,alength) | $less(X19,X20) | sK0 = X19 | sK0 = X20 (3:15:3) [resolution 174,96]
[SA] new: 1525. $less(X0,alength) | $less(0,X0) | 0 = sK0 | sK0 = X0 (4:12) [resolution 408,46]
[SA] new: 1526. $less(X1,alength) | $less(sK0,X1) | sK0 = sK0 | sK0 = X1 (4:12) [resolution 408,13]
[SA] new: 1527. $less(X2,alength) | $less(X3,X2) | sK0 = X3 | sK0 = X2 | sK0 = X3 | 0 = sK0 | $less(0,X3) | $less(sK0,sK0) (4:24) [resolution 408,333]
[SA] new: 1528. $less(X4,alength) | $less(X5,X4) | sK0 = X5 | sK0 = X4 | alength = i(main_end) | ~$less(X5,alength) | $less(sK0,i(main_end)) (4:23) [resolution 408,277]
[SA] new: 1529. $less(X6,alength) | $less(X7,X6) | sK0 = X7 | sK0 = X6 | alength = X7 | $less(sK0,X7) | $less(sK0,alength) (4:21) [resolution 408,254]
[SA] new: 1530. $less(X8,alength) | $less(X9,X8) | sK0 = X9 | sK0 = X8 | sK0 = X9 | $less(sK0,X9) | $less(sK0,alength) (4:21) [resolution 408,205]
[SA] new: 1531. $less(X10,alength) | $less(X11,X10) | sK0 = X11 | sK0 = X10 | X11 = X12 | $less(X12,X11) | $less(sK0,X12) (4:21) [resolution 408,99]
[SA] new: 1532. $less(X13,alength) | $less(X14,X13) | sK0 = X14 | sK0 = X13 | ~$less(X15,sK0) | $less(X15,X14) (4:18) [resolution 408,14]
[SA] new: 1533. $less(X16,alength) | $less(i(main_end),X16) | i(main_end) = sK0 | sK0 = X16 | $less(X17,0) | i(main_end) = X17 | i(l8(sK2(X17))) = X17 | $less(sK0,X17) (4:30) [resolution 408,215]
[SA] new: 1534. $less(X18,alength) | $less(i(main_end),X18) | i(main_end) = sK0 | sK0 = X18 | $less(sK0,0) | i(l8(sK2(sK0))) = sK0 (4:23) [resolution 408,84]
[SA] new: 1535. $less(X19,alength) | $less(i(main_end),X19) | i(main_end) = sK0 | sK0 = X19 | $less(sK0,0) | 'Sub'(sK2(sK0),nl8) (4:21) [resolution 408,83]
[SA] new: 1536. $less(X20,alength) | $less(i(l8(p(nl8))),X20) | i(l8(p(nl8))) = sK0 | sK0 = X20 | zero = nl8 | $less(sK0,alength) (4:24) [resolution 408,194]
[SA] new: 1537. $less(X21,alength) | $less(0,X21) | 0 = sK0 | sK0 = X21 | sK0 = X22 | 0 = sK0 | $less(X22,sK0) | $less(sK0,X22) (4:24) [resolution 408,324]
[SA] new: 1538. $less(X23,alength) | $less(0,X23) | 0 = sK0 | sK0 = X23 | alength = 0 | $less(sK0,alength) (5:18) [resolution 408,262]
[SA] new: 1539. $less(X24,alength) | $less(0,X24) | 0 = sK0 | sK0 = X24 | 0 = sK0 | $less(sK0,alength) (4:18) [resolution 408,116]
[SA] new: 1540. $less(X25,alength) | $less(0,X25) | 0 = sK0 | sK0 = X25 | 0 = sK0 | $less(sK0,sK0) (4:18) [resolution 408,114]
[SA] new: 1541. $less(X26,alength) | $less(sK0,X26) | sK0 = sK0 | sK0 = X26 | alength = X27 | $less(X27,alength) | $less(sK0,X27) (4:21) [resolution 408,247]
[SA] new: 1542. $less(X28,alength) | $less(sK0,X28) | sK0 = sK0 | sK0 = X28 | sK0 = X29 | $less(X29,alength) | $less(sK0,X29) (4:21) [resolution 408,199]
[SA] new: 1543. $less(X30,alength) | $less(sK0,X30) | sK0 = sK0 | sK0 = X30 | $less(sK0,alength) (4:15) [resolution 408,96]
[SA] new: 1544. $less(sK0,X0) | $less(X0,alength) | sK0 = X0 | alength = sK0 (4:12) [resolution 408,13]
[SA] new: 1545. $less(sK0,X1) | $less(X1,X2) | sK0 = X1 | sK0 = X2 | alength = sK0 | 0 = sK0 | $less(0,alength) | $less(X2,sK0) (4:24) [resolution 408,333]
[SA] new: 1546. $less(sK0,X3) | $less(X3,X4) | sK0 = X3 | sK0 = X4 | alength = i(main_end) | ~$less(alength,alength) | $less(X4,i(main_end)) (4:23) [resolution 408,277]
[SA] new: 1547. $less(sK0,X5) | $less(X5,X6) | sK0 = X5 | sK0 = X6 | alength = alength | $less(sK0,alength) | $less(X6,alength) (4:21) [resolution 408,254]
[SA] new: 1548. $less(sK0,X7) | $less(X7,X8) | sK0 = X7 | sK0 = X8 | alength = sK0 | $less(sK0,alength) | $less(X8,alength) (4:21) [resolution 408,205]
[SA] new: 1549. $less(sK0,X9) | $less(X9,X10) | sK0 = X9 | sK0 = X10 | alength = X11 | $less(X11,alength) | $less(X10,X11) (4:21) [resolution 408,99]
[SA] new: 1550. $less(sK0,X12) | $less(X12,X13) | sK0 = X12 | sK0 = X13 | ~$less(X14,X13) | $less(X14,alength) (4:18) [resolution 408,14]
[SA] new: 1551. $less(sK0,X15) | $less(X15,i(main_end)) | sK0 = X15 | i(main_end) = sK0 (4:14) [resolution 408,82]
[SA] new: 1552. $less(sK0,sK0) | $less(0,alength) | sK0 = sK0 | 0 = sK0 (4:12) [resolution 408,46]
[SA] new: 1553. $less(sK0,X0) | $less(X0,alength) | sK0 = X0 | sK0 = X0 (4:12) [resolution 408,13]
[SA] new: 1554. $less(sK0,X1) | $less(X2,alength) | sK0 = X1 | sK0 = X2 | sK0 = X2 | 0 = sK0 | $less(0,X2) | $less(X1,sK0) (4:24) [resolution 408,333]
[SA] new: 1555. $less(sK0,X3) | $less(X4,alength) | sK0 = X3 | sK0 = X4 | alength = i(main_end) | ~$less(X4,alength) | $less(X3,i(main_end)) (4:23) [resolution 408,277]
[SA] new: 1556. $less(sK0,X5) | $less(X6,alength) | sK0 = X5 | sK0 = X6 | alength = X6 | $less(sK0,X6) | $less(X5,alength) (4:21) [resolution 408,254]
[SA] new: 1557. $less(sK0,X7) | $less(X8,alength) | sK0 = X7 | sK0 = X8 | sK0 = X8 | $less(sK0,X8) | $less(X7,alength) (4:21) [resolution 408,205]
[SA] new: 1558. $less(sK0,X9) | $less(X10,alength) | sK0 = X9 | sK0 = X10 | X10 = X11 | $less(X11,X10) | $less(X9,X11) (4:21) [resolution 408,99]
[SA] new: 1559. $less(sK0,X12) | $less(X13,alength) | sK0 = X12 | sK0 = X13 | ~$less(X14,X12) | $less(X14,X13) (4:18) [resolution 408,14]
[SA] new: 1560. $less(sK0,X15) | $less(i(main_end),alength) | sK0 = X15 | i(main_end) = sK0 | $less(X16,0) | i(main_end) = X16 | i(l8(sK2(X16))) = X16 | $less(X15,X16) (4:30) [resolution 408,215]
[SA] new: 1561. $less(sK0,X17) | $less(i(main_end),alength) | sK0 = X17 | i(main_end) = sK0 | $less(X17,0) | i(l8(sK2(X17))) = X17 (4:23) [resolution 408,84]
[SA] new: 1562. $less(sK0,X18) | $less(i(main_end),alength) | sK0 = X18 | i(main_end) = sK0 | $less(X18,0) | 'Sub'(sK2(X18),nl8) (4:21) [resolution 408,83]
[SA] new: 1563. $less(sK0,X19) | $less(i(l8(p(nl8))),alength) | sK0 = X19 | i(l8(p(nl8))) = sK0 | zero = nl8 | $less(X19,alength) (4:24) [resolution 408,194]
[SA] new: 1564. $less(sK0,X20) | $less(0,alength) | sK0 = X20 | 0 = sK0 | sK0 = X21 | 0 = sK0 | $less(X21,sK0) | $less(X20,X21) (4:24) [resolution 408,324]
[SA] new: 1565. $less(sK0,X22) | $less(0,alength) | sK0 = X22 | 0 = sK0 | alength = 0 | $less(X22,alength) (5:18) [resolution 408,262]
[SA] new: 1566. $less(sK0,X23) | $less(0,alength) | sK0 = X23 | 0 = sK0 | 0 = sK0 | $less(X23,alength) (4:18) [resolution 408,116]
[SA] new: 1567. $less(sK0,X24) | $less(0,alength) | sK0 = X24 | 0 = sK0 | 0 = sK0 | $less(X24,sK0) (4:18) [resolution 408,114]
[SA] new: 1568. $less(sK0,X25) | $less(sK0,alength) | sK0 = X25 | sK0 = sK0 | alength = X26 | $less(X26,alength) | $less(X25,X26) (4:21) [resolution 408,247]
[SA] new: 1569. $less(sK0,X27) | $less(sK0,alength) | sK0 = X27 | sK0 = sK0 | sK0 = X28 | $less(X28,alength) | $less(X27,X28) (4:21) [resolution 408,199]
[SA] new: 1570. $less(sK0,X29) | $less(sK0,alength) | sK0 = X29 | sK0 = sK0 | $less(X29,alength) (4:15) [resolution 408,96]
[SA] new: 1571. $less(sK0,i(main_end)) | $less(alength,alength) | i(main_end) = sK0 | alength = sK0 (4:14) [resolution 408,82]
[SA] new: 1572. $less(sK0,i(main_end)) | $less(sK0,alength) | i(main_end) = sK0 | sK0 = sK0 | alength = i(main_end) (6:18) [resolution 408,812]
[SA] new: 1573. $less(sK0,i(main_end)) | $less(sK0,alength) | i(main_end) = sK0 | sK0 = sK0 | i(main_end) = sK0 (6:18) [resolution 408,722]
[SA] new: 1574. $less(sK0,alength) | $less(alength,sK0) | alength = sK0 | sK0 = sK0 (4:12) [factoring 408]
[SA] new: 1575. $less(sK0,sK0) | $less(sK0,alength) | sK0 = sK0 | sK0 = sK0 (4:12) [factoring 408]
[SA] new: 1576. $less(sK0,alength) | $less(alength,alength) | alength = sK0 | alength = sK0 (4:12) [factoring 408]
[SA] new: 1577. $less(sK0,alength) | $less(alength,alength) | alength = sK0 (4:9) [duplicate literal removal 1576]
[SA] new: 1578. $less(sK0,sK0) | $less(sK0,alength) | sK0 = sK0 (4:9) [duplicate literal removal 1575]
[SA] new: 1579. $less(sK0,i(main_end)) | $less(sK0,alength) | i(main_end) = sK0 | sK0 = sK0 (6:14) [duplicate literal removal 1573]
[SA] new: 1580. $less(sK0,X24) | $less(0,alength) | sK0 = X24 | 0 = sK0 | $less(X24,sK0) (4:15) [duplicate literal removal 1567]
[SA] new: 1581. $less(sK0,X23) | $less(0,alength) | sK0 = X23 | 0 = sK0 | $less(X23,alength) (4:15) [duplicate literal removal 1566]
[SA] new: 1582. $less(sK0,X20) | $less(0,alength) | sK0 = X20 | 0 = sK0 | sK0 = X21 | $less(X21,sK0) | $less(X20,X21) (4:21) [duplicate literal removal 1564]
[SA] new: 1583. $less(sK0,X7) | $less(X8,alength) | sK0 = X7 | sK0 = X8 | $less(sK0,X8) | $less(X7,alength) (4:18) [duplicate literal removal 1557]
[SA] new: 1584. $less(sK0,X1) | $less(X2,alength) | sK0 = X1 | sK0 = X2 | 0 = sK0 | $less(0,X2) | $less(X1,sK0) (4:21) [duplicate literal removal 1554]
[SA] new: 1585. $less(sK0,X0) | $less(X0,alength) | sK0 = X0 (4:9) [duplicate literal removal 1553]
[SA] new: 1586. $less(X25,alength) | $less(0,X25) | 0 = sK0 | sK0 = X25 | $less(sK0,sK0) (4:15) [duplicate literal removal 1540]
[SA] new: 1587. $less(X24,alength) | $less(0,X24) | 0 = sK0 | sK0 = X24 | $less(sK0,alength) (4:15) [duplicate literal removal 1539]
[SA] new: 1588. $less(X21,alength) | $less(0,X21) | 0 = sK0 | sK0 = X21 | sK0 = X22 | $less(X22,sK0) | $less(sK0,X22) (4:21) [duplicate literal removal 1537]
[SA] new: 1589. $less(X8,alength) | $less(X9,X8) | sK0 = X9 | sK0 = X8 | $less(sK0,X9) | $less(sK0,alength) (4:18) [duplicate literal removal 1530]
[SA] new: 1590. $less(X2,alength) | $less(X3,X2) | sK0 = X3 | sK0 = X2 | 0 = sK0 | $less(0,X3) | $less(sK0,sK0) (4:21) [duplicate literal removal 1527]
[SA] passive: 1532. $less(X13,alength) | $less(X14,X13) | sK0 = X14 | sK0 = X13 | ~$less(X15,sK0) | $less(X15,X14) (4:18) [resolution 408,14]
[SA] passive: 1550. $less(sK0,X12) | $less(X12,X13) | sK0 = X12 | sK0 = X13 | ~$less(X14,X13) | $less(X14,alength) (4:18) [resolution 408,14]
[SA] passive: 1558. $less(sK0,X9) | $less(X10,alength) | sK0 = X9 | sK0 = X10 | X10 = X11 | $less(X11,X10) | $less(X9,X11) (4:21) [resolution 408,99]
[SA] passive: 1559. $less(sK0,X12) | $less(X13,alength) | sK0 = X12 | sK0 = X13 | ~$less(X14,X12) | $less(X14,X13) (4:18) [resolution 408,14]
[SA] new: 1591. $less(sK0,X15) | sK0 = X15 | i(main_end) = sK0 | $less(X16,0) | i(main_end) = X16 | i(l8(sK2(X16))) = X16 | $less(X15,X16) (4:26) [subsumption resolution 1560,82]
[SA] passive: 1591. $less(sK0,X15) | sK0 = X15 | i(main_end) = sK0 | $less(X16,0) | i(main_end) = X16 | i(l8(sK2(X16))) = X16 | $less(X15,X16) (4:26) [subsumption resolution 1560,82]
[SA] new: 1592. $less(sK0,X17) | sK0 = X17 | i(main_end) = sK0 | $less(X17,0) | i(l8(sK2(X17))) = X17 (4:19) [subsumption resolution 1561,82]
[SA] passive: 1592. $less(sK0,X17) | sK0 = X17 | i(main_end) = sK0 | $less(X17,0) | i(l8(sK2(X17))) = X17 (4:19) [subsumption resolution 1561,82]
[SA] new: 1593. $less(sK0,X18) | sK0 = X18 | i(main_end) = sK0 | $less(X18,0) | 'Sub'(sK2(X18),nl8) (4:17) [subsumption resolution 1562,82]
[SA] passive: 1593. $less(sK0,X18) | sK0 = X18 | i(main_end) = sK0 | $less(X18,0) | 'Sub'(sK2(X18),nl8) (4:17) [subsumption resolution 1562,82]
[SA] active: 278. $less(0,X0) | $less(X0,alength) | 0 = sK0 | 0 = X0 (4:12:2) [resolution 116,15]
[SA] new: 1594. $less(0,alength) | 0 = sK0 | 0 = 0 (5:9) [resolution 278,13]
[SA] new: 1595. $less(X0,alength) | 0 = sK0 | 0 = X0 | sK0 = X0 | 0 = sK0 | $less(0,X0) | $less(0,sK0) (5:21) [resolution 278,333]
[SA] new: 1596. $less(X1,alength) | 0 = sK0 | 0 = X1 | alength = i(main_end) | ~$less(X1,alength) | $less(0,i(main_end)) (5:20) [resolution 278,277]
[SA] new: 1597. $less(X2,alength) | 0 = sK0 | 0 = X2 | alength = X2 | $less(sK0,X2) | $less(0,alength) (5:18) [resolution 278,254]
[SA] new: 1598. $less(X3,alength) | 0 = sK0 | 0 = X3 | sK0 = X3 | $less(sK0,X3) | $less(0,alength) (5:18) [resolution 278,205]
[SA] new: 1599. $less(X4,alength) | 0 = sK0 | 0 = X4 | X4 = X5 | $less(X5,X4) | $less(0,X5) (5:18) [resolution 278,99]
[SA] new: 1600. $less(X6,alength) | 0 = sK0 | 0 = X6 | ~$less(X7,0) | $less(X7,X6) (5:15) [resolution 278,14]
[SA] new: 1601. $less(i(main_end),alength) | 0 = sK0 | i(main_end) = 0 | $less(X8,0) | i(main_end) = X8 | i(l8(sK2(X8))) = X8 | $less(0,X8) (5:27) [resolution 278,215]
[SA] new: 1602. $less(i(main_end),alength) | 0 = sK0 | i(main_end) = 0 | $less(0,0) | i(l8(sK2(0))) = 0 (5:20) [resolution 278,84]
[SA] new: 1603. $less(i(main_end),alength) | 0 = sK0 | i(main_end) = 0 | $less(0,0) | 'Sub'(sK2(0),nl8) (5:18) [resolution 278,83]
[SA] new: 1604. $less(i(l8(p(nl8))),alength) | 0 = sK0 | i(l8(p(nl8))) = 0 | zero = nl8 | $less(0,alength) (5:21) [resolution 278,194]
[SA] new: 1605. $less(0,alength) | 0 = sK0 | 0 = 0 | sK0 = X9 | 0 = sK0 | $less(X9,sK0) | $less(0,X9) (5:21) [resolution 278,324]
[SA] new: 1606. $less(0,alength) | 0 = sK0 | 0 = 0 | alength = 0 | $less(0,alength) (5:15) [resolution 278,262]
[SA] new: 1607. $less(0,alength) | 0 = sK0 | 0 = 0 | 0 = sK0 | $less(0,alength) (5:15) [resolution 278,116]
[SA] new: 1608. $less(0,alength) | 0 = sK0 | 0 = 0 | 0 = sK0 | $less(0,sK0) (5:15) [resolution 278,114]
[SA] new: 1609. $less(sK0,alength) | 0 = sK0 | 0 = sK0 | alength = X10 | $less(X10,alength) | $less(0,X10) (5:18) [resolution 278,247]
[SA] new: 1610. $less(sK0,alength) | 0 = sK0 | 0 = sK0 | sK0 = X11 | $less(X11,alength) | $less(0,X11) (5:18) [resolution 278,199]
[SA] new: 1611. $less(sK0,alength) | 0 = sK0 | 0 = sK0 | $less(0,alength) (5:12) [resolution 278,96]
[SA] new: 1612. $less(0,alength) | 0 = sK0 | alength = 0 (5:9) [resolution 278,13]
[SA] new: 1613. $less(0,X0) | 0 = sK0 | 0 = X0 | alength = sK0 | 0 = sK0 | $less(0,alength) | $less(X0,sK0) (5:21) [resolution 278,333]
[SA] new: 1614. $less(0,X1) | 0 = sK0 | 0 = X1 | alength = i(main_end) | ~$less(alength,alength) | $less(X1,i(main_end)) (5:20) [resolution 278,277]
[SA] new: 1615. $less(0,X2) | 0 = sK0 | 0 = X2 | alength = alength | $less(sK0,alength) | $less(X2,alength) (5:18) [resolution 278,254]
[SA] new: 1616. $less(0,X3) | 0 = sK0 | 0 = X3 | alength = sK0 | $less(sK0,alength) | $less(X3,alength) (5:18) [resolution 278,205]
[SA] new: 1617. $less(0,X4) | 0 = sK0 | 0 = X4 | alength = X5 | $less(X5,alength) | $less(X4,X5) (5:18) [resolution 278,99]
[SA] new: 1618. $less(0,X6) | 0 = sK0 | 0 = X6 | ~$less(X7,X6) | $less(X7,alength) (5:15) [resolution 278,14]
[SA] new: 1619. $less(0,i(main_end)) | 0 = sK0 | i(main_end) = 0 (5:11) [resolution 278,82]
[SA] new: 1620. $less(0,X0) | 0 = sK0 | 0 = X0 | alength = sK0 | $less(0,alength) | $less(X0,sK0) (5:18) [duplicate literal removal 1613]
[SA] new: 1621. $less(sK0,alength) | 0 = sK0 | $less(0,alength) (5:9) [duplicate literal removal 1611]
[SA] new: 1622. $less(sK0,alength) | 0 = sK0 | sK0 = X11 | $less(X11,alength) | $less(0,X11) (5:15) [duplicate literal removal 1610]
[SA] new: 1623. $less(sK0,alength) | 0 = sK0 | alength = X10 | $less(X10,alength) | $less(0,X10) (5:15) [duplicate literal removal 1609]
[SA] new: 1624. $less(0,alength) | 0 = sK0 | 0 = 0 | $less(0,sK0) (5:12) [duplicate literal removal 1608]
[SA] new: 1625. $less(0,alength) | 0 = sK0 | 0 = 0 (5:9) [duplicate literal removal 1607]
[SA] new: 1626. $less(0,alength) | 0 = sK0 | 0 = 0 | alength = 0 (5:12) [duplicate literal removal 1606]
[SA] new: 1627. $less(0,alength) | 0 = sK0 | 0 = 0 | sK0 = X9 | $less(X9,sK0) | $less(0,X9) (5:18) [duplicate literal removal 1605]
[SA] new: 1628. $less(i(main_end),alength) | 0 = sK0 | i(main_end) = 0 | 'Sub'(sK2(0),nl8) (5:15) [evaluation 1603]
[SA] new: 1629. $less(i(main_end),alength) | 0 = sK0 | i(main_end) = 0 | i(l8(sK2(0))) = 0 (5:17) [evaluation 1602]
[SA] new: 1630. $less(X0,alength) | 0 = sK0 | 0 = X0 | sK0 = X0 | $less(0,X0) | $less(0,sK0) (5:18) [duplicate literal removal 1595]
[SA] passive: 1600. $less(X6,alength) | 0 = sK0 | 0 = X6 | ~$less(X7,0) | $less(X7,X6) (5:15) [resolution 278,14]
[SA] new: 1631. 0 = sK0 | i(main_end) = 0 | $less(X8,0) | i(main_end) = X8 | i(l8(sK2(X8))) = X8 | $less(0,X8) (5:23) [subsumption resolution 1601,82]
[SA] passive: 1631. 0 = sK0 | i(main_end) = 0 | $less(X8,0) | i(main_end) = X8 | i(l8(sK2(X8))) = X8 | $less(0,X8) (5:23) [subsumption resolution 1601,82]
[SA] new: 1632. 0 = sK0 | i(main_end) = 0 | 'Sub'(sK2(0),nl8) (5:11) [subsumption resolution 1628,82]
[SA] passive: 1632. 0 = sK0 | i(main_end) = 0 | 'Sub'(sK2(0),nl8) (5:11) [subsumption resolution 1628,82]
[SA] passive: 1618. $less(0,X6) | 0 = sK0 | 0 = X6 | ~$less(X7,X6) | $less(X7,alength) (5:15) [resolution 278,14]
[SA] passive: 1619. $less(0,i(main_end)) | 0 = sK0 | i(main_end) = 0 (5:11) [resolution 278,82]
[SA] active: 440. a(l8(X0),X1) = a(l8(s(X0)),X1) | i(l8(X0)) = X1 (3:15:1) [forward demodulation 439,169]
[SA] new: 1633. a(l8(X0),X1) = a(l8(p(X0)),X1) | i(l8(p(X0))) = X1 | zero = X0 (4:19) [superposition 440,44]
[SA] new: 1634. a(l8(X2),X3) = a(l11(X2),X3) | i(l8(X2)) = X3 (4:14) [superposition 440,57]
[SA] new: 1636. v = a(l8(X5),i(l8(p(s(X5))))) | i(l8(X5)) = i(l8(p(s(X5)))) | zero = s(X5) (4:23) [superposition 440,189]
[SA] new: 1637. v = a(l8(nl8),i(main_end)) | i(main_end) = i(l8(nl8)) (5:13) [superposition 440,354]
[SA] new: 1638. v = a(l8(s(zero)),1) | i(l8(s(zero))) = 1 (5:13) [superposition 440,355]
[SA] new: 1639. v = a(l8(s(nl8)),$sum(i(main_end),1)) | i(l8(s(nl8))) = $sum(i(main_end),1) (5:19) [superposition 440,359]
[SA] new: 1640. v = a(l8(zero),0) | i(l8(zero)) = 0 (5:11) [superposition 440,353]
[SA] new: 1641. v = a(l8(s(s(zero))),2) | i(l8(s(s(zero)))) = 2 (5:15) [superposition 440,356]
[SA] new: 1642. v = a(l8(s(s(s(zero)))),3) | i(l8(s(s(s(zero))))) = 3 (5:17) [superposition 440,357]
[SA] new: 1643. v = a(l8(s(s(s(s(zero))))),4) | i(l8(s(s(s(s(zero)))))) = 4 (6:19) [superposition 440,358]
[SA] new: 1644. a(l8(X0),X1) = a(l11(X0),X1) | i(l8(X0)) = X1 (4:14) [superposition 57,440]
[SA] new: 1646. v = a(l8(X3),i(l8(p(s(X3))))) | zero = s(X3) | i(l8(X3)) = i(l8(p(s(X3)))) (4:23) [superposition 189,440]
[SA] new: 1647. v = a(l8(nl8),i(main_end)) | i(main_end) = i(l8(nl8)) (5:13) [superposition 354,440]
[SA] new: 1648. v = a(l8(s(zero)),1) | i(l8(s(zero))) = 1 (5:13) [superposition 355,440]
[SA] new: 1649. v = a(l8(s(nl8)),$sum(i(main_end),1)) | i(l8(s(nl8))) = $sum(i(main_end),1) (5:19) [superposition 359,440]
[SA] new: 1650. v = a(l8(zero),0) | i(l8(zero)) = 0 (5:11) [superposition 353,440]
[SA] new: 1651. v = a(l8(s(s(zero))),2) | i(l8(s(s(zero)))) = 2 (5:15) [superposition 356,440]
[SA] new: 1652. v = a(l8(s(s(s(zero)))),3) | i(l8(s(s(s(zero))))) = 3 (5:17) [superposition 357,440]
[SA] new: 1653. v = a(l8(s(s(s(s(zero))))),4) | i(l8(s(s(s(s(zero)))))) = 4 (6:19) [superposition 358,440]
[SA] new: 1654. v = a(l8(X3),i(l8(p(s(X3))))) | i(l8(X3)) = i(l8(p(s(X3)))) (4:19) [term algebras distinctness 1646]
[SA] new: 1655. v = a(l8(X5),i(l8(p(s(X5))))) | i(l8(X5)) = i(l8(p(s(X5)))) (4:19) [term algebras distinctness 1636]
[SA] new: 1656. v = a(l8(X5),i(l8(X5))) | i(l8(X5)) = i(l8(p(s(X5)))) (4:17) [forward demodulation 1655,169]
[SA] new: 1657. v = a(l8(X3),i(l8(X3))) | i(l8(X3)) = i(l8(p(s(X3)))) (4:17) [forward demodulation 1654,169]
[SA] active: 1619. $less(0,i(main_end)) | 0 = sK0 | i(main_end) = 0 (5:11:1) [resolution 278,82]
[SA] new: 1658. 0 = sK0 | i(main_end) = 0 | $less(X0,0) | i(main_end) = X0 | i(l8(sK2(X0))) = X0 | $less(0,X0) (6:23) [resolution 1619,215]
[SA] new: 1659. 0 = sK0 | i(main_end) = 0 | $less(0,0) | i(l8(sK2(0))) = 0 (6:16) [resolution 1619,84]
[SA] new: 1660. 0 = sK0 | i(main_end) = 0 | $less(0,0) | 'Sub'(sK2(0),nl8) (6:14) [resolution 1619,83]
[SA] new: 1661. 0 = sK0 | i(main_end) = 0 | i(main_end) = sK0 | 0 = sK0 | $less(0,i(main_end)) | $less(0,sK0) (6:21) [resolution 1619,333]
[SA] new: 1662. 0 = sK0 | i(main_end) = 0 | alength = i(main_end) | ~$less(i(main_end),alength) | $less(0,i(main_end)) (6:19) [resolution 1619,277]
[SA] new: 1663. 0 = sK0 | i(main_end) = 0 | alength = i(main_end) | $less(sK0,i(main_end)) | $less(0,alength) (6:18) [resolution 1619,254]
[SA] new: 1664. 0 = sK0 | i(main_end) = 0 | i(main_end) = sK0 | $less(sK0,i(main_end)) | $less(0,alength) (6:18) [resolution 1619,205]
[SA] new: 1665. 0 = sK0 | i(main_end) = 0 | i(main_end) = X1 | $less(X1,i(main_end)) | $less(0,X1) (6:18) [resolution 1619,99]
[SA] new: 1666. 0 = sK0 | i(main_end) = 0 | ~$less(X2,0) | $less(X2,i(main_end)) (6:14) [resolution 1619,14]
[SA] new: 1667. 0 = sK0 | i(main_end) = 0 | i(main_end) = sK0 | $less(0,i(main_end)) | $less(0,sK0) (6:18) [duplicate literal removal 1661]
[SA] new: 1668. 0 = sK0 | i(main_end) = 0 | 'Sub'(sK2(0),nl8) (6:11) [evaluation 1660]
[SA] new: 1669. 0 = sK0 | i(main_end) = 0 | i(l8(sK2(0))) = 0 (6:13) [evaluation 1659]
[SA] passive: 1665. 0 = sK0 | i(main_end) = 0 | i(main_end) = X1 | $less(X1,i(main_end)) | $less(0,X1) (6:18) [resolution 1619,99]
[SA] passive: 1666. 0 = sK0 | i(main_end) = 0 | ~$less(X2,0) | $less(X2,i(main_end)) (6:14) [resolution 1619,14]
[SA] active: 273. ~$less(X0,alength) | alength = i(main_end) | $less(X0,0) | i(l8(sK2(X0))) = X0 (3:16:1) [resolution 144,84]
[SA] new: 1670. alength = i(main_end) | $less(X0,0) | i(l8(sK2(X0))) = X0 | $less(sK0,X0) | sK0 = X0 (4:19) [resolution 273,106]
[SA] new: 1671. alength = i(main_end) | $less(X1,0) | i(l8(sK2(X1))) = X1 | $less(0,X1) | alength = X1 | 0 = sK0 (4:22) [resolution 273,178]
[SA] new: 1672. alength = i(main_end) | $less(X2,0) | i(l8(sK2(X2))) = X2 | $less(sK0,X2) | alength = X2 (4:19) [resolution 273,179]
[SA] new: 1673. alength = i(main_end) | $less(X3,0) | i(l8(sK2(X3))) = X3 | $less(sK0,X4) | $less(X4,X3) | X3 = X4 | sK0 = X3 (4:25) [resolution 273,198]
[SA] new: 1674. alength = i(main_end) | $less(X5,0) | i(l8(sK2(X5))) = X5 | $less(sK0,X6) | $less(X6,X5) | alength = X5 | sK0 = X6 (4:25) [resolution 273,204]
[SA] new: 1675. alength = i(main_end) | $less(X7,0) | i(l8(sK2(X7))) = X7 | $less(sK0,X8) | $less(X8,X7) | X7 = X8 | alength = X7 (4:25) [resolution 273,246]
[SA] new: 1676. alength = i(main_end) | $less(X9,0) | i(l8(sK2(X9))) = X9 | $less(sK0,X10) | $less(X10,X9) | alength = X9 | alength = X10 (4:25) [resolution 273,253]
[SA] new: 1677. alength = i(main_end) | $less(X11,0) | i(l8(sK2(X11))) = X11 | $less(0,X11) | alength = X11 | alength = 0 (5:22) [resolution 273,261]
[SA] new: 1678. alength = i(main_end) | $less(X12,0) | i(l8(sK2(X12))) = X12 | $less(0,X12) | 0 = sK0 | 0 = X12 (5:22) [resolution 273,278]
[SA] new: 1679. alength = i(main_end) | $less(X13,0) | i(l8(sK2(X13))) = X13 | $less(0,X13) | 0 = sK0 | sK0 = X13 (4:22) [resolution 273,330]
[SA] new: 1680. alength = i(main_end) | $less(X14,0) | i(l8(sK2(X14))) = X14 | $less(sK0,X15) | $less(X15,X14) | X14 = X15 | sK0 = X15 (4:25) [resolution 273,388]
[SA] new: 1681. alength = i(main_end) | $less(X16,0) | i(l8(sK2(X16))) = X16 | $less(sK0,X17) | $less(X17,X16) | sK0 = X17 | sK0 = X16 (4:25) [resolution 273,408]
[SA] new: 1682. alength = i(main_end) | $less(X18,0) | i(l8(sK2(X18))) = X18 | $less(alength,X18) | alength = X18 (4:19) [resolution 273,15]
[SA] new: 1683. alength = i(main_end) | $less(X19,0) | i(l8(sK2(X19))) = X19 | $less(alength,X19) | alength = X19 (4:19) [resolution 273,15]
[SA] new: 1684. alength = i(main_end) | $less(X20,0) | i(l8(sK2(X20))) = X20 | $less(X21,X20) | $less(alength,X21) | alength = X20 | alength = X21 (4:25) [resolution 273,174]
[SA] new: 1685. alength = i(main_end) | $less(X22,0) | i(l8(sK2(X22))) = X22 | $less(alength,X23) | $less(X23,X22) | X22 = X23 | alength = X22 (4:25) [resolution 273,174]
[SA] new: 1686. alength = i(main_end) | $less(X24,0) | i(l8(sK2(X24))) = X24 | $less(X25,X24) | $less(alength,X25) | alength = X25 | X24 = X25 (4:25) [resolution 273,174]
[SA] new: 1687. alength = i(main_end) | $less(X26,0) | i(l8(sK2(X26))) = X26 | $less(sK0,X26) | $less(alength,alength) | alength = X26 | alength = sK0 (4:25) [resolution 273,198]
[SA] new: 1688. alength = i(main_end) | $less(X27,0) | i(l8(sK2(X27))) = X27 | $less(sK0,X27) | $less(alength,alength) | alength = alength | sK0 = X27 (4:25) [resolution 273,204]
[SA] new: 1689. alength = i(main_end) | $less(X28,0) | i(l8(sK2(X28))) = X28 | $less(i(main_end),X28) | $less(alength,0) | alength = X28 | alength = i(main_end) | alength = i(l8(sK2(alength))) (4:33) [resolution 273,214]
[SA] new: 1690. alength = i(main_end) | $less(X29,0) | i(l8(sK2(X29))) = X29 | $less(sK0,X29) | $less(alength,alength) | alength = X29 | alength = alength (4:25) [resolution 273,246]
[SA] new: 1691. alength = i(main_end) | $less(X30,0) | i(l8(sK2(X30))) = X30 | $less(sK0,X30) | $less(alength,alength) | alength = alength | alength = X30 (4:25) [resolution 273,253]
[SA] new: 1692. alength = i(main_end) | $less(X31,0) | i(l8(sK2(X31))) = X31 | $less(sK0,X31) | $less(alength,alength) | alength = X31 | sK0 = X31 (4:25) [resolution 273,388]
[SA] new: 1693. alength = i(main_end) | $less(X32,0) | i(l8(sK2(X32))) = X32 | $less(sK0,X32) | $less(alength,alength) | sK0 = X32 | alength = sK0 (4:25) [resolution 273,408]
[SA] new: 1694. alength = i(main_end) | $less(i(l8(p(nl8))),0) | i(l8(p(nl8))) = i(l8(sK2(i(l8(p(nl8)))))) | zero = nl8 (4:25) [resolution 273,95]
[SA] new: 1695. alength = i(main_end) | $less(i(l8(p(p(nl8)))),0) | i(l8(p(p(nl8)))) = i(l8(sK2(i(l8(p(p(nl8))))))) | zero = p(nl8) | zero = nl8 (4:32) [resolution 273,314]
[SA] new: 1696. alength = i(main_end) | $less(i(l8(sK2(X33))),0) | i(l8(sK2(X33))) = i(l8(sK2(i(l8(sK2(X33)))))) | $less(i(main_end),X33) | i(main_end) = X33 | $less(X33,0) (4:33) [resolution 273,195]
[SA] new: 1697. alength = i(main_end) | $less(i(l8(sK2(sK0))),0) | i(l8(sK2(sK0))) = i(l8(sK2(i(l8(sK2(sK0)))))) | alength = i(main_end) (5:26) [resolution 273,272]
[SA] new: 1698. alength = i(main_end) | $less(i(l8(sK2(sK0))),0) | i(l8(sK2(sK0))) = i(l8(sK2(i(l8(sK2(sK0)))))) | i(main_end) = sK0 (5:26) [resolution 273,228]
[SA] new: 1699. alength = i(main_end) | $less(i(l8(sK2(alength))),0) | i(l8(sK2(alength))) = i(l8(sK2(i(l8(sK2(alength)))))) | alength = i(main_end) | $less(alength,0) (4:29) [resolution 273,270]
[SA] new: 1700. alength = i(main_end) | $less(i(main_end),0) | i(main_end) = i(l8(sK2(i(main_end)))) | i(l8(sK1(X34,sK2(alength)))) = X34 | $less(alength,0) | a(main_end,X34) = a(l11(sK2(alength)),X34) | alength = i(main_end) (4:40) [resolution 273,225]
[SA] new: 1701. alength = i(main_end) | $less(i(main_end),0) | i(main_end) = i(l8(sK2(i(main_end)))) | $less(alength,X35) | $less(X35,0) | alength = X35 | i(main_end) = X35 | i(l8(sK2(X35))) = X35 (4:35) [resolution 273,214]
[SA] new: 1702. alength = i(main_end) | $less(i(main_end),0) | i(main_end) = i(l8(sK2(i(main_end)))) | alength = i(l8(sK2(alength))) | $less(alength,0) | alength = i(main_end) (4:29) [resolution 273,127]
[SA] new: 1703. alength = i(main_end) | $less(i(l8(p(nl8))),0) | i(l8(p(nl8))) = i(l8(sK2(i(l8(p(nl8)))))) | alength = alength | $less(alength,alength) | zero = nl8 (4:31) [resolution 273,193]
[SA] new: 1704. alength = i(main_end) | $less(0,0) | i(l8(sK2(0))) = 0 | alength = 0 (4:16) [resolution 273,244]
[SA] new: 1705. alength = i(main_end) | $less(0,0) | i(l8(sK2(0))) = 0 | 0 = sK0 (4:16) [resolution 273,113]
[SA] new: 1706. alength = i(main_end) | $less(0,0) | i(l8(sK2(0))) = 0 | $less(alength,sK0) | alength = sK0 | 0 = sK0 (4:22) [resolution 273,177]
[SA] new: 1707. alength = i(main_end) | $less(0,0) | i(l8(sK2(0))) = 0 | $less(alength,alength) | alength = alength | 0 = sK0 (4:22) [resolution 273,178]
[SA] new: 1708. alength = i(main_end) | $less(0,0) | i(l8(sK2(0))) = 0 | $less(alength,sK0) | 0 = sK0 | alength = 0 (4:22) [resolution 273,229]
[SA] new: 1709. alength = i(main_end) | $less(0,0) | i(l8(sK2(0))) = 0 | $less(alength,alength) | alength = alength | alength = 0 (5:22) [resolution 273,261]
[SA] new: 1710. alength = i(main_end) | $less(0,0) | i(l8(sK2(0))) = 0 | $less(alength,alength) | 0 = sK0 | alength = 0 (5:22) [resolution 273,278]
[SA] new: 1711. alength = i(main_end) | $less(0,0) | i(l8(sK2(0))) = 0 | $less(alength,alength) | 0 = sK0 | alength = sK0 (4:22) [resolution 273,330]
[SA] new: 1712. alength = i(main_end) | $less(0,0) | i(l8(sK2(0))) = 0 | $less(alength,sK0) | alength = 0 | alength = sK0 (4:22) [resolution 273,380]
[SA] new: 1713. alength = i(main_end) | $less(sK0,0) | i(l8(sK2(sK0))) = sK0 (4:13) [resolution 273,45]
[SA] new: 1714. alength = i(main_end) | $less(sK0,0) | i(l8(sK2(sK0))) = sK0 | $less(alength,alength) | alength = sK0 (4:19) [resolution 273,106]
[SA] new: 1715. alength = i(main_end) | $less(sK0,0) | i(l8(sK2(sK0))) = sK0 | $less(alength,alength) | alength = alength (4:19) [resolution 273,179]
[SA] new: 1716. alength = i(main_end) | $less(sK0,0) | i(l8(sK2(sK0))) = sK0 | $less(alength,X36) | $less(X36,alength) | alength = X36 | sK0 = X36 (4:25) [resolution 273,198]
[SA] new: 1717. alength = i(main_end) | $less(sK0,0) | i(l8(sK2(sK0))) = sK0 | $less(X37,alength) | $less(alength,X37) | alength = X37 | alength = sK0 (4:25) [resolution 273,204]
[SA] new: 1718. alength = i(main_end) | $less(sK0,0) | i(l8(sK2(sK0))) = sK0 | $less(alength,X38) | $less(X38,alength) | alength = X38 | alength = X38 (4:25) [resolution 273,246]
[SA] new: 1719. alength = i(main_end) | $less(sK0,0) | i(l8(sK2(sK0))) = sK0 | $less(X39,alength) | $less(alength,X39) | alength = X39 | alength = alength (4:25) [resolution 273,253]
[SA] new: 1720. alength = i(main_end) | $less(sK0,0) | i(l8(sK2(sK0))) = sK0 | $less(X40,alength) | $less(alength,X40) | alength = X40 | alength = sK0 (4:25) [resolution 273,388]
[SA] new: 1721. alength = i(main_end) | $less(sK0,0) | i(l8(sK2(sK0))) = sK0 | $less(X41,alength) | $less(alength,X41) | alength = sK0 | sK0 = X41 (4:25) [resolution 273,408]
[SA] new: 1722. alength = i(main_end) | $less(sK0,0) | i(l8(sK2(sK0))) = sK0 | $less(alength,X38) | $less(X38,alength) | alength = X38 (4:22) [duplicate literal removal 1718]
[SA] new: 1723. alength = i(main_end) | i(l8(sK2(0))) = 0 | $less(alength,sK0) | alength = 0 | alength = sK0 (4:19) [evaluation 1712]
[SA] new: 1724. alength = i(main_end) | i(l8(sK2(0))) = 0 | $less(alength,alength) | 0 = sK0 | alength = sK0 (4:19) [evaluation 1711]
[SA] new: 1725. alength = i(main_end) | i(l8(sK2(0))) = 0 | $less(alength,alength) | 0 = sK0 | alength = 0 (5:19) [evaluation 1710]
[SA] new: 1726. alength = i(main_end) | i(l8(sK2(0))) = 0 | $less(alength,sK0) | 0 = sK0 | alength = 0 (4:19) [evaluation 1708]
[SA] new: 1727. alength = i(main_end) | i(l8(sK2(0))) = 0 | $less(alength,sK0) | alength = sK0 | 0 = sK0 (4:19) [evaluation 1706]
[SA] new: 1728. alength = i(main_end) | i(l8(sK2(0))) = 0 | 0 = sK0 (4:13) [evaluation 1705]
[SA] new: 1729. alength = i(main_end) | i(l8(sK2(0))) = 0 | alength = 0 (4:13) [evaluation 1704]
[SA] new: 1730. alength = i(main_end) | $less(i(main_end),0) | i(main_end) = i(l8(sK2(i(main_end)))) | alength = i(l8(sK2(alength))) | $less(alength,0) (4:25) [duplicate literal removal 1702]
[SA] new: 1731. alength = i(main_end) | $less(i(main_end),0) | i(main_end) = i(l8(sK2(i(main_end)))) | i(l8(sK1(X34,sK2(alength)))) = X34 | $less(alength,0) | a(main_end,X34) = a(l11(sK2(alength)),X34) (4:36) [duplicate literal removal 1700]
[SA] new: 1732. alength = i(main_end) | $less(i(l8(sK2(alength))),0) | i(l8(sK2(alength))) = i(l8(sK2(i(l8(sK2(alength)))))) | $less(alength,0) (4:25) [duplicate literal removal 1699]
[SA] new: 1733. alength = i(main_end) | $less(i(l8(sK2(sK0))),0) | i(l8(sK2(sK0))) = i(l8(sK2(i(l8(sK2(sK0)))))) (5:22) [duplicate literal removal 1697]
[SA] new: 1734. alength = i(main_end) | $less(X28,0) | i(l8(sK2(X28))) = X28 | $less(i(main_end),X28) | $less(alength,0) | alength = X28 | alength = i(l8(sK2(alength))) (4:29) [duplicate literal removal 1689]
[SA] passive: 1671. alength = i(main_end) | $less(X1,0) | i(l8(sK2(X1))) = X1 | $less(0,X1) | alength = X1 | 0 = sK0 (4:22) [resolution 273,178]
[SA] passive: 1673. alength = i(main_end) | $less(X3,0) | i(l8(sK2(X3))) = X3 | $less(sK0,X4) | $less(X4,X3) | X3 = X4 | sK0 = X3 (4:25) [resolution 273,198]
[SA] passive: 1674. alength = i(main_end) | $less(X5,0) | i(l8(sK2(X5))) = X5 | $less(sK0,X6) | $less(X6,X5) | alength = X5 | sK0 = X6 (4:25) [resolution 273,204]
[SA] passive: 1675. alength = i(main_end) | $less(X7,0) | i(l8(sK2(X7))) = X7 | $less(sK0,X8) | $less(X8,X7) | X7 = X8 | alength = X7 (4:25) [resolution 273,246]
[SA] passive: 1676. alength = i(main_end) | $less(X9,0) | i(l8(sK2(X9))) = X9 | $less(sK0,X10) | $less(X10,X9) | alength = X9 | alength = X10 (4:25) [resolution 273,253]
[SA] passive: 1677. alength = i(main_end) | $less(X11,0) | i(l8(sK2(X11))) = X11 | $less(0,X11) | alength = X11 | alength = 0 (5:22) [resolution 273,261]
[SA] passive: 1679. alength = i(main_end) | $less(X13,0) | i(l8(sK2(X13))) = X13 | $less(0,X13) | 0 = sK0 | sK0 = X13 (4:22) [resolution 273,330]
[SA] passive: 1680. alength = i(main_end) | $less(X14,0) | i(l8(sK2(X14))) = X14 | $less(sK0,X15) | $less(X15,X14) | X14 = X15 | sK0 = X15 (4:25) [resolution 273,388]
[SA] passive: 1681. alength = i(main_end) | $less(X16,0) | i(l8(sK2(X16))) = X16 | $less(sK0,X17) | $less(X17,X16) | sK0 = X17 | sK0 = X16 (4:25) [resolution 273,408]
[SA] passive: 1684. alength = i(main_end) | $less(X20,0) | i(l8(sK2(X20))) = X20 | $less(X21,X20) | $less(alength,X21) | alength = X20 | alength = X21 (4:25) [resolution 273,174]
[SA] passive: 1685. alength = i(main_end) | $less(X22,0) | i(l8(sK2(X22))) = X22 | $less(alength,X23) | $less(X23,X22) | X22 = X23 | alength = X22 (4:25) [resolution 273,174]
[SA] passive: 1686. alength = i(main_end) | $less(X24,0) | i(l8(sK2(X24))) = X24 | $less(X25,X24) | $less(alength,X25) | alength = X25 | X24 = X25 (4:25) [resolution 273,174]
[SA] passive: 1695. alength = i(main_end) | $less(i(l8(p(p(nl8)))),0) | i(l8(p(p(nl8)))) = i(l8(sK2(i(l8(p(p(nl8))))))) | zero = p(nl8) | zero = nl8 (4:32) [resolution 273,314]
[SA] passive: 1696. alength = i(main_end) | $less(i(l8(sK2(X33))),0) | i(l8(sK2(X33))) = i(l8(sK2(i(l8(sK2(X33)))))) | $less(i(main_end),X33) | i(main_end) = X33 | $less(X33,0) (4:33) [resolution 273,195]
[SA] passive: 1733. alength = i(main_end) | $less(i(l8(sK2(sK0))),0) | i(l8(sK2(sK0))) = i(l8(sK2(i(l8(sK2(sK0)))))) (5:22) [duplicate literal removal 1697]
[SA] passive: 1732. alength = i(main_end) | $less(i(l8(sK2(alength))),0) | i(l8(sK2(alength))) = i(l8(sK2(i(l8(sK2(alength)))))) | $less(alength,0) (4:25) [duplicate literal removal 1699]
[SA] active: 1632. 'Sub'(sK2(0),nl8) | i(main_end) = 0 | 0 = sK0 (5:11:1) [subsumption resolution 1628,82]
[SA] new: 1735. i(main_end) = 0 | 0 = sK0 | a(main_end,X0) = a(l11(sK2(0)),X0) | i(l8(sK1(X0,sK2(0)))) = X0 (6:24) [resolution 1632,150]
[SA] new: 1736. i(main_end) = 0 | 0 = sK0 | $less(i(l8(sK2(0))),alength) (6:13) [resolution 1632,56]
[SA] passive: 1736. i(main_end) = 0 | 0 = sK0 | $less(i(l8(sK2(0))),alength) (6:13) [resolution 1632,56]
[SA] active: 365. ~$less(X1,alength) | alength = i(l8(sK2(alength))) | alength = i(main_end) | $less(X1,0) (3:16:1) [resolution 142,14]
[SA] new: 1737. alength = i(l8(sK2(alength))) | alength = i(main_end) | $less(X0,0) | $less(sK0,X0) | sK0 = X0 (4:19) [resolution 365,106]
[SA] new: 1738. alength = i(l8(sK2(alength))) | alength = i(main_end) | $less(X1,0) | $less(0,X1) | alength = X1 | 0 = sK0 (4:22) [resolution 365,178]
[SA] new: 1739. alength = i(l8(sK2(alength))) | alength = i(main_end) | $less(X2,0) | $less(sK0,X2) | alength = X2 (4:19) [resolution 365,179]
[SA] new: 1740. alength = i(l8(sK2(alength))) | alength = i(main_end) | $less(X3,0) | $less(sK0,X4) | $less(X4,X3) | X3 = X4 | sK0 = X3 (4:25) [resolution 365,198]
[SA] new: 1741. alength = i(l8(sK2(alength))) | alength = i(main_end) | $less(X5,0) | $less(sK0,X6) | $less(X6,X5) | alength = X5 | sK0 = X6 (4:25) [resolution 365,204]
[SA] new: 1742. alength = i(l8(sK2(alength))) | alength = i(main_end) | $less(X7,0) | $less(sK0,X8) | $less(X8,X7) | X7 = X8 | alength = X7 (4:25) [resolution 365,246]
[SA] new: 1743. alength = i(l8(sK2(alength))) | alength = i(main_end) | $less(X9,0) | $less(sK0,X10) | $less(X10,X9) | alength = X9 | alength = X10 (4:25) [resolution 365,253]
[SA] new: 1744. alength = i(l8(sK2(alength))) | alength = i(main_end) | $less(X11,0) | $less(0,X11) | alength = X11 | alength = 0 (5:22) [resolution 365,261]
[SA] new: 1745. alength = i(l8(sK2(alength))) | alength = i(main_end) | $less(X12,0) | $less(0,X12) | 0 = sK0 | 0 = X12 (5:22) [resolution 365,278]
[SA] new: 1746. alength = i(l8(sK2(alength))) | alength = i(main_end) | $less(X13,0) | $less(0,X13) | 0 = sK0 | sK0 = X13 (4:22) [resolution 365,330]
[SA] new: 1747. alength = i(l8(sK2(alength))) | alength = i(main_end) | $less(X14,0) | $less(sK0,X15) | $less(X15,X14) | X14 = X15 | sK0 = X15 (4:25) [resolution 365,388]
[SA] new: 1748. alength = i(l8(sK2(alength))) | alength = i(main_end) | $less(X16,0) | $less(sK0,X17) | $less(X17,X16) | sK0 = X17 | sK0 = X16 (4:25) [resolution 365,408]
[SA] new: 1749. alength = i(l8(sK2(alength))) | alength = i(main_end) | $less(X18,0) | $less(alength,X18) | alength = X18 (4:19) [resolution 365,15]
[SA] new: 1750. alength = i(l8(sK2(alength))) | alength = i(main_end) | $less(X19,0) | $less(alength,X19) | alength = X19 (4:19) [resolution 365,15]
[SA] new: 1751. alength = i(l8(sK2(alength))) | alength = i(main_end) | $less(X20,0) | $less(X21,X20) | $less(alength,X21) | alength = X20 | alength = X21 (4:25) [resolution 365,174]
[SA] new: 1752. alength = i(l8(sK2(alength))) | alength = i(main_end) | $less(X22,0) | $less(alength,X23) | $less(X23,X22) | X22 = X23 | alength = X22 (4:25) [resolution 365,174]
[SA] new: 1753. alength = i(l8(sK2(alength))) | alength = i(main_end) | $less(X24,0) | $less(X25,X24) | $less(alength,X25) | alength = X25 | X24 = X25 (4:25) [resolution 365,174]
[SA] new: 1754. alength = i(l8(sK2(alength))) | alength = i(main_end) | $less(X26,0) | $less(sK0,X26) | $less(alength,alength) | alength = X26 | alength = sK0 (4:25) [resolution 365,198]
[SA] new: 1755. alength = i(l8(sK2(alength))) | alength = i(main_end) | $less(X27,0) | $less(sK0,X27) | $less(alength,alength) | alength = alength | sK0 = X27 (4:25) [resolution 365,204]
[SA] new: 1756. alength = i(l8(sK2(alength))) | alength = i(main_end) | $less(X28,0) | $less(i(main_end),X28) | $less(alength,0) | alength = X28 | alength = i(main_end) | alength = i(l8(sK2(alength))) (4:33) [resolution 365,214]
[SA] new: 1757. alength = i(l8(sK2(alength))) | alength = i(main_end) | $less(X29,0) | $less(sK0,X29) | $less(alength,alength) | alength = X29 | alength = alength (4:25) [resolution 365,246]
[SA] new: 1758. alength = i(l8(sK2(alength))) | alength = i(main_end) | $less(X30,0) | $less(sK0,X30) | $less(alength,alength) | alength = alength | alength = X30 (4:25) [resolution 365,253]
[SA] new: 1759. alength = i(l8(sK2(alength))) | alength = i(main_end) | $less(X31,0) | $less(sK0,X31) | $less(alength,alength) | alength = X31 | sK0 = X31 (4:25) [resolution 365,388]
[SA] new: 1760. alength = i(l8(sK2(alength))) | alength = i(main_end) | $less(X32,0) | $less(sK0,X32) | $less(alength,alength) | sK0 = X32 | alength = sK0 (4:25) [resolution 365,408]
[SA] new: 1761. alength = i(l8(sK2(alength))) | alength = i(main_end) | $less(i(l8(p(nl8))),0) | zero = nl8 (4:19) [resolution 365,95]
[SA] new: 1762. alength = i(l8(sK2(alength))) | alength = i(main_end) | $less(i(l8(p(p(nl8)))),0) | zero = p(nl8) | zero = nl8 (4:24) [resolution 365,314]
[SA] new: 1763. alength = i(l8(sK2(alength))) | alength = i(main_end) | $less(i(l8(sK2(X33))),0) | $less(i(main_end),X33) | i(main_end) = X33 | $less(X33,0) (4:27) [resolution 365,195]
[SA] new: 1764. alength = i(l8(sK2(alength))) | alength = i(main_end) | $less(i(l8(sK2(sK0))),0) | alength = i(main_end) (5:20) [resolution 365,272]
[SA] new: 1765. alength = i(l8(sK2(alength))) | alength = i(main_end) | $less(i(l8(sK2(sK0))),0) | i(main_end) = sK0 (5:20) [resolution 365,228]
[SA] new: 1766. alength = i(l8(sK2(alength))) | alength = i(main_end) | $less(i(l8(sK2(alength))),0) | alength = i(main_end) | $less(alength,0) (4:23) [resolution 365,270]
[SA] new: 1767. alength = i(l8(sK2(alength))) | alength = i(main_end) | $less(i(main_end),0) | i(l8(sK1(X34,sK2(alength)))) = X34 | $less(alength,0) | a(main_end,X34) = a(l11(sK2(alength)),X34) | alength = i(main_end) (4:38) [resolution 365,225]
[SA] new: 1768. alength = i(l8(sK2(alength))) | alength = i(main_end) | $less(i(main_end),0) | $less(alength,X35) | $less(X35,0) | alength = X35 | i(main_end) = X35 | i(l8(sK2(X35))) = X35 (4:33) [resolution 365,214]
[SA] new: 1769. alength = i(l8(sK2(alength))) | alength = i(main_end) | $less(i(main_end),0) | alength = i(l8(sK2(alength))) | $less(alength,0) | alength = i(main_end) (4:27) [resolution 365,127]
[SA] new: 1770. alength = i(l8(sK2(alength))) | alength = i(main_end) | $less(i(l8(p(nl8))),0) | alength = alength | $less(alength,alength) | zero = nl8 (4:25) [resolution 365,193]
[SA] new: 1771. alength = i(l8(sK2(alength))) | alength = i(main_end) | $less(0,0) | alength = 0 (4:16) [resolution 365,244]
[SA] new: 1772. alength = i(l8(sK2(alength))) | alength = i(main_end) | $less(0,0) | 0 = sK0 (4:16) [resolution 365,113]
[SA] new: 1773. alength = i(l8(sK2(alength))) | alength = i(main_end) | $less(0,0) | $less(alength,sK0) | alength = sK0 | 0 = sK0 (4:22) [resolution 365,177]
[SA] new: 1774. alength = i(l8(sK2(alength))) | alength = i(main_end) | $less(0,0) | $less(alength,alength) | alength = alength | 0 = sK0 (4:22) [resolution 365,178]
[SA] new: 1775. alength = i(l8(sK2(alength))) | alength = i(main_end) | $less(0,0) | $less(alength,sK0) | 0 = sK0 | alength = 0 (4:22) [resolution 365,229]
[SA] new: 1776. alength = i(l8(sK2(alength))) | alength = i(main_end) | $less(0,0) | $less(alength,alength) | alength = alength | alength = 0 (5:22) [resolution 365,261]
[SA] new: 1777. alength = i(l8(sK2(alength))) | alength = i(main_end) | $less(0,0) | $less(alength,alength) | 0 = sK0 | alength = 0 (5:22) [resolution 365,278]
[SA] new: 1778. alength = i(l8(sK2(alength))) | alength = i(main_end) | $less(0,0) | $less(alength,alength) | 0 = sK0 | alength = sK0 (4:22) [resolution 365,330]
[SA] new: 1779. alength = i(l8(sK2(alength))) | alength = i(main_end) | $less(0,0) | $less(alength,sK0) | alength = 0 | alength = sK0 (4:22) [resolution 365,380]
[SA] new: 1780. alength = i(l8(sK2(alength))) | alength = i(main_end) | $less(sK0,0) (4:13) [resolution 365,45]
[SA] new: 1781. alength = i(l8(sK2(alength))) | alength = i(main_end) | $less(sK0,0) | $less(alength,alength) | alength = sK0 (4:19) [resolution 365,106]
[SA] new: 1782. alength = i(l8(sK2(alength))) | alength = i(main_end) | $less(sK0,0) | $less(alength,alength) | alength = alength (4:19) [resolution 365,179]
[SA] new: 1783. alength = i(l8(sK2(alength))) | alength = i(main_end) | $less(sK0,0) | $less(alength,X36) | $less(X36,alength) | alength = X36 | sK0 = X36 (4:25) [resolution 365,198]
[SA] new: 1784. alength = i(l8(sK2(alength))) | alength = i(main_end) | $less(sK0,0) | $less(X37,alength) | $less(alength,X37) | alength = X37 | alength = sK0 (4:25) [resolution 365,204]
[SA] new: 1785. alength = i(l8(sK2(alength))) | alength = i(main_end) | $less(sK0,0) | $less(alength,X38) | $less(X38,alength) | alength = X38 | alength = X38 (4:25) [resolution 365,246]
[SA] new: 1786. alength = i(l8(sK2(alength))) | alength = i(main_end) | $less(sK0,0) | $less(X39,alength) | $less(alength,X39) | alength = X39 | alength = alength (4:25) [resolution 365,253]
[SA] new: 1787. alength = i(l8(sK2(alength))) | alength = i(main_end) | $less(sK0,0) | $less(X40,alength) | $less(alength,X40) | alength = X40 | alength = sK0 (4:25) [resolution 365,388]
[SA] new: 1788. alength = i(l8(sK2(alength))) | alength = i(main_end) | $less(sK0,0) | $less(X41,alength) | $less(alength,X41) | alength = sK0 | sK0 = X41 (4:25) [resolution 365,408]
[SA] new: 1789. alength = i(l8(sK2(alength))) | alength = i(main_end) | $less(sK0,0) | $less(alength,X38) | $less(X38,alength) | alength = X38 (4:22) [duplicate literal removal 1785]
[SA] new: 1790. alength = i(l8(sK2(alength))) | alength = i(main_end) | $less(alength,sK0) | alength = 0 | alength = sK0 (4:19) [evaluation 1779]
[SA] new: 1791. alength = i(l8(sK2(alength))) | alength = i(main_end) | $less(alength,alength) | 0 = sK0 | alength = sK0 (4:19) [evaluation 1778]
[SA] new: 1792. alength = i(l8(sK2(alength))) | alength = i(main_end) | $less(alength,alength) | 0 = sK0 | alength = 0 (5:19) [evaluation 1777]
[SA] new: 1793. alength = i(l8(sK2(alength))) | alength = i(main_end) | $less(alength,sK0) | 0 = sK0 | alength = 0 (4:19) [evaluation 1775]
[SA] new: 1794. alength = i(l8(sK2(alength))) | alength = i(main_end) | $less(alength,sK0) | alength = sK0 | 0 = sK0 (4:19) [evaluation 1773]
[SA] new: 1795. alength = i(l8(sK2(alength))) | alength = i(main_end) | 0 = sK0 (4:13) [evaluation 1772]
[SA] new: 1796. alength = i(l8(sK2(alength))) | alength = i(main_end) | alength = 0 (4:13) [evaluation 1771]
[SA] new: 1797. alength = i(l8(sK2(alength))) | alength = i(main_end) | $less(i(main_end),0) | $less(alength,0) (4:17) [duplicate literal removal 1769]
[SA] new: 1798. alength = i(l8(sK2(alength))) | alength = i(main_end) | $less(i(main_end),0) | i(l8(sK1(X34,sK2(alength)))) = X34 | $less(alength,0) | a(main_end,X34) = a(l11(sK2(alength)),X34) (4:34) [duplicate literal removal 1767]
[SA] new: 1799. alength = i(l8(sK2(alength))) | alength = i(main_end) | $less(i(l8(sK2(alength))),0) | $less(alength,0) (4:19) [duplicate literal removal 1766]
[SA] new: 1800. alength = i(l8(sK2(alength))) | alength = i(main_end) | $less(i(l8(sK2(sK0))),0) (5:16) [duplicate literal removal 1764]
[SA] new: 1801. alength = i(l8(sK2(alength))) | alength = i(main_end) | $less(X28,0) | $less(i(main_end),X28) | $less(alength,0) | alength = X28 (4:23) [duplicate literal removal 1756]
[SA] passive: 1737. alength = i(l8(sK2(alength))) | alength = i(main_end) | $less(X0,0) | $less(sK0,X0) | sK0 = X0 (4:19) [resolution 365,106]
[SA] passive: 1739. alength = i(l8(sK2(alength))) | alength = i(main_end) | $less(X2,0) | $less(sK0,X2) | alength = X2 (4:19) [resolution 365,179]
[SA] passive: 1740. alength = i(l8(sK2(alength))) | alength = i(main_end) | $less(X3,0) | $less(sK0,X4) | $less(X4,X3) | X3 = X4 | sK0 = X3 (4:25) [resolution 365,198]
[SA] passive: 1741. alength = i(l8(sK2(alength))) | alength = i(main_end) | $less(X5,0) | $less(sK0,X6) | $less(X6,X5) | alength = X5 | sK0 = X6 (4:25) [resolution 365,204]
[SA] passive: 1742. alength = i(l8(sK2(alength))) | alength = i(main_end) | $less(X7,0) | $less(sK0,X8) | $less(X8,X7) | X7 = X8 | alength = X7 (4:25) [resolution 365,246]
[SA] passive: 1743. alength = i(l8(sK2(alength))) | alength = i(main_end) | $less(X9,0) | $less(sK0,X10) | $less(X10,X9) | alength = X9 | alength = X10 (4:25) [resolution 365,253]
[SA] passive: 1747. alength = i(l8(sK2(alength))) | alength = i(main_end) | $less(X14,0) | $less(sK0,X15) | $less(X15,X14) | X14 = X15 | sK0 = X15 (4:25) [resolution 365,388]
[SA] passive: 1748. alength = i(l8(sK2(alength))) | alength = i(main_end) | $less(X16,0) | $less(sK0,X17) | $less(X17,X16) | sK0 = X17 | sK0 = X16 (4:25) [resolution 365,408]
[SA] passive: 1749. alength = i(l8(sK2(alength))) | alength = i(main_end) | $less(X18,0) | $less(alength,X18) | alength = X18 (4:19) [resolution 365,15]
[SA] passive: 1751. alength = i(l8(sK2(alength))) | alength = i(main_end) | $less(X20,0) | $less(X21,X20) | $less(alength,X21) | alength = X20 | alength = X21 (4:25) [resolution 365,174]
[SA] passive: 1752. alength = i(l8(sK2(alength))) | alength = i(main_end) | $less(X22,0) | $less(alength,X23) | $less(X23,X22) | X22 = X23 | alength = X22 (4:25) [resolution 365,174]
[SA] passive: 1753. alength = i(l8(sK2(alength))) | alength = i(main_end) | $less(X24,0) | $less(X25,X24) | $less(alength,X25) | alength = X25 | X24 = X25 (4:25) [resolution 365,174]
[SA] passive: 1761. alength = i(l8(sK2(alength))) | alength = i(main_end) | $less(i(l8(p(nl8))),0) | zero = nl8 (4:19) [resolution 365,95]
[SA] passive: 1762. alength = i(l8(sK2(alength))) | alength = i(main_end) | $less(i(l8(p(p(nl8)))),0) | zero = p(nl8) | zero = nl8 (4:24) [resolution 365,314]
[SA] passive: 1763. alength = i(l8(sK2(alength))) | alength = i(main_end) | $less(i(l8(sK2(X33))),0) | $less(i(main_end),X33) | i(main_end) = X33 | $less(X33,0) (4:27) [resolution 365,195]
[SA] passive: 1800. alength = i(l8(sK2(alength))) | alength = i(main_end) | $less(i(l8(sK2(sK0))),0) (5:16) [duplicate literal removal 1764]
[SA] new: 1802. alength = i(l8(sK2(alength))) | alength = i(main_end) (4:10) [subsumption resolution 1780,46]
[SA] passive: 1802. alength = i(l8(sK2(alength))) | alength = i(main_end) (4:10) [subsumption resolution 1780,46]
[SA] active: 1802. alength = i(l8(sK2(alength))) | alength = i(main_end) (4:10:1) [subsumption resolution 1780,46]
[SA] new: 1803. $less(alength,alength) | alength = i(main_end) | $less(alength,0) | alength = i(main_end) (5:14) [superposition 270,1802]
[SA] new: 1804. $less(alength,alength) | $less(i(main_end),alength) | alength = i(main_end) | $less(alength,0) | alength = i(main_end) (5:18) [superposition 195,1802]
[SA] new: 1805. alength != i(l8(X0)) | sK2(alength) = X0 | alength = i(main_end) (5:13) [superposition 48,1802]
[SA] new: 1806. alength != i(l8(X1)) | sK2(alength) = X1 | alength = i(main_end) (5:13) [superposition 48,1802]
[SA] new: 1807. v = a(l11(sK2(alength)),alength) | alength = i(main_end) (5:11) [superposition 70,1802]
[SA] new: 1808. i(l8(s(sK2(alength)))) = $sum(alength,1) | alength = i(main_end) (5:13) [superposition 74,1802]
[SA] new: 1809. alength != 0 | zero = sK2(alength) | alength = i(main_end) (5:11) [superposition 108,1802]
[SA] new: 1810. alength != i(main_end) | nl8 = sK2(alength) | alength = i(main_end) (5:12) [superposition 109,1802]
[SA] new: 1811. alength != 1 | s(zero) = sK2(alength) | alength = i(main_end) (5:12) [superposition 121,1802]
[SA] new: 1812. alength != 2 | s(s(zero)) = sK2(alength) | alength = i(main_end) (5:13) [superposition 130,1802]
[SA] new: 1813. alength != $sum(i(main_end),1) | s(nl8) = sK2(alength) | alength = i(main_end) (5:15) [superposition 158,1802]
[SA] new: 1814. alength != 3 | s(s(s(zero))) = sK2(alength) | alength = i(main_end) (5:14) [superposition 183,1802]
[SA] new: 1815. v = a(l8(s(sK2(alength))),alength) | alength = i(main_end) (5:12) [superposition 349,1802]
[SA] new: 1816. $less(alength,alength) | $less(i(main_end),alength) | alength = i(main_end) | $less(alength,0) (5:14) [duplicate literal removal 1804]
[SA] new: 1817. $less(alength,alength) | alength = i(main_end) | $less(alength,0) (5:10) [duplicate literal removal 1803]
[SA] new: 1818. alength = i(main_end) | $less(alength,0) (5:7) [subsumption resolution 1817,13]
[SA] passive: 1818. alength = i(main_end) | $less(alength,0) (5:7) [subsumption resolution 1817,13]
[SA] passive: 1805. alength != i(l8(X0)) | sK2(alength) = X0 | alength = i(main_end) (5:13) [superposition 48,1802]
[SA] passive: 1807. v = a(l11(sK2(alength)),alength) | alength = i(main_end) (5:11) [superposition 70,1802]
[SA] passive: 1808. i(l8(s(sK2(alength)))) = $sum(alength,1) | alength = i(main_end) (5:13) [superposition 74,1802]
[SA] passive: 1809. alength != 0 | zero = sK2(alength) | alength = i(main_end) (5:11) [superposition 108,1802]
[SA] passive: 1811. alength != 1 | s(zero) = sK2(alength) | alength = i(main_end) (5:12) [superposition 121,1802]
[SA] passive: 1812. alength != 2 | s(s(zero)) = sK2(alength) | alength = i(main_end) (5:13) [superposition 130,1802]
[SA] passive: 1813. alength != $sum(i(main_end),1) | s(nl8) = sK2(alength) | alength = i(main_end) (5:15) [superposition 158,1802]
[SA] passive: 1814. alength != 3 | s(s(s(zero))) = sK2(alength) | alength = i(main_end) (5:14) [superposition 183,1802]
[SA] passive: 1815. v = a(l8(s(sK2(alength))),alength) | alength = i(main_end) (5:12) [superposition 349,1802]
[SA] active: 234. $less(i(main_end),sK0) | 0 = sK0 | i(l8(sK2(0))) = 0 | i(main_end) = 0 (3:17:1) [evaluation 231]
[SA] new: 1819. 0 = sK0 | i(l8(sK2(0))) = 0 | i(main_end) = 0 | alength = i(main_end) (6:17) [resolution 234,812]
[SA] new: 1820. 0 = sK0 | i(l8(sK2(0))) = 0 | i(main_end) = 0 | i(main_end) = sK0 (6:17) [resolution 234,722]
[SA] new: 1821. 0 = sK0 | i(l8(sK2(0))) = 0 | i(main_end) = 0 | alength = X0 | $less(X0,alength) | $less(i(main_end),X0) (4:23) [resolution 234,247]
[SA] new: 1822. 0 = sK0 | i(l8(sK2(0))) = 0 | i(main_end) = 0 | sK0 = X1 | $less(X1,alength) | $less(i(main_end),X1) (4:23) [resolution 234,199]
[SA] new: 1823. 0 = sK0 | i(l8(sK2(0))) = 0 | i(main_end) = 0 | $less(i(main_end),alength) (4:17) [resolution 234,96]
[SA] new: 1824. 0 = sK0 | i(l8(sK2(0))) = 0 | i(main_end) = 0 | sK0 = sK0 | 0 = sK0 | $less(0,sK0) | $less(i(main_end),sK0) (4:26) [resolution 234,333]
[SA] new: 1825. 0 = sK0 | i(l8(sK2(0))) = 0 | i(main_end) = 0 | alength = i(main_end) | ~$less(sK0,alength) | $less(i(main_end),i(main_end)) (4:25) [resolution 234,277]
[SA] new: 1826. 0 = sK0 | i(l8(sK2(0))) = 0 | i(main_end) = 0 | alength = sK0 | $less(sK0,sK0) | $less(i(main_end),alength) (4:23) [resolution 234,254]
[SA] new: 1827. 0 = sK0 | i(l8(sK2(0))) = 0 | i(main_end) = 0 | sK0 = sK0 | $less(sK0,sK0) | $less(i(main_end),alength) (4:23) [resolution 234,205]
[SA] new: 1828. 0 = sK0 | i(l8(sK2(0))) = 0 | i(main_end) = 0 | sK0 = X2 | $less(X2,sK0) | $less(i(main_end),X2) (4:23) [resolution 234,99]
[SA] new: 1829. 0 = sK0 | i(l8(sK2(0))) = 0 | i(main_end) = 0 | ~$less(X3,i(main_end)) | $less(X3,sK0) (4:20) [resolution 234,14]
[SA] new: 1830. 0 = sK0 | i(l8(sK2(0))) = 0 | i(main_end) = 0 | sK0 = sK0 | $less(0,sK0) | $less(i(main_end),sK0) (4:23) [duplicate literal removal 1824]
[SA] active: 1818. $less(alength,0) | alength = i(main_end) (5:7:1) [subsumption resolution 1817,13]
[SA] new: 1831. alength = i(main_end) | sK0 = X0 | 0 = sK0 | $less(X0,sK0) | $less(alength,X0) (6:16) [resolution 1818,324]
[SA] new: 1832. alength = i(main_end) | alength = 0 | $less(alength,alength) (6:10) [resolution 1818,262]
[SA] new: 1833. alength = i(main_end) | 0 = sK0 | $less(alength,alength) (6:10) [resolution 1818,116]
[SA] new: 1834. alength = i(main_end) | 0 = sK0 | $less(alength,sK0) (6:10) [resolution 1818,114]
[SA] new: 1835. alength = i(main_end) | 0 = sK0 | 0 = sK0 | $less(0,0) | $less(alength,sK0) (6:16) [resolution 1818,333]
[SA] new: 1836. alength = i(main_end) | alength = i(main_end) | ~$less(0,alength) | $less(alength,i(main_end)) (6:15) [resolution 1818,277]
[SA] new: 1837. alength = i(main_end) | alength = 0 | $less(sK0,0) | $less(alength,alength) (6:13) [resolution 1818,254]
[SA] new: 1838. alength = i(main_end) | 0 = sK0 | $less(sK0,0) | $less(alength,alength) (6:13) [resolution 1818,205]
[SA] new: 1839. alength = i(main_end) | 0 = X1 | $less(X1,0) | $less(alength,X1) (6:13) [resolution 1818,99]
[SA] new: 1840. alength = i(main_end) | ~$less(X2,alength) | $less(X2,0) (6:10) [resolution 1818,14]
[SA] new: 1841. alength = i(main_end) | ~$less(0,alength) | $less(alength,i(main_end)) (6:11) [duplicate literal removal 1836]
[SA] new: 1842. alength = i(main_end) | 0 = sK0 | $less(0,0) | $less(alength,sK0) (6:13) [duplicate literal removal 1835]
[SA] new: 1843. alength = i(main_end) | 0 = sK0 | $less(alength,sK0) (6:10) [evaluation 1842]
[SA] passive: 1831. alength = i(main_end) | sK0 = X0 | 0 = sK0 | $less(X0,sK0) | $less(alength,X0) (6:16) [resolution 1818,324]
[SA] new: 1844. alength = i(main_end) | alength = 0 (6:7) [subsumption resolution 1832,13]
[SA] passive: 1844. alength = i(main_end) | alength = 0 (6:7) [subsumption resolution 1832,13]
[SA] new: 1845. alength = i(main_end) | 0 = sK0 (6:7) [subsumption resolution 1833,13]
[SA] passive: 1845. alength = i(main_end) | 0 = sK0 (6:7) [subsumption resolution 1833,13]
[SA] passive: 1839. alength = i(main_end) | 0 = X1 | $less(X1,0) | $less(alength,X1) (6:13) [resolution 1818,99]
[SA] passive: 1840. alength = i(main_end) | ~$less(X2,alength) | $less(X2,0) (6:10) [resolution 1818,14]
[SA] active: 337. $less(i(main_end),sK0) | i(main_end) = sK0 | 0 = sK0 | i(l8(sK2(0))) = 0 (3:17:1) [evaluation 325]
[SA] new: 1846. i(main_end) = sK0 | 0 = sK0 | i(l8(sK2(0))) = 0 | alength = i(main_end) (6:17) [resolution 337,812]
[SA] new: 1847. i(main_end) = sK0 | 0 = sK0 | i(l8(sK2(0))) = 0 | i(main_end) = sK0 (6:17) [resolution 337,722]
[SA] new: 1848. i(main_end) = sK0 | 0 = sK0 | i(l8(sK2(0))) = 0 | alength = X0 | $less(X0,alength) | $less(i(main_end),X0) (4:23) [resolution 337,247]
[SA] new: 1849. i(main_end) = sK0 | 0 = sK0 | i(l8(sK2(0))) = 0 | sK0 = X1 | $less(X1,alength) | $less(i(main_end),X1) (4:23) [resolution 337,199]
[SA] new: 1850. i(main_end) = sK0 | 0 = sK0 | i(l8(sK2(0))) = 0 | $less(i(main_end),alength) (4:17) [resolution 337,96]
[SA] new: 1851. i(main_end) = sK0 | 0 = sK0 | i(l8(sK2(0))) = 0 | sK0 = sK0 | 0 = sK0 | $less(0,sK0) | $less(i(main_end),sK0) (4:26) [resolution 337,333]
[SA] new: 1852. i(main_end) = sK0 | 0 = sK0 | i(l8(sK2(0))) = 0 | alength = i(main_end) | ~$less(sK0,alength) | $less(i(main_end),i(main_end)) (4:25) [resolution 337,277]
[SA] new: 1853. i(main_end) = sK0 | 0 = sK0 | i(l8(sK2(0))) = 0 | alength = sK0 | $less(sK0,sK0) | $less(i(main_end),alength) (4:23) [resolution 337,254]
[SA] new: 1854. i(main_end) = sK0 | 0 = sK0 | i(l8(sK2(0))) = 0 | sK0 = sK0 | $less(sK0,sK0) | $less(i(main_end),alength) (4:23) [resolution 337,205]
[SA] new: 1855. i(main_end) = sK0 | 0 = sK0 | i(l8(sK2(0))) = 0 | sK0 = X2 | $less(X2,sK0) | $less(i(main_end),X2) (4:23) [resolution 337,99]
[SA] new: 1856. i(main_end) = sK0 | 0 = sK0 | i(l8(sK2(0))) = 0 | ~$less(X3,i(main_end)) | $less(X3,sK0) (4:20) [resolution 337,14]
[SA] new: 1857. i(main_end) = sK0 | 0 = sK0 | i(l8(sK2(0))) = 0 | sK0 = sK0 | $less(0,sK0) | $less(i(main_end),sK0) (4:23) [duplicate literal removal 1851]
[SA] new: 1858. i(main_end) = sK0 | 0 = sK0 | i(l8(sK2(0))) = 0 (6:13) [duplicate literal removal 1847]
[SA] active: 1844. alength = i(main_end) | alength = 0 (6:7:1) [subsumption resolution 1832,13]
[SA] new: 1859. ~$less(alength,alength) | alength = 0 (7:6) [superposition 82,1844]
[SA] new: 1860. ~$less(X0,alength) | $less(X0,0) | 'Sub'(sK2(X0),nl8) | alength = 0 (7:13) [superposition 83,1844]
[SA] new: 1861. ~$less(X1,alength) | $less(X1,0) | i(l8(sK2(X1))) = X1 | alength = 0 (7:15) [superposition 84,1844]
[SA] new: 1862. v = a(l11(nl8),alength) | alength = 0 (7:9) [superposition 88,1844]
[SA] new: 1863. $less(alength,alength) | alength = i(main_end) | alength = 0 (7:10) [superposition 100,1844]
[SA] new: 1864. $less(alength,X3) | i(l8(sK2(X3))) = X3 | $less(X3,0) | i(main_end) = X3 | alength = 0 (7:19) [superposition 127,1844]
[SA] new: 1865. $less(X4,alength) | ~$less(X4,alength) | alength = i(main_end) | alength = 0 (7:13) [superposition 144,1844]
[SA] new: 1866. alength != 0 | zero = nl8 | alength = 0 (7:9) [superposition 146,1844]
[SA] new: 1867. alength != 1 | s(zero) = nl8 | alength = 0 (7:10) [superposition 153,1844]
[SA] new: 1868. alength != 2 | s(s(zero)) = nl8 | alength = 0 (7:11) [superposition 154,1844]
[SA] new: 1869. alength != $sum(alength,1) | s(nl8) = nl8 | alength = 0 (7:12) [superposition 155,1844]
[SA] new: 1870. i(l8(s(s(nl8)))) = $sum($sum(alength,1),1) | alength = 0 (7:14) [superposition 157,1844]
[SA] new: 1871. i(l8(X5)) != $sum(alength,1) | s(nl8) = X5 | alength = 0 (7:14) [superposition 158,1844]
[SA] new: 1872. v = a(l11(s(nl8)),$sum(alength,1)) | alength = 0 (7:12) [superposition 160,1844]
[SA] new: 1873. alength != -1 | alength = 0 (7:6) [superposition 163,1844]
[SA] new: 1874. $less(X6,alength) | i(main_end) = X6 | $less(alength,X6) | alength = i(main_end) | alength = 0 (7:17) [superposition 176,1844]
[SA] new: 1875. alength != 3 | s(s(s(zero))) = nl8 | alength = 0 (7:12) [superposition 180,1844]
[SA] new: 1876. $less(sK0,alength) | i(main_end) = sK0 | alength = 0 (7:10) [superposition 206,1844]
[SA] new: 1877. $less(alength,X7) | $less(X7,X8) | $less(X8,0) | X7 = X8 | i(main_end) = X8 | i(l8(sK2(X8))) = X8 | alength = 0 (7:25) [superposition 214,1844]
[SA] new: 1878. ~$less(X9,alength) | $less(X10,0) | i(main_end) = X10 | i(l8(sK2(X10))) = X10 | $less(X9,X10) | alength = 0 (7:22) [superposition 215,1844]
[SA] new: 1879. $less(X11,alength) | ~$less(X11,sK0) | i(main_end) = sK0 | alength = 0 (7:13) [superposition 224,1844]
[SA] new: 1880. $less(alength,X12) | i(l8(sK1(X13,sK2(X12)))) = X13 | $less(X12,0) | a(main_end,X13) = a(l11(sK2(X12)),X13) | i(main_end) = X12 | alength = 0 (7:30) [superposition 225,1844]
[SA] new: 1881. $less(alength,sK0) | 0 = sK0 | i(l8(sK2(0))) = 0 | i(main_end) = 0 | alength = 0 (7:19) [superposition 234,1844]
[SA] new: 1882. $less(sK0,alength) | alength = i(main_end) | alength = 0 (7:10) [superposition 255,1844]
[SA] new: 1883. $less(X14,alength) | ~$less(X14,sK0) | alength = i(main_end) | alength = 0 (7:13) [superposition 268,1844]
[SA] new: 1884. alength != 4 | s(s(s(s(zero)))) = nl8 | alength = 0 (7:13) [superposition 306,1844]
[SA] new: 1885. $less(alength,sK0) | i(main_end) = sK0 | 0 = sK0 | i(l8(sK2(0))) = 0 | alength = 0 (7:19) [superposition 337,1844]
[SA] new: 1886. v = a(l8(s(nl8)),alength) | alength = 0 (7:10) [superposition 354,1844]
[SA] new: 1887. v = a(l8(s(s(nl8))),$sum(alength,1)) | alength = 0 (7:13) [superposition 359,1844]
[SA] new: 1888. $less(X15,alength) | $less(alength,X15) | alength = X15 | i(main_end) = X15 | alength = 0 (7:16) [superposition 389,1844]
[SA] new: 1889. $less(X16,alength) | $less(alength,X16) | alength = i(main_end) | alength = X16 | alength = 0 (7:16) [superposition 409,1844]
[SA] new: 1890. $less(0,alength) | alength = i(main_end) | 0 = sK0 | alength = 0 (7:13) [superposition 583,1844]
[SA] new: 1891. ~$less(alength,sK0) | i(main_end) = sK0 | alength = 0 (7:10) [superposition 722,1844]
[SA] new: 1892. $less(0,alength) | i(main_end) = sK0 | 0 = sK0 | alength = 0 (7:13) [superposition 761,1844]
[SA] new: 1893. ~$less(alength,sK0) | alength = i(main_end) | alength = 0 (7:10) [superposition 812,1844]
[SA] new: 1894. $less(0,alength) | alength = i(main_end) | alength = 0 | alength = 0 (7:13) [superposition 1429,1844]
[SA] new: 1895. $less(0,alength) | 0 = sK0 | i(main_end) = 0 | alength = 0 (7:13) [superposition 1619,1844]
[SA] new: 1896. $less(0,alength) | alength = i(main_end) | alength = 0 (7:10) [duplicate literal removal 1894]
[SA] passive: 1860. ~$less(X0,alength) | $less(X0,0) | 'Sub'(sK2(X0),nl8) | alength = 0 (7:13) [superposition 83,1844]
[SA] passive: 1861. ~$less(X1,alength) | $less(X1,0) | i(l8(sK2(X1))) = X1 | alength = 0 (7:15) [superposition 84,1844]
[SA] passive: 1862. v = a(l11(nl8),alength) | alength = 0 (7:9) [superposition 88,1844]
[SA] passive: 1864. $less(alength,X3) | i(l8(sK2(X3))) = X3 | $less(X3,0) | i(main_end) = X3 | alength = 0 (7:19) [superposition 127,1844]
[SA] passive: 1867. alength != 1 | s(zero) = nl8 | alength = 0 (7:10) [superposition 153,1844]
[SA] passive: 1868. alength != 2 | s(s(zero)) = nl8 | alength = 0 (7:11) [superposition 154,1844]
[SA] passive: 1869. alength != $sum(alength,1) | s(nl8) = nl8 | alength = 0 (7:12) [superposition 155,1844]
[SA] passive: 1870. i(l8(s(s(nl8)))) = $sum($sum(alength,1),1) | alength = 0 (7:14) [superposition 157,1844]
[SA] passive: 1871. i(l8(X5)) != $sum(alength,1) | s(nl8) = X5 | alength = 0 (7:14) [superposition 158,1844]
[SA] passive: 1872. v = a(l11(s(nl8)),$sum(alength,1)) | alength = 0 (7:12) [superposition 160,1844]
[SA] passive: 1873. alength != -1 | alength = 0 (7:6) [superposition 163,1844]
[SA] passive: 1875. alength != 3 | s(s(s(zero))) = nl8 | alength = 0 (7:12) [superposition 180,1844]
[SA] passive: 1877. $less(alength,X7) | $less(X7,X8) | $less(X8,0) | X7 = X8 | i(main_end) = X8 | i(l8(sK2(X8))) = X8 | alength = 0 (7:25) [superposition 214,1844]
[SA] passive: 1878. ~$less(X9,alength) | $less(X10,0) | i(main_end) = X10 | i(l8(sK2(X10))) = X10 | $less(X9,X10) | alength = 0 (7:22) [superposition 215,1844]
[SA] passive: 1880. $less(alength,X12) | i(l8(sK1(X13,sK2(X12)))) = X13 | $less(X12,0) | a(main_end,X13) = a(l11(sK2(X12)),X13) | i(main_end) = X12 | alength = 0 (7:30) [superposition 225,1844]
[SA] passive: 1884. alength != 4 | s(s(s(s(zero)))) = nl8 | alength = 0 (7:13) [superposition 306,1844]
[SA] passive: 1886. v = a(l8(s(nl8)),alength) | alength = 0 (7:10) [superposition 354,1844]
[SA] passive: 1887. v = a(l8(s(s(nl8))),$sum(alength,1)) | alength = 0 (7:13) [superposition 359,1844]
[SA] passive: 1891. ~$less(alength,sK0) | i(main_end) = sK0 | alength = 0 (7:10) [superposition 722,1844]
[SA] active: 276. ~$less(X2,alength) | alength = i(main_end) | i(main_end) = X3 | $less(X3,i(main_end)) | $less(X2,X3) (3:18:1) [resolution 144,99]
[SA] new: 1897. alength = i(main_end) | i(main_end) = X0 | $less(X0,i(main_end)) | $less(X1,X0) | $less(sK0,X1) | sK0 = X1 (4:21) [resolution 276,106]
[SA] new: 1898. alength = i(main_end) | i(main_end) = X2 | $less(X2,i(main_end)) | $less(X3,X2) | $less(0,X3) | alength = X3 | 0 = sK0 (4:24) [resolution 276,178]
[SA] new: 1899. alength = i(main_end) | i(main_end) = X4 | $less(X4,i(main_end)) | $less(X5,X4) | $less(sK0,X5) | alength = X5 (4:21) [resolution 276,179]
[SA] new: 1900. alength = i(main_end) | i(main_end) = X6 | $less(X6,i(main_end)) | $less(X7,X6) | $less(sK0,X8) | $less(X8,X7) | X7 = X8 | sK0 = X7 (4:27) [resolution 276,198]
[SA] new: 1901. alength = i(main_end) | i(main_end) = X9 | $less(X9,i(main_end)) | $less(X10,X9) | $less(sK0,X11) | $less(X11,X10) | alength = X10 | sK0 = X11 (4:27) [resolution 276,204]
[SA] new: 1902. alength = i(main_end) | i(main_end) = X12 | $less(X12,i(main_end)) | $less(X13,X12) | $less(sK0,X14) | $less(X14,X13) | X13 = X14 | alength = X13 (4:27) [resolution 276,246]
[SA] new: 1903. alength = i(main_end) | i(main_end) = X15 | $less(X15,i(main_end)) | $less(X16,X15) | $less(sK0,X17) | $less(X17,X16) | alength = X16 | alength = X17 (4:27) [resolution 276,253]
[SA] new: 1904. alength = i(main_end) | i(main_end) = X18 | $less(X18,i(main_end)) | $less(X19,X18) | $less(0,X19) | alength = X19 | alength = 0 (5:24) [resolution 276,261]
[SA] new: 1905. alength = i(main_end) | i(main_end) = X20 | $less(X20,i(main_end)) | $less(X21,X20) | $less(0,X21) | 0 = sK0 | 0 = X21 (5:24) [resolution 276,278]
[SA] new: 1906. alength = i(main_end) | i(main_end) = X22 | $less(X22,i(main_end)) | $less(X23,X22) | $less(0,X23) | 0 = sK0 | sK0 = X23 (4:24) [resolution 276,330]
[SA] new: 1907. alength = i(main_end) | i(main_end) = X24 | $less(X24,i(main_end)) | $less(X25,X24) | $less(sK0,X26) | $less(X26,X25) | X25 = X26 | sK0 = X26 (4:27) [resolution 276,388]
[SA] new: 1908. alength = i(main_end) | i(main_end) = X27 | $less(X27,i(main_end)) | $less(X28,X27) | $less(sK0,X29) | $less(X29,X28) | sK0 = X29 | sK0 = X28 (4:27) [resolution 276,408]
[SA] new: 1909. alength = i(main_end) | i(main_end) = X30 | $less(X30,i(main_end)) | $less(X31,X30) | $less(alength,X31) | alength = X31 (4:21) [resolution 276,15]
[SA] new: 1910. alength = i(main_end) | i(main_end) = X32 | $less(X32,i(main_end)) | $less(X33,X32) | $less(alength,X33) | alength = X33 (4:21) [resolution 276,15]
[SA] new: 1911. alength = i(main_end) | i(main_end) = X34 | $less(X34,i(main_end)) | $less(X35,X34) | $less(X36,X35) | $less(alength,X36) | alength = X35 | alength = X36 (4:27) [resolution 276,174]
[SA] new: 1912. alength = i(main_end) | i(main_end) = X37 | $less(X37,i(main_end)) | $less(X38,X37) | $less(alength,X39) | $less(X39,X38) | X38 = X39 | alength = X38 (4:27) [resolution 276,174]
[SA] new: 1913. alength = i(main_end) | i(main_end) = X40 | $less(X40,i(main_end)) | $less(X41,X40) | $less(X42,X41) | $less(alength,X42) | alength = X42 | X41 = X42 (4:27) [resolution 276,174]
[SA] new: 1914. alength = i(main_end) | i(main_end) = X43 | $less(X43,i(main_end)) | $less(X44,X43) | $less(sK0,X44) | $less(alength,alength) | alength = X44 | alength = sK0 (4:27) [resolution 276,198]
[SA] new: 1915. alength = i(main_end) | i(main_end) = X45 | $less(X45,i(main_end)) | $less(X46,X45) | $less(sK0,X46) | $less(alength,alength) | alength = alength | sK0 = X46 (4:27) [resolution 276,204]
[SA] new: 1916. alength = i(main_end) | i(main_end) = X47 | $less(X47,i(main_end)) | $less(X48,X47) | $less(i(main_end),X48) | $less(alength,0) | alength = X48 | alength = i(main_end) | alength = i(l8(sK2(alength))) (4:35) [resolution 276,214]
[SA] new: 1917. alength = i(main_end) | i(main_end) = X49 | $less(X49,i(main_end)) | $less(X50,X49) | $less(sK0,X50) | $less(alength,alength) | alength = X50 | alength = alength (4:27) [resolution 276,246]
[SA] new: 1918. alength = i(main_end) | i(main_end) = X51 | $less(X51,i(main_end)) | $less(X52,X51) | $less(sK0,X52) | $less(alength,alength) | alength = alength | alength = X52 (4:27) [resolution 276,253]
[SA] new: 1919. alength = i(main_end) | i(main_end) = X53 | $less(X53,i(main_end)) | $less(X54,X53) | $less(sK0,X54) | $less(alength,alength) | alength = X54 | sK0 = X54 (4:27) [resolution 276,388]
[SA] new: 1920. alength = i(main_end) | i(main_end) = X55 | $less(X55,i(main_end)) | $less(X56,X55) | $less(sK0,X56) | $less(alength,alength) | sK0 = X56 | alength = sK0 (4:27) [resolution 276,408]
[SA] new: 1921. alength = i(main_end) | i(main_end) = X57 | $less(X57,i(main_end)) | $less(i(l8(p(nl8))),X57) | zero = nl8 (4:21) [resolution 276,95]
[SA] new: 1922. alength = i(main_end) | i(main_end) = X58 | $less(X58,i(main_end)) | $less(i(l8(p(p(nl8)))),X58) | zero = p(nl8) | zero = nl8 (4:26) [resolution 276,314]
[SA] new: 1923. alength = i(main_end) | i(main_end) = X59 | $less(X59,i(main_end)) | $less(i(l8(sK2(X60))),X59) | $less(i(main_end),X60) | i(main_end) = X60 | $less(X60,0) (4:29) [resolution 276,195]
[SA] new: 1924. alength = i(main_end) | i(main_end) = X61 | $less(X61,i(main_end)) | $less(i(l8(sK2(sK0))),X61) | alength = i(main_end) (5:22) [resolution 276,272]
[SA] new: 1925. alength = i(main_end) | i(main_end) = X62 | $less(X62,i(main_end)) | $less(i(l8(sK2(sK0))),X62) | i(main_end) = sK0 (5:22) [resolution 276,228]
[SA] new: 1926. alength = i(main_end) | i(main_end) = X63 | $less(X63,i(main_end)) | $less(i(l8(sK2(alength))),X63) | alength = i(main_end) | $less(alength,0) (4:25) [resolution 276,270]
[SA] new: 1927. alength = i(main_end) | i(main_end) = X64 | $less(X64,i(main_end)) | $less(i(main_end),X64) | i(l8(sK1(X65,sK2(alength)))) = X65 | $less(alength,0) | a(main_end,X65) = a(l11(sK2(alength)),X65) | alength = i(main_end) (4:40) [resolution 276,225]
[SA] new: 1928. alength = i(main_end) | i(main_end) = X66 | $less(X66,i(main_end)) | $less(i(main_end),X66) | $less(alength,X67) | $less(X67,0) | alength = X67 | i(main_end) = X67 | i(l8(sK2(X67))) = X67 (4:35) [resolution 276,214]
[SA] new: 1929. alength = i(main_end) | i(main_end) = X68 | $less(X68,i(main_end)) | $less(i(main_end),X68) | alength = i(l8(sK2(alength))) | $less(alength,0) | alength = i(main_end) (4:29) [resolution 276,127]
[SA] new: 1930. alength = i(main_end) | i(main_end) = X69 | $less(X69,i(main_end)) | $less(i(l8(p(nl8))),X69) | alength = alength | $less(alength,alength) | zero = nl8 (4:27) [resolution 276,193]
[SA] new: 1931. alength = i(main_end) | i(main_end) = X70 | $less(X70,i(main_end)) | $less(0,X70) | alength = 0 (4:18) [resolution 276,244]
[SA] new: 1932. alength = i(main_end) | i(main_end) = X71 | $less(X71,i(main_end)) | $less(0,X71) | 0 = sK0 (4:18) [resolution 276,113]
[SA] new: 1933. alength = i(main_end) | i(main_end) = X72 | $less(X72,i(main_end)) | $less(0,X72) | $less(alength,sK0) | alength = sK0 | 0 = sK0 (4:24) [resolution 276,177]
[SA] new: 1934. alength = i(main_end) | i(main_end) = X73 | $less(X73,i(main_end)) | $less(0,X73) | $less(alength,alength) | alength = alength | 0 = sK0 (4:24) [resolution 276,178]
[SA] new: 1935. alength = i(main_end) | i(main_end) = X74 | $less(X74,i(main_end)) | $less(0,X74) | $less(alength,sK0) | 0 = sK0 | alength = 0 (4:24) [resolution 276,229]
[SA] new: 1936. alength = i(main_end) | i(main_end) = X75 | $less(X75,i(main_end)) | $less(0,X75) | $less(alength,alength) | alength = alength | alength = 0 (5:24) [resolution 276,261]
[SA] new: 1937. alength = i(main_end) | i(main_end) = X76 | $less(X76,i(main_end)) | $less(0,X76) | $less(alength,alength) | 0 = sK0 | alength = 0 (5:24) [resolution 276,278]
[SA] new: 1938. alength = i(main_end) | i(main_end) = X77 | $less(X77,i(main_end)) | $less(0,X77) | $less(alength,alength) | 0 = sK0 | alength = sK0 (4:24) [resolution 276,330]
[SA] new: 1939. alength = i(main_end) | i(main_end) = X78 | $less(X78,i(main_end)) | $less(0,X78) | $less(alength,sK0) | alength = 0 | alength = sK0 (4:24) [resolution 276,380]
[SA] new: 1940. alength = i(main_end) | i(main_end) = X79 | $less(X79,i(main_end)) | $less(sK0,X79) (4:15) [resolution 276,45]
[SA] new: 1941. alength = i(main_end) | i(main_end) = X80 | $less(X80,i(main_end)) | $less(sK0,X80) | $less(alength,alength) | alength = sK0 (4:21) [resolution 276,106]
[SA] new: 1942. alength = i(main_end) | i(main_end) = X81 | $less(X81,i(main_end)) | $less(sK0,X81) | $less(alength,alength) | alength = alength (4:21) [resolution 276,179]
[SA] new: 1943. alength = i(main_end) | i(main_end) = X82 | $less(X82,i(main_end)) | $less(sK0,X82) | $less(alength,X83) | $less(X83,alength) | alength = X83 | sK0 = X83 (4:27) [resolution 276,198]
[SA] new: 1944. alength = i(main_end) | i(main_end) = X84 | $less(X84,i(main_end)) | $less(sK0,X84) | $less(X85,alength) | $less(alength,X85) | alength = X85 | alength = sK0 (4:27) [resolution 276,204]
[SA] new: 1945. alength = i(main_end) | i(main_end) = X86 | $less(X86,i(main_end)) | $less(sK0,X86) | $less(alength,X87) | $less(X87,alength) | alength = X87 | alength = X87 (4:27) [resolution 276,246]
[SA] new: 1946. alength = i(main_end) | i(main_end) = X88 | $less(X88,i(main_end)) | $less(sK0,X88) | $less(X89,alength) | $less(alength,X89) | alength = X89 | alength = alength (4:27) [resolution 276,253]
[SA] new: 1947. alength = i(main_end) | i(main_end) = X90 | $less(X90,i(main_end)) | $less(sK0,X90) | $less(X91,alength) | $less(alength,X91) | alength = X91 | alength = sK0 (4:27) [resolution 276,388]
[SA] new: 1948. alength = i(main_end) | i(main_end) = X92 | $less(X92,i(main_end)) | $less(sK0,X92) | $less(X93,alength) | $less(alength,X93) | alength = sK0 | sK0 = X93 (4:27) [resolution 276,408]
[SA] new: 1949. alength = i(main_end) | i(main_end) = X86 | $less(X86,i(main_end)) | $less(sK0,X86) | $less(alength,X87) | $less(X87,alength) | alength = X87 (4:24) [duplicate literal removal 1945]
[SA] new: 1950. alength = i(main_end) | i(main_end) = X68 | $less(X68,i(main_end)) | $less(i(main_end),X68) | alength = i(l8(sK2(alength))) | $less(alength,0) (4:25) [duplicate literal removal 1929]
[SA] new: 1951. alength = i(main_end) | i(main_end) = X64 | $less(X64,i(main_end)) | $less(i(main_end),X64) | i(l8(sK1(X65,sK2(alength)))) = X65 | $less(alength,0) | a(main_end,X65) = a(l11(sK2(alength)),X65) (4:36) [duplicate literal removal 1927]
[SA] new: 1952. alength = i(main_end) | i(main_end) = X63 | $less(X63,i(main_end)) | $less(i(l8(sK2(alength))),X63) | $less(alength,0) (4:21) [duplicate literal removal 1926]
[SA] new: 1953. alength = i(main_end) | i(main_end) = X61 | $less(X61,i(main_end)) | $less(i(l8(sK2(sK0))),X61) (5:18) [duplicate literal removal 1924]
[SA] new: 1954. alength = i(main_end) | i(main_end) = X47 | $less(X47,i(main_end)) | $less(X48,X47) | $less(i(main_end),X48) | $less(alength,0) | alength = X48 | alength = i(l8(sK2(alength))) (4:31) [duplicate literal removal 1916]
[SA] passive: 1897. alength = i(main_end) | i(main_end) = X0 | $less(X0,i(main_end)) | $less(X1,X0) | $less(sK0,X1) | sK0 = X1 (4:21) [resolution 276,106]
[SA] passive: 1899. alength = i(main_end) | i(main_end) = X4 | $less(X4,i(main_end)) | $less(X5,X4) | $less(sK0,X5) | alength = X5 (4:21) [resolution 276,179]
[SA] passive: 1900. alength = i(main_end) | i(main_end) = X6 | $less(X6,i(main_end)) | $less(X7,X6) | $less(sK0,X8) | $less(X8,X7) | X7 = X8 | sK0 = X7 (4:27) [resolution 276,198]
[SA] passive: 1901. alength = i(main_end) | i(main_end) = X9 | $less(X9,i(main_end)) | $less(X10,X9) | $less(sK0,X11) | $less(X11,X10) | alength = X10 | sK0 = X11 (4:27) [resolution 276,204]
[SA] passive: 1902. alength = i(main_end) | i(main_end) = X12 | $less(X12,i(main_end)) | $less(X13,X12) | $less(sK0,X14) | $less(X14,X13) | X13 = X14 | alength = X13 (4:27) [resolution 276,246]
[SA] passive: 1903. alength = i(main_end) | i(main_end) = X15 | $less(X15,i(main_end)) | $less(X16,X15) | $less(sK0,X17) | $less(X17,X16) | alength = X16 | alength = X17 (4:27) [resolution 276,253]
[SA] passive: 1907. alength = i(main_end) | i(main_end) = X24 | $less(X24,i(main_end)) | $less(X25,X24) | $less(sK0,X26) | $less(X26,X25) | X25 = X26 | sK0 = X26 (4:27) [resolution 276,388]
[SA] passive: 1908. alength = i(main_end) | i(main_end) = X27 | $less(X27,i(main_end)) | $less(X28,X27) | $less(sK0,X29) | $less(X29,X28) | sK0 = X29 | sK0 = X28 (4:27) [resolution 276,408]
[SA] passive: 1911. alength = i(main_end) | i(main_end) = X34 | $less(X34,i(main_end)) | $less(X35,X34) | $less(X36,X35) | $less(alength,X36) | alength = X35 | alength = X36 (4:27) [resolution 276,174]
[SA] passive: 1912. alength = i(main_end) | i(main_end) = X37 | $less(X37,i(main_end)) | $less(X38,X37) | $less(alength,X39) | $less(X39,X38) | X38 = X39 | alength = X38 (4:27) [resolution 276,174]
[SA] passive: 1913. alength = i(main_end) | i(main_end) = X40 | $less(X40,i(main_end)) | $less(X41,X40) | $less(X42,X41) | $less(alength,X42) | alength = X42 | X41 = X42 (4:27) [resolution 276,174]
[SA] passive: 1921. alength = i(main_end) | i(main_end) = X57 | $less(X57,i(main_end)) | $less(i(l8(p(nl8))),X57) | zero = nl8 (4:21) [resolution 276,95]
[SA] passive: 1922. alength = i(main_end) | i(main_end) = X58 | $less(X58,i(main_end)) | $less(i(l8(p(p(nl8)))),X58) | zero = p(nl8) | zero = nl8 (4:26) [resolution 276,314]
[SA] passive: 1923. alength = i(main_end) | i(main_end) = X59 | $less(X59,i(main_end)) | $less(i(l8(sK2(X60))),X59) | $less(i(main_end),X60) | i(main_end) = X60 | $less(X60,0) (4:29) [resolution 276,195]
[SA] passive: 1953. alength = i(main_end) | i(main_end) = X61 | $less(X61,i(main_end)) | $less(i(l8(sK2(sK0))),X61) (5:18) [duplicate literal removal 1924]
[SA] active: 1873. alength != -1 | alength = 0 (7:6:1) [superposition 163,1844]
[SA] active: 315. 'Sub'(p(p(p(X1))),X1) | zero = p(p(X1)) | zero = p(X1) | zero = X1 (3:18:1) [resolution 164,92]
[SA] new: 1955. zero = p(p(nl8)) | zero = p(nl8) | zero = nl8 | a(main_end,X0) = a(l11(p(p(p(nl8)))),X0) | i(l8(sK1(X0,p(p(p(nl8)))))) = X0 (4:33) [resolution 315,150]
[SA] new: 1956. zero = p(p(nl8)) | zero = p(nl8) | zero = nl8 | $less(i(l8(p(p(p(nl8))))),alength) (4:20) [resolution 315,56]
[SA] new: 1957. zero = p(p(p(X1))) | zero = p(p(X1)) | zero = p(X1) | 'Sub'(p(p(p(p(X1)))),X1) | zero = X1 (4:25) [resolution 315,92]
[SA] new: 1958. 'Sub'(p(p(X0)),s(X0)) | zero = p(p(s(X0))) | zero = p(s(X0)) | zero = s(X0) (4:21) [superposition 315,169]
[SA] new: 1959. 'Sub'(p(p(X0)),s(X0)) | zero = p(p(s(X0))) | zero = p(s(X0)) (4:17) [term algebras distinctness 1958]
[SA] passive: 1955. zero = p(p(nl8)) | zero = p(nl8) | zero = nl8 | a(main_end,X0) = a(l11(p(p(p(nl8)))),X0) | i(l8(sK1(X0,p(p(p(nl8)))))) = X0 (4:33) [resolution 315,150]
[SA] passive: 1956. zero = p(p(nl8)) | zero = p(nl8) | zero = nl8 | $less(i(l8(p(p(p(nl8))))),alength) (4:20) [resolution 315,56]
[SA] passive: 1957. zero = p(p(p(X1))) | zero = p(p(X1)) | zero = p(X1) | 'Sub'(p(p(p(p(X1)))),X1) | zero = X1 (4:25) [resolution 315,92]
[SA] new: 1960. zero = p(X0) | 'Sub'(p(p(X0)),s(X0)) | zero = p(s(X0)) (4:15) [forward demodulation 1959,169]
[SA] new: 1961. zero = X0 | zero = p(X0) | 'Sub'(p(p(X0)),s(X0)) (4:13) [forward demodulation 1960,169]
[SA] passive: 1961. zero = X0 | zero = p(X0) | 'Sub'(p(p(X0)),s(X0)) (4:13) [forward demodulation 1960,169]
[SA] active: 1845. alength = i(main_end) | 0 = sK0 (6:7:1) [subsumption resolution 1833,13]
[SA] new: 1962. ~$less(alength,alength) | 0 = sK0 (7:6) [superposition 82,1845]
[SA] new: 1963. ~$less(X0,alength) | $less(X0,0) | 'Sub'(sK2(X0),nl8) | 0 = sK0 (7:13) [superposition 83,1845]
[SA] new: 1964. ~$less(X1,alength) | $less(X1,0) | i(l8(sK2(X1))) = X1 | 0 = sK0 (7:15) [superposition 84,1845]
[SA] new: 1965. v = a(l11(nl8),alength) | 0 = sK0 (7:9) [superposition 88,1845]
[SA] new: 1966. $less(alength,alength) | alength = i(main_end) | 0 = sK0 (7:10) [superposition 100,1845]
[SA] new: 1967. $less(alength,X3) | i(l8(sK2(X3))) = X3 | $less(X3,0) | i(main_end) = X3 | 0 = sK0 (7:19) [superposition 127,1845]
[SA] new: 1968. $less(X4,alength) | ~$less(X4,alength) | alength = i(main_end) | 0 = sK0 (7:13) [superposition 144,1845]
[SA] new: 1969. alength != 0 | zero = nl8 | 0 = sK0 (7:9) [superposition 146,1845]
[SA] new: 1970. alength != 1 | s(zero) = nl8 | 0 = sK0 (7:10) [superposition 153,1845]
[SA] new: 1971. alength != 2 | s(s(zero)) = nl8 | 0 = sK0 (7:11) [superposition 154,1845]
[SA] new: 1972. alength != $sum(alength,1) | s(nl8) = nl8 | 0 = sK0 (7:12) [superposition 155,1845]
[SA] new: 1973. i(l8(s(s(nl8)))) = $sum($sum(alength,1),1) | 0 = sK0 (7:14) [superposition 157,1845]
[SA] new: 1974. i(l8(X5)) != $sum(alength,1) | s(nl8) = X5 | 0 = sK0 (7:14) [superposition 158,1845]
[SA] new: 1975. v = a(l11(s(nl8)),$sum(alength,1)) | 0 = sK0 (7:12) [superposition 160,1845]
[SA] new: 1976. alength != -1 | 0 = sK0 (7:6) [superposition 163,1845]
[SA] new: 1977. $less(X6,alength) | i(main_end) = X6 | $less(alength,X6) | alength = i(main_end) | 0 = sK0 (7:17) [superposition 176,1845]
[SA] new: 1978. alength != 3 | s(s(s(zero))) = nl8 | 0 = sK0 (7:12) [superposition 180,1845]
[SA] new: 1979. $less(sK0,alength) | i(main_end) = sK0 | 0 = sK0 (7:10) [superposition 206,1845]
[SA] new: 1980. $less(alength,X7) | $less(X7,X8) | $less(X8,0) | X7 = X8 | i(main_end) = X8 | i(l8(sK2(X8))) = X8 | 0 = sK0 (7:25) [superposition 214,1845]
[SA] new: 1981. ~$less(X9,alength) | $less(X10,0) | i(main_end) = X10 | i(l8(sK2(X10))) = X10 | $less(X9,X10) | 0 = sK0 (7:22) [superposition 215,1845]
[SA] new: 1982. $less(X11,alength) | ~$less(X11,sK0) | i(main_end) = sK0 | 0 = sK0 (7:13) [superposition 224,1845]
[SA] new: 1983. $less(alength,X12) | i(l8(sK1(X13,sK2(X12)))) = X13 | $less(X12,0) | a(main_end,X13) = a(l11(sK2(X12)),X13) | i(main_end) = X12 | 0 = sK0 (7:30) [superposition 225,1845]
[SA] new: 1984. $less(alength,sK0) | 0 = sK0 | i(l8(sK2(0))) = 0 | i(main_end) = 0 | 0 = sK0 (7:19) [superposition 234,1845]
[SA] new: 1985. $less(sK0,alength) | alength = i(main_end) | 0 = sK0 (7:10) [superposition 255,1845]
[SA] new: 1986. $less(X14,alength) | ~$less(X14,sK0) | alength = i(main_end) | 0 = sK0 (7:13) [superposition 268,1845]
[SA] new: 1987. alength != 4 | s(s(s(s(zero)))) = nl8 | 0 = sK0 (7:13) [superposition 306,1845]
[SA] new: 1988. $less(alength,sK0) | i(main_end) = sK0 | 0 = sK0 | i(l8(sK2(0))) = 0 | 0 = sK0 (7:19) [superposition 337,1845]
[SA] new: 1989. v = a(l8(s(nl8)),alength) | 0 = sK0 (7:10) [superposition 354,1845]
[SA] new: 1990. v = a(l8(s(s(nl8))),$sum(alength,1)) | 0 = sK0 (7:13) [superposition 359,1845]
[SA] new: 1991. $less(X15,alength) | $less(alength,X15) | alength = X15 | i(main_end) = X15 | 0 = sK0 (7:16) [superposition 389,1845]
[SA] new: 1992. $less(X16,alength) | $less(alength,X16) | alength = i(main_end) | alength = X16 | 0 = sK0 (7:16) [superposition 409,1845]
[SA] new: 1993. $less(0,alength) | alength = i(main_end) | 0 = sK0 | 0 = sK0 (7:13) [superposition 583,1845]
[SA] new: 1994. ~$less(alength,sK0) | i(main_end) = sK0 | 0 = sK0 (7:10) [superposition 722,1845]
[SA] new: 1995. $less(0,alength) | i(main_end) = sK0 | 0 = sK0 | 0 = sK0 (7:13) [superposition 761,1845]
[SA] new: 1996. ~$less(alength,sK0) | alength = i(main_end) | 0 = sK0 (7:10) [superposition 812,1845]
[SA] new: 1997. $less(0,alength) | alength = i(main_end) | alength = 0 | 0 = sK0 (7:13) [superposition 1429,1845]
[SA] new: 1998. $less(0,alength) | 0 = sK0 | i(main_end) = 0 | 0 = sK0 (7:13) [superposition 1619,1845]
[SA] new: 1999. $less(0,alength) | 0 = sK0 | i(main_end) = 0 (7:10) [duplicate literal removal 1998]
[SA] new: 2000. $less(0,alength) | i(main_end) = sK0 | 0 = sK0 (7:10) [duplicate literal removal 1995]
[SA] new: 2001. $less(0,alength) | alength = i(main_end) | 0 = sK0 (7:10) [duplicate literal removal 1993]
[SA] new: 2002. $less(alength,sK0) | i(main_end) = sK0 | 0 = sK0 | i(l8(sK2(0))) = 0 (7:16) [duplicate literal removal 1988]
[SA] new: 2003. $less(alength,sK0) | 0 = sK0 | i(l8(sK2(0))) = 0 | i(main_end) = 0 (7:16) [duplicate literal removal 1984]
[SA] passive: 1963. ~$less(X0,alength) | $less(X0,0) | 'Sub'(sK2(X0),nl8) | 0 = sK0 (7:13) [superposition 83,1845]
[SA] passive: 1964. ~$less(X1,alength) | $less(X1,0) | i(l8(sK2(X1))) = X1 | 0 = sK0 (7:15) [superposition 84,1845]
[SA] passive: 1965. v = a(l11(nl8),alength) | 0 = sK0 (7:9) [superposition 88,1845]
[SA] passive: 1967. $less(alength,X3) | i(l8(sK2(X3))) = X3 | $less(X3,0) | i(main_end) = X3 | 0 = sK0 (7:19) [superposition 127,1845]
[SA] passive: 1969. alength != 0 | zero = nl8 | 0 = sK0 (7:9) [superposition 146,1845]
[SA] passive: 1970. alength != 1 | s(zero) = nl8 | 0 = sK0 (7:10) [superposition 153,1845]
[SA] passive: 1971. alength != 2 | s(s(zero)) = nl8 | 0 = sK0 (7:11) [superposition 154,1845]
[SA] passive: 1972. alength != $sum(alength,1) | s(nl8) = nl8 | 0 = sK0 (7:12) [superposition 155,1845]
[SA] passive: 1973. i(l8(s(s(nl8)))) = $sum($sum(alength,1),1) | 0 = sK0 (7:14) [superposition 157,1845]
[SA] passive: 1974. i(l8(X5)) != $sum(alength,1) | s(nl8) = X5 | 0 = sK0 (7:14) [superposition 158,1845]
[SA] passive: 1975. v = a(l11(s(nl8)),$sum(alength,1)) | 0 = sK0 (7:12) [superposition 160,1845]
[SA] passive: 1976. alength != -1 | 0 = sK0 (7:6) [superposition 163,1845]
[SA] passive: 1978. alength != 3 | s(s(s(zero))) = nl8 | 0 = sK0 (7:12) [superposition 180,1845]
[SA] passive: 1980. $less(alength,X7) | $less(X7,X8) | $less(X8,0) | X7 = X8 | i(main_end) = X8 | i(l8(sK2(X8))) = X8 | 0 = sK0 (7:25) [superposition 214,1845]
[SA] passive: 1981. ~$less(X9,alength) | $less(X10,0) | i(main_end) = X10 | i(l8(sK2(X10))) = X10 | $less(X9,X10) | 0 = sK0 (7:22) [superposition 215,1845]
[SA] passive: 1983. $less(alength,X12) | i(l8(sK1(X13,sK2(X12)))) = X13 | $less(X12,0) | a(main_end,X13) = a(l11(sK2(X12)),X13) | i(main_end) = X12 | 0 = sK0 (7:30) [superposition 225,1845]
[SA] passive: 1987. alength != 4 | s(s(s(s(zero)))) = nl8 | 0 = sK0 (7:13) [superposition 306,1845]
[SA] passive: 1989. v = a(l8(s(nl8)),alength) | 0 = sK0 (7:10) [superposition 354,1845]
[SA] passive: 1990. v = a(l8(s(s(nl8))),$sum(alength,1)) | 0 = sK0 (7:13) [superposition 359,1845]
[SA] passive: 1994. ~$less(alength,sK0) | i(main_end) = sK0 | 0 = sK0 (7:10) [superposition 722,1845]
[SA] active: 323. $less(0,X1) | $less(X1,X0) | $less(X0,sK0) | X0 = X1 | 0 = sK0 | sK0 = X0 (3:18:3) [resolution 177,99]
[SA] new: 2004. $less(0,X0) | $less(X0,sK0) | 0 = X0 | 0 = sK0 | sK0 = X0 (4:15) [resolution 323,13]
[SA] new: 2005. $less(X1,X2) | $less(X2,sK0) | X1 = X2 | 0 = sK0 | sK0 = X2 | sK0 = X1 | 0 = sK0 | $less(0,X1) | $less(0,sK0) (4:27) [resolution 323,333]
[SA] new: 2006. $less(X3,X4) | $less(X4,sK0) | X3 = X4 | 0 = sK0 | sK0 = X4 | alength = i(main_end) | ~$less(X3,alength) | $less(0,i(main_end)) (4:26) [resolution 323,277]
[SA] new: 2007. $less(X5,X6) | $less(X6,sK0) | X5 = X6 | 0 = sK0 | sK0 = X6 | alength = X5 | $less(sK0,X5) | $less(0,alength) (4:24) [resolution 323,254]
[SA] new: 2008. $less(X7,X8) | $less(X8,sK0) | X7 = X8 | 0 = sK0 | sK0 = X8 | sK0 = X7 | $less(sK0,X7) | $less(0,alength) (4:24) [resolution 323,205]
[SA] new: 2009. $less(X9,X10) | $less(X10,sK0) | X9 = X10 | 0 = sK0 | sK0 = X10 | X9 = X11 | $less(X11,X9) | $less(0,X11) (4:24) [resolution 323,99]
[SA] new: 2010. $less(X12,X13) | $less(X13,sK0) | X12 = X13 | 0 = sK0 | sK0 = X13 | ~$less(X14,0) | $less(X14,X12) (4:21) [resolution 323,14]
[SA] new: 2011. $less(alength,X15) | $less(X15,sK0) | alength = X15 | 0 = sK0 | sK0 = X15 | alength = i(main_end) | i(main_end) = X16 | $less(X16,i(main_end)) | $less(0,X16) (4:30) [resolution 323,276]
[SA] new: 2012. $less(alength,X17) | $less(X17,sK0) | alength = X17 | 0 = sK0 | sK0 = X17 | alength = i(l8(sK2(alength))) | alength = i(main_end) | $less(0,0) (4:28) [resolution 323,365]
[SA] new: 2013. $less(alength,X18) | $less(X18,sK0) | alength = X18 | 0 = sK0 | sK0 = X18 | alength = i(main_end) | $less(0,0) | i(l8(sK2(0))) = 0 (4:28) [resolution 323,273]
[SA] new: 2014. $less(i(main_end),X19) | $less(X19,sK0) | i(main_end) = X19 | 0 = sK0 | sK0 = X19 | $less(X20,0) | i(main_end) = X20 | i(l8(sK2(X20))) = X20 | $less(0,X20) (4:33) [resolution 323,215]
[SA] new: 2015. $less(i(main_end),X21) | $less(X21,sK0) | i(main_end) = X21 | 0 = sK0 | sK0 = X21 | $less(0,0) | i(l8(sK2(0))) = 0 (4:26) [resolution 323,84]
[SA] new: 2016. $less(i(main_end),X22) | $less(X22,sK0) | i(main_end) = X22 | 0 = sK0 | sK0 = X22 | $less(0,0) | 'Sub'(sK2(0),nl8) (4:24) [resolution 323,83]
[SA] new: 2017. $less(i(l8(p(nl8))),X23) | $less(X23,sK0) | i(l8(p(nl8))) = X23 | 0 = sK0 | sK0 = X23 | zero = nl8 | $less(0,alength) (4:27) [resolution 323,194]
[SA] new: 2018. $less(0,X24) | $less(X24,sK0) | 0 = X24 | 0 = sK0 | sK0 = X24 | sK0 = X25 | 0 = sK0 | $less(X25,sK0) | $less(0,X25) (4:27) [resolution 323,324]
[SA] new: 2019. $less(0,X26) | $less(X26,sK0) | 0 = X26 | 0 = sK0 | sK0 = X26 | alength = 0 | $less(0,alength) (5:21) [resolution 323,262]
[SA] new: 2020. $less(0,X27) | $less(X27,sK0) | 0 = X27 | 0 = sK0 | sK0 = X27 | 0 = sK0 | $less(0,alength) (4:21) [resolution 323,116]
[SA] new: 2021. $less(0,X28) | $less(X28,sK0) | 0 = X28 | 0 = sK0 | sK0 = X28 | 0 = sK0 | $less(0,sK0) (4:21) [resolution 323,114]
[SA] new: 2022. $less(sK0,X29) | $less(X29,sK0) | sK0 = X29 | 0 = sK0 | sK0 = X29 | alength = X30 | $less(X30,alength) | $less(0,X30) (4:24) [resolution 323,247]
[SA] new: 2023. $less(sK0,X31) | $less(X31,sK0) | sK0 = X31 | 0 = sK0 | sK0 = X31 | sK0 = X32 | $less(X32,alength) | $less(0,X32) (4:24) [resolution 323,199]
[SA] new: 2024. $less(sK0,X33) | $less(X33,sK0) | sK0 = X33 | 0 = sK0 | sK0 = X33 | $less(0,alength) (4:18) [resolution 323,96]
[SA] new: 2025. $less(0,sK0) | $less(0,sK0) | 0 = sK0 | 0 = sK0 | 0 = sK0 (4:15) [resolution 323,46]
[SA] new: 2026. $less(0,X0) | $less(X0,sK0) | X0 = X0 | 0 = sK0 | sK0 = X0 (4:15) [resolution 323,13]
[SA] new: 2027. $less(0,X1) | $less(X2,sK0) | X1 = X2 | 0 = sK0 | sK0 = X2 | sK0 = X2 | 0 = sK0 | $less(0,X2) | $less(X1,sK0) (4:27) [resolution 323,333]
[SA] new: 2028. $less(0,X3) | $less(X4,sK0) | X3 = X4 | 0 = sK0 | sK0 = X4 | alength = i(main_end) | ~$less(X4,alength) | $less(X3,i(main_end)) (4:26) [resolution 323,277]
[SA] new: 2029. $less(0,X5) | $less(X6,sK0) | X5 = X6 | 0 = sK0 | sK0 = X6 | alength = X6 | $less(sK0,X6) | $less(X5,alength) (4:24) [resolution 323,254]
[SA] new: 2030. $less(0,X7) | $less(X8,sK0) | X7 = X8 | 0 = sK0 | sK0 = X8 | sK0 = X8 | $less(sK0,X8) | $less(X7,alength) (4:24) [resolution 323,205]
[SA] new: 2031. $less(0,X9) | $less(X10,sK0) | X9 = X10 | 0 = sK0 | sK0 = X10 | X10 = X11 | $less(X11,X10) | $less(X9,X11) (4:24) [resolution 323,99]
[SA] new: 2032. $less(0,X12) | $less(X13,sK0) | X12 = X13 | 0 = sK0 | sK0 = X13 | ~$less(X14,X12) | $less(X14,X13) (4:21) [resolution 323,14]
[SA] new: 2033. $less(0,X15) | $less(alength,sK0) | alength = X15 | 0 = sK0 | alength = sK0 | alength = i(main_end) | i(main_end) = X16 | $less(X16,i(main_end)) | $less(X15,X16) (4:30) [resolution 323,276]
[SA] new: 2034. $less(0,X17) | $less(alength,sK0) | alength = X17 | 0 = sK0 | alength = sK0 | alength = i(l8(sK2(alength))) | alength = i(main_end) | $less(X17,0) (4:28) [resolution 323,365]
[SA] new: 2035. $less(0,X18) | $less(alength,sK0) | alength = X18 | 0 = sK0 | alength = sK0 | alength = i(main_end) | $less(X18,0) | i(l8(sK2(X18))) = X18 (4:28) [resolution 323,273]
[SA] new: 2036. $less(0,X19) | $less(i(main_end),sK0) | i(main_end) = X19 | 0 = sK0 | i(main_end) = sK0 | $less(X20,0) | i(main_end) = X20 | i(l8(sK2(X20))) = X20 | $less(X19,X20) (4:34) [resolution 323,215]
[SA] new: 2037. $less(0,X21) | $less(i(main_end),sK0) | i(main_end) = X21 | 0 = sK0 | i(main_end) = sK0 | $less(X21,0) | i(l8(sK2(X21))) = X21 (4:27) [resolution 323,84]
[SA] new: 2038. $less(0,X22) | $less(i(main_end),sK0) | i(main_end) = X22 | 0 = sK0 | i(main_end) = sK0 | $less(X22,0) | 'Sub'(sK2(X22),nl8) (4:25) [resolution 323,83]
[SA] new: 2039. $less(0,X23) | $less(i(l8(p(nl8))),sK0) | i(l8(p(nl8))) = X23 | 0 = sK0 | i(l8(p(nl8))) = sK0 | zero = nl8 | $less(X23,alength) (4:30) [resolution 323,194]
[SA] new: 2040. $less(0,X24) | $less(0,sK0) | 0 = X24 | 0 = sK0 | 0 = sK0 | sK0 = X25 | 0 = sK0 | $less(X25,sK0) | $less(X24,X25) (4:27) [resolution 323,324]
[SA] new: 2041. $less(0,X26) | $less(0,sK0) | 0 = X26 | 0 = sK0 | 0 = sK0 | alength = 0 | $less(X26,alength) (5:21) [resolution 323,262]
[SA] new: 2042. $less(0,X27) | $less(0,sK0) | 0 = X27 | 0 = sK0 | 0 = sK0 | 0 = sK0 | $less(X27,alength) (4:21) [resolution 323,116]
[SA] new: 2043. $less(0,X28) | $less(0,sK0) | 0 = X28 | 0 = sK0 | 0 = sK0 | 0 = sK0 | $less(X28,sK0) (4:21) [resolution 323,114]
[SA] new: 2044. $less(0,X29) | $less(sK0,sK0) | sK0 = X29 | 0 = sK0 | sK0 = sK0 | alength = X30 | $less(X30,alength) | $less(X29,X30) (4:24) [resolution 323,247]
[SA] new: 2045. $less(0,X31) | $less(sK0,sK0) | sK0 = X31 | 0 = sK0 | sK0 = sK0 | sK0 = X32 | $less(X32,alength) | $less(X31,X32) (4:24) [resolution 323,199]
[SA] new: 2046. $less(0,X33) | $less(sK0,sK0) | sK0 = X33 | 0 = sK0 | sK0 = sK0 | $less(X33,alength) (4:18) [resolution 323,96]
[SA] new: 2047. $less(0,i(main_end)) | $less(alength,sK0) | alength = i(main_end) | 0 = sK0 | alength = sK0 (4:17) [resolution 323,82]
[SA] new: 2048. $less(0,i(main_end)) | $less(sK0,sK0) | i(main_end) = sK0 | 0 = sK0 | sK0 = sK0 | alength = i(main_end) (6:21) [resolution 323,812]
[SA] new: 2049. $less(0,i(main_end)) | $less(sK0,sK0) | i(main_end) = sK0 | 0 = sK0 | sK0 = sK0 | i(main_end) = sK0 (6:21) [resolution 323,722]
[SA] new: 2050. $less(0,X0) | $less(X0,X1) | X0 = X1 | 0 = sK0 | sK0 = X1 | alength = X2 | $less(X2,alength) | $less(X1,X2) (4:24) [resolution 323,247]
[SA] new: 2051. $less(0,X3) | $less(X3,X4) | X3 = X4 | 0 = sK0 | sK0 = X4 | sK0 = X5 | $less(X5,alength) | $less(X4,X5) (4:24) [resolution 323,199]
[SA] new: 2052. $less(0,X6) | $less(X6,X7) | X6 = X7 | 0 = sK0 | sK0 = X7 | $less(X7,alength) (4:18) [resolution 323,96]
[SA] new: 2053. $less(0,X8) | $less(X8,sK0) | sK0 = X8 | 0 = sK0 | sK0 = sK0 (4:15) [resolution 323,13]
[SA] new: 2054. $less(0,X9) | $less(X9,X10) | X9 = X10 | 0 = sK0 | sK0 = X10 | sK0 = sK0 | 0 = sK0 | $less(0,sK0) | $less(X10,sK0) (4:27) [resolution 323,333]
[SA] new: 2055. $less(0,X11) | $less(X11,X12) | X11 = X12 | 0 = sK0 | sK0 = X12 | alength = i(main_end) | ~$less(sK0,alength) | $less(X12,i(main_end)) (4:26) [resolution 323,277]
[SA] new: 2056. $less(0,X13) | $less(X13,X14) | X13 = X14 | 0 = sK0 | sK0 = X14 | alength = sK0 | $less(sK0,sK0) | $less(X14,alength) (4:24) [resolution 323,254]
[SA] new: 2057. $less(0,X15) | $less(X15,X16) | X15 = X16 | 0 = sK0 | sK0 = X16 | sK0 = sK0 | $less(sK0,sK0) | $less(X16,alength) (4:24) [resolution 323,205]
[SA] new: 2058. $less(0,X17) | $less(X17,X18) | X17 = X18 | 0 = sK0 | sK0 = X18 | sK0 = X19 | $less(X19,sK0) | $less(X18,X19) (4:24) [resolution 323,99]
[SA] new: 2059. $less(0,X20) | $less(X20,X21) | X20 = X21 | 0 = sK0 | sK0 = X21 | ~$less(X22,X21) | $less(X22,sK0) (4:21) [resolution 323,14]
[SA] new: 2060. $less(0,X23) | $less(X23,i(main_end)) | i(main_end) = X23 | 0 = sK0 | i(main_end) = sK0 | alength = i(main_end) (6:22) [resolution 323,812]
[SA] new: 2061. $less(0,X24) | $less(X24,i(main_end)) | i(main_end) = X24 | 0 = sK0 | i(main_end) = sK0 | i(main_end) = sK0 (6:22) [resolution 323,722]
[SA] new: 2062. $less(0,0) | $less(0,sK0) | 0 = 0 | 0 = sK0 | 0 = sK0 (4:15) [factoring 323]
[SA] new: 2063. $less(0,sK0) | $less(sK0,0) | 0 = sK0 | 0 = sK0 | 0 = sK0 (4:15) [factoring 323]
[SA] new: 2064. $less(0,sK0) | $less(sK0,sK0) | sK0 = sK0 | 0 = sK0 | sK0 = sK0 (4:15) [factoring 323]
[SA] new: 2065. $less(0,sK0) | $less(sK0,sK0) | sK0 = sK0 | 0 = sK0 (4:12) [duplicate literal removal 2064]
[SA] new: 2066. $less(0,sK0) | $less(sK0,0) | 0 = sK0 (4:9) [duplicate literal removal 2063]
[SA] new: 2067. $less(0,0) | $less(0,sK0) | 0 = 0 | 0 = sK0 (4:12) [duplicate literal removal 2062]
[SA] new: 2068. $less(0,X24) | $less(X24,i(main_end)) | i(main_end) = X24 | 0 = sK0 | i(main_end) = sK0 (6:18) [duplicate literal removal 2061]
[SA] new: 2069. $less(0,X9) | $less(X9,X10) | X9 = X10 | 0 = sK0 | sK0 = X10 | sK0 = sK0 | $less(0,sK0) | $less(X10,sK0) (4:24) [duplicate literal removal 2054]
[SA] new: 2070. $less(0,i(main_end)) | $less(sK0,sK0) | i(main_end) = sK0 | 0 = sK0 | sK0 = sK0 (6:17) [duplicate literal removal 2049]
[SA] new: 2071. $less(0,X28) | $less(0,sK0) | 0 = X28 | 0 = sK0 | $less(X28,sK0) (4:15) [duplicate literal removal 2043]
[SA] new: 2072. $less(0,X27) | $less(0,sK0) | 0 = X27 | 0 = sK0 | $less(X27,alength) (4:15) [duplicate literal removal 2042]
[SA] new: 2073. $less(0,X26) | $less(0,sK0) | 0 = X26 | 0 = sK0 | alength = 0 | $less(X26,alength) (5:18) [duplicate literal removal 2041]
[SA] new: 2074. $less(0,X24) | $less(0,sK0) | 0 = X24 | 0 = sK0 | sK0 = X25 | $less(X25,sK0) | $less(X24,X25) (4:21) [duplicate literal removal 2040]
[SA] new: 2075. $less(0,X7) | $less(X8,sK0) | X7 = X8 | 0 = sK0 | sK0 = X8 | $less(sK0,X8) | $less(X7,alength) (4:21) [duplicate literal removal 2030]
[SA] new: 2076. $less(0,X1) | $less(X2,sK0) | X1 = X2 | 0 = sK0 | sK0 = X2 | $less(0,X2) | $less(X1,sK0) (4:21) [duplicate literal removal 2027]
[SA] new: 2077. $less(0,sK0) | 0 = sK0 (4:6) [duplicate literal removal 2025]
[SA] new: 2078. $less(sK0,X33) | $less(X33,sK0) | sK0 = X33 | 0 = sK0 | $less(0,alength) (4:15) [duplicate literal removal 2024]
[SA] new: 2079. $less(sK0,X31) | $less(X31,sK0) | sK0 = X31 | 0 = sK0 | sK0 = X32 | $less(X32,alength) | $less(0,X32) (4:21) [duplicate literal removal 2023]
[SA] new: 2080. $less(sK0,X29) | $less(X29,sK0) | sK0 = X29 | 0 = sK0 | alength = X30 | $less(X30,alength) | $less(0,X30) (4:21) [duplicate literal removal 2022]
[SA] new: 2081. $less(0,X28) | $less(X28,sK0) | 0 = X28 | 0 = sK0 | sK0 = X28 | $less(0,sK0) (4:18) [duplicate literal removal 2021]
[SA] new: 2082. $less(0,X27) | $less(X27,sK0) | 0 = X27 | 0 = sK0 | sK0 = X27 | $less(0,alength) (4:18) [duplicate literal removal 2020]
[SA] new: 2083. $less(0,X24) | $less(X24,sK0) | 0 = X24 | 0 = sK0 | sK0 = X24 | sK0 = X25 | $less(X25,sK0) | $less(0,X25) (4:24) [duplicate literal removal 2018]
[SA] new: 2084. $less(i(main_end),X22) | $less(X22,sK0) | i(main_end) = X22 | 0 = sK0 | sK0 = X22 | 'Sub'(sK2(0),nl8) (4:21) [evaluation 2016]
[SA] new: 2085. $less(i(main_end),X21) | $less(X21,sK0) | i(main_end) = X21 | 0 = sK0 | sK0 = X21 | i(l8(sK2(0))) = 0 (4:23) [evaluation 2015]
[SA] new: 2086. $less(alength,X18) | $less(X18,sK0) | alength = X18 | 0 = sK0 | sK0 = X18 | alength = i(main_end) | i(l8(sK2(0))) = 0 (4:25) [evaluation 2013]
[SA] new: 2087. $less(alength,X17) | $less(X17,sK0) | alength = X17 | 0 = sK0 | sK0 = X17 | alength = i(l8(sK2(alength))) | alength = i(main_end) (4:25) [evaluation 2012]
[SA] new: 2088. $less(X1,X2) | $less(X2,sK0) | X1 = X2 | 0 = sK0 | sK0 = X2 | sK0 = X1 | $less(0,X1) | $less(0,sK0) (4:24) [duplicate literal removal 2005]
[SA] passive: 2009. $less(X9,X10) | $less(X10,sK0) | X9 = X10 | 0 = sK0 | sK0 = X10 | X9 = X11 | $less(X11,X9) | $less(0,X11) (4:24) [resolution 323,99]
[SA] passive: 2010. $less(X12,X13) | $less(X13,sK0) | X12 = X13 | 0 = sK0 | sK0 = X13 | ~$less(X14,0) | $less(X14,X12) (4:21) [resolution 323,14]
[SA] passive: 2014. $less(i(main_end),X19) | $less(X19,sK0) | i(main_end) = X19 | 0 = sK0 | sK0 = X19 | $less(X20,0) | i(main_end) = X20 | i(l8(sK2(X20))) = X20 | $less(0,X20) (4:33) [resolution 323,215]
[SA] passive: 2085. $less(i(main_end),X21) | $less(X21,sK0) | i(main_end) = X21 | 0 = sK0 | sK0 = X21 | i(l8(sK2(0))) = 0 (4:23) [evaluation 2015]
[SA] passive: 2084. $less(i(main_end),X22) | $less(X22,sK0) | i(main_end) = X22 | 0 = sK0 | sK0 = X22 | 'Sub'(sK2(0),nl8) (4:21) [evaluation 2016]
[SA] passive: 2031. $less(0,X9) | $less(X10,sK0) | X9 = X10 | 0 = sK0 | sK0 = X10 | X10 = X11 | $less(X11,X10) | $less(X9,X11) (4:24) [resolution 323,99]
[SA] passive: 2032. $less(0,X12) | $less(X13,sK0) | X12 = X13 | 0 = sK0 | sK0 = X13 | ~$less(X14,X12) | $less(X14,X13) (4:21) [resolution 323,14]
[SA] new: 2089. $less(0,X19) | i(main_end) = X19 | 0 = sK0 | i(main_end) = sK0 | $less(X20,0) | i(main_end) = X20 | i(l8(sK2(X20))) = X20 | $less(X19,X20) (4:30) [subsumption resolution 2036,722]
[SA] passive: 2089. $less(0,X19) | i(main_end) = X19 | 0 = sK0 | i(main_end) = sK0 | $less(X20,0) | i(main_end) = X20 | i(l8(sK2(X20))) = X20 | $less(X19,X20) (4:30) [subsumption resolution 2036,722]
[SA] new: 2090. $less(0,X22) | i(main_end) = X22 | 0 = sK0 | i(main_end) = sK0 | $less(X22,0) | 'Sub'(sK2(X22),nl8) (4:21) [subsumption resolution 2038,722]
[SA] passive: 2090. $less(0,X22) | i(main_end) = X22 | 0 = sK0 | i(main_end) = sK0 | $less(X22,0) | 'Sub'(sK2(X22),nl8) (4:21) [subsumption resolution 2038,722]
[SA] passive: 2039. $less(0,X23) | $less(i(l8(p(nl8))),sK0) | i(l8(p(nl8))) = X23 | 0 = sK0 | i(l8(p(nl8))) = sK0 | zero = nl8 | $less(X23,alength) (4:30) [resolution 323,194]
[SA] passive: 2050. $less(0,X0) | $less(X0,X1) | X0 = X1 | 0 = sK0 | sK0 = X1 | alength = X2 | $less(X2,alength) | $less(X1,X2) (4:24) [resolution 323,247]
[SA] passive: 2051. $less(0,X3) | $less(X3,X4) | X3 = X4 | 0 = sK0 | sK0 = X4 | sK0 = X5 | $less(X5,alength) | $less(X4,X5) (4:24) [resolution 323,199]
[SA] passive: 2058. $less(0,X17) | $less(X17,X18) | X17 = X18 | 0 = sK0 | sK0 = X18 | sK0 = X19 | $less(X19,sK0) | $less(X18,X19) (4:24) [resolution 323,99]
[SA] passive: 2059. $less(0,X20) | $less(X20,X21) | X20 = X21 | 0 = sK0 | sK0 = X21 | ~$less(X22,X21) | $less(X22,sK0) (4:21) [resolution 323,14]
[SA] active: 1976. alength != -1 | 0 = sK0 (7:6:1) [superposition 163,1845]
[SA] active: 332. $less(0,X1) | $less(X2,sK0) | $less(X1,X2) | sK0 = X2 | 0 = sK0 | sK0 = X1 (3:18:3) [resolution 177,99]
[SA] new: 2091. $less(X0,sK0) | $less(0,X0) | sK0 = X0 | 0 = sK0 | 0 = sK0 (4:15) [resolution 332,13]
[SA] new: 2092. $less(X1,sK0) | $less(X2,X1) | sK0 = X1 | 0 = sK0 | sK0 = X2 | sK0 = X2 | 0 = sK0 | $less(0,X2) | $less(0,sK0) (4:27) [resolution 332,333]
[SA] new: 2093. $less(X3,sK0) | $less(X4,X3) | sK0 = X3 | 0 = sK0 | sK0 = X4 | alength = i(main_end) | ~$less(X4,alength) | $less(0,i(main_end)) (4:26) [resolution 332,277]
[SA] new: 2094. $less(X5,sK0) | $less(X6,X5) | sK0 = X5 | 0 = sK0 | sK0 = X6 | alength = X6 | $less(sK0,X6) | $less(0,alength) (4:24) [resolution 332,254]
[SA] new: 2095. $less(X7,sK0) | $less(X8,X7) | sK0 = X7 | 0 = sK0 | sK0 = X8 | sK0 = X8 | $less(sK0,X8) | $less(0,alength) (4:24) [resolution 332,205]
[SA] new: 2096. $less(X9,sK0) | $less(X10,X9) | sK0 = X9 | 0 = sK0 | sK0 = X10 | X10 = X11 | $less(X11,X10) | $less(0,X11) (4:24) [resolution 332,99]
[SA] new: 2097. $less(X12,sK0) | $less(X13,X12) | sK0 = X12 | 0 = sK0 | sK0 = X13 | ~$less(X14,0) | $less(X14,X13) (4:21) [resolution 332,14]
[SA] new: 2098. $less(X15,sK0) | $less(alength,X15) | sK0 = X15 | 0 = sK0 | alength = sK0 | alength = i(main_end) | i(main_end) = X16 | $less(X16,i(main_end)) | $less(0,X16) (4:30) [resolution 332,276]
[SA] new: 2099. $less(X17,sK0) | $less(alength,X17) | sK0 = X17 | 0 = sK0 | alength = sK0 | alength = i(l8(sK2(alength))) | alength = i(main_end) | $less(0,0) (4:28) [resolution 332,365]
[SA] new: 2100. $less(X18,sK0) | $less(alength,X18) | sK0 = X18 | 0 = sK0 | alength = sK0 | alength = i(main_end) | $less(0,0) | i(l8(sK2(0))) = 0 (4:28) [resolution 332,273]
[SA] new: 2101. $less(X19,sK0) | $less(i(main_end),X19) | sK0 = X19 | 0 = sK0 | i(main_end) = sK0 | $less(X20,0) | i(main_end) = X20 | i(l8(sK2(X20))) = X20 | $less(0,X20) (4:33) [resolution 332,215]
[SA] new: 2102. $less(X21,sK0) | $less(i(main_end),X21) | sK0 = X21 | 0 = sK0 | i(main_end) = sK0 | $less(0,0) | i(l8(sK2(0))) = 0 (4:26) [resolution 332,84]
[SA] new: 2103. $less(X22,sK0) | $less(i(main_end),X22) | sK0 = X22 | 0 = sK0 | i(main_end) = sK0 | $less(0,0) | 'Sub'(sK2(0),nl8) (4:24) [resolution 332,83]
[SA] new: 2104. $less(X23,sK0) | $less(i(l8(p(nl8))),X23) | sK0 = X23 | 0 = sK0 | i(l8(p(nl8))) = sK0 | zero = nl8 | $less(0,alength) (4:27) [resolution 332,194]
[SA] new: 2105. $less(X24,sK0) | $less(0,X24) | sK0 = X24 | 0 = sK0 | 0 = sK0 | sK0 = X25 | 0 = sK0 | $less(X25,sK0) | $less(0,X25) (4:27) [resolution 332,324]
[SA] new: 2106. $less(X26,sK0) | $less(0,X26) | sK0 = X26 | 0 = sK0 | 0 = sK0 | alength = 0 | $less(0,alength) (5:21) [resolution 332,262]
[SA] new: 2107. $less(X27,sK0) | $less(0,X27) | sK0 = X27 | 0 = sK0 | 0 = sK0 | 0 = sK0 | $less(0,alength) (4:21) [resolution 332,116]
[SA] new: 2108. $less(X28,sK0) | $less(0,X28) | sK0 = X28 | 0 = sK0 | 0 = sK0 | 0 = sK0 | $less(0,sK0) (4:21) [resolution 332,114]
[SA] new: 2109. $less(X29,sK0) | $less(sK0,X29) | sK0 = X29 | 0 = sK0 | sK0 = sK0 | alength = X30 | $less(X30,alength) | $less(0,X30) (4:24) [resolution 332,247]
[SA] new: 2110. $less(X31,sK0) | $less(sK0,X31) | sK0 = X31 | 0 = sK0 | sK0 = sK0 | sK0 = X32 | $less(X32,alength) | $less(0,X32) (4:24) [resolution 332,199]
[SA] new: 2111. $less(X33,sK0) | $less(sK0,X33) | sK0 = X33 | 0 = sK0 | sK0 = sK0 | $less(0,alength) (4:18) [resolution 332,96]
[SA] new: 2112. $less(0,X0) | $less(X0,X1) | sK0 = X1 | 0 = sK0 | sK0 = X0 | alength = X2 | $less(X2,alength) | $less(X1,X2) (4:24) [resolution 332,247]
[SA] new: 2113. $less(0,X3) | $less(X3,X4) | sK0 = X4 | 0 = sK0 | sK0 = X3 | sK0 = X5 | $less(X5,alength) | $less(X4,X5) (4:24) [resolution 332,199]
[SA] new: 2114. $less(0,X6) | $less(X6,X7) | sK0 = X7 | 0 = sK0 | sK0 = X6 | $less(X7,alength) (4:18) [resolution 332,96]
[SA] new: 2115. $less(0,X8) | $less(X8,sK0) | sK0 = sK0 | 0 = sK0 | sK0 = X8 (4:15) [resolution 332,13]
[SA] new: 2116. $less(0,X9) | $less(X9,X10) | sK0 = X10 | 0 = sK0 | sK0 = X9 | sK0 = sK0 | 0 = sK0 | $less(0,sK0) | $less(X10,sK0) (4:27) [resolution 332,333]
[SA] new: 2117. $less(0,X11) | $less(X11,X12) | sK0 = X12 | 0 = sK0 | sK0 = X11 | alength = i(main_end) | ~$less(sK0,alength) | $less(X12,i(main_end)) (4:26) [resolution 332,277]
[SA] new: 2118. $less(0,X13) | $less(X13,X14) | sK0 = X14 | 0 = sK0 | sK0 = X13 | alength = sK0 | $less(sK0,sK0) | $less(X14,alength) (4:24) [resolution 332,254]
[SA] new: 2119. $less(0,X15) | $less(X15,X16) | sK0 = X16 | 0 = sK0 | sK0 = X15 | sK0 = sK0 | $less(sK0,sK0) | $less(X16,alength) (4:24) [resolution 332,205]
[SA] new: 2120. $less(0,X17) | $less(X17,X18) | sK0 = X18 | 0 = sK0 | sK0 = X17 | sK0 = X19 | $less(X19,sK0) | $less(X18,X19) (4:24) [resolution 332,99]
[SA] new: 2121. $less(0,X20) | $less(X20,X21) | sK0 = X21 | 0 = sK0 | sK0 = X20 | ~$less(X22,X21) | $less(X22,sK0) (4:21) [resolution 332,14]
[SA] new: 2122. $less(0,X23) | $less(X23,i(main_end)) | i(main_end) = sK0 | 0 = sK0 | sK0 = X23 | alength = i(main_end) (6:21) [resolution 332,812]
[SA] new: 2123. $less(0,X24) | $less(X24,i(main_end)) | i(main_end) = sK0 | 0 = sK0 | sK0 = X24 | i(main_end) = sK0 (6:21) [resolution 332,722]
[SA] new: 2124. $less(0,sK0) | $less(0,sK0) | 0 = sK0 | 0 = sK0 | sK0 = sK0 (4:15) [resolution 332,46]
[SA] new: 2125. $less(0,X0) | $less(X0,sK0) | sK0 = X0 | 0 = sK0 | sK0 = X0 (4:15) [resolution 332,13]
[SA] new: 2126. $less(0,X1) | $less(X2,sK0) | sK0 = X2 | 0 = sK0 | sK0 = X1 | sK0 = X2 | 0 = sK0 | $less(0,X2) | $less(X1,sK0) (4:27) [resolution 332,333]
[SA] new: 2127. $less(0,X3) | $less(X4,sK0) | sK0 = X4 | 0 = sK0 | sK0 = X3 | alength = i(main_end) | ~$less(X4,alength) | $less(X3,i(main_end)) (4:26) [resolution 332,277]
[SA] new: 2128. $less(0,X5) | $less(X6,sK0) | sK0 = X6 | 0 = sK0 | sK0 = X5 | alength = X6 | $less(sK0,X6) | $less(X5,alength) (4:24) [resolution 332,254]
[SA] new: 2129. $less(0,X7) | $less(X8,sK0) | sK0 = X8 | 0 = sK0 | sK0 = X7 | sK0 = X8 | $less(sK0,X8) | $less(X7,alength) (4:24) [resolution 332,205]
[SA] new: 2130. $less(0,X9) | $less(X10,sK0) | sK0 = X10 | 0 = sK0 | sK0 = X9 | X10 = X11 | $less(X11,X10) | $less(X9,X11) (4:24) [resolution 332,99]
[SA] new: 2131. $less(0,X12) | $less(X13,sK0) | sK0 = X13 | 0 = sK0 | sK0 = X12 | ~$less(X14,X12) | $less(X14,X13) (4:21) [resolution 332,14]
[SA] new: 2132. $less(0,X15) | $less(alength,sK0) | alength = sK0 | 0 = sK0 | sK0 = X15 | alength = i(main_end) | i(main_end) = X16 | $less(X16,i(main_end)) | $less(X15,X16) (4:30) [resolution 332,276]
[SA] new: 2133. $less(0,X17) | $less(alength,sK0) | alength = sK0 | 0 = sK0 | sK0 = X17 | alength = i(l8(sK2(alength))) | alength = i(main_end) | $less(X17,0) (4:28) [resolution 332,365]
[SA] new: 2134. $less(0,X18) | $less(alength,sK0) | alength = sK0 | 0 = sK0 | sK0 = X18 | alength = i(main_end) | $less(X18,0) | i(l8(sK2(X18))) = X18 (4:28) [resolution 332,273]
[SA] new: 2135. $less(0,X19) | $less(i(main_end),sK0) | i(main_end) = sK0 | 0 = sK0 | sK0 = X19 | $less(X20,0) | i(main_end) = X20 | i(l8(sK2(X20))) = X20 | $less(X19,X20) (4:33) [resolution 332,215]
[SA] new: 2136. $less(0,X21) | $less(i(main_end),sK0) | i(main_end) = sK0 | 0 = sK0 | sK0 = X21 | $less(X21,0) | i(l8(sK2(X21))) = X21 (4:26) [resolution 332,84]
[SA] new: 2137. $less(0,X22) | $less(i(main_end),sK0) | i(main_end) = sK0 | 0 = sK0 | sK0 = X22 | $less(X22,0) | 'Sub'(sK2(X22),nl8) (4:24) [resolution 332,83]
[SA] new: 2138. $less(0,X23) | $less(i(l8(p(nl8))),sK0) | i(l8(p(nl8))) = sK0 | 0 = sK0 | sK0 = X23 | zero = nl8 | $less(X23,alength) (4:27) [resolution 332,194]
[SA] new: 2139. $less(0,X24) | $less(0,sK0) | 0 = sK0 | 0 = sK0 | sK0 = X24 | sK0 = X25 | 0 = sK0 | $less(X25,sK0) | $less(X24,X25) (4:27) [resolution 332,324]
[SA] new: 2140. $less(0,X26) | $less(0,sK0) | 0 = sK0 | 0 = sK0 | sK0 = X26 | alength = 0 | $less(X26,alength) (5:21) [resolution 332,262]
[SA] new: 2141. $less(0,X27) | $less(0,sK0) | 0 = sK0 | 0 = sK0 | sK0 = X27 | 0 = sK0 | $less(X27,alength) (4:21) [resolution 332,116]
[SA] new: 2142. $less(0,X28) | $less(0,sK0) | 0 = sK0 | 0 = sK0 | sK0 = X28 | 0 = sK0 | $less(X28,sK0) (4:21) [resolution 332,114]
[SA] new: 2143. $less(0,X29) | $less(sK0,sK0) | sK0 = sK0 | 0 = sK0 | sK0 = X29 | alength = X30 | $less(X30,alength) | $less(X29,X30) (4:24) [resolution 332,247]
[SA] new: 2144. $less(0,X31) | $less(sK0,sK0) | sK0 = sK0 | 0 = sK0 | sK0 = X31 | sK0 = X32 | $less(X32,alength) | $less(X31,X32) (4:24) [resolution 332,199]
[SA] new: 2145. $less(0,X33) | $less(sK0,sK0) | sK0 = sK0 | 0 = sK0 | sK0 = X33 | $less(X33,alength) (4:18) [resolution 332,96]
[SA] new: 2146. $less(0,i(main_end)) | $less(alength,sK0) | alength = sK0 | 0 = sK0 | i(main_end) = sK0 (4:17) [resolution 332,82]
[SA] new: 2147. $less(0,i(main_end)) | $less(sK0,sK0) | sK0 = sK0 | 0 = sK0 | i(main_end) = sK0 | alength = i(main_end) (6:21) [resolution 332,812]
[SA] new: 2148. $less(0,i(main_end)) | $less(sK0,sK0) | sK0 = sK0 | 0 = sK0 | i(main_end) = sK0 | i(main_end) = sK0 (6:21) [resolution 332,722]
[SA] new: 2149. $less(0,sK0) | $less(sK0,0) | 0 = sK0 | 0 = sK0 | sK0 = sK0 (4:15) [factoring 332]
[SA] new: 2150. $less(0,0) | $less(0,sK0) | 0 = sK0 | 0 = sK0 | 0 = sK0 (4:15) [factoring 332]
[SA] new: 2151. $less(0,sK0) | $less(sK0,sK0) | sK0 = sK0 | 0 = sK0 | sK0 = sK0 (4:15) [factoring 332]
[SA] new: 2152. $less(0,sK0) | $less(sK0,sK0) | sK0 = sK0 | 0 = sK0 (4:12) [duplicate literal removal 2151]
[SA] new: 2153. $less(0,0) | $less(0,sK0) | 0 = sK0 (4:9) [duplicate literal removal 2150]
[SA] new: 2154. $less(0,sK0) | 0 = sK0 (4:6) [evaluation 2153]
[SA] new: 2155. $less(0,sK0) | $less(sK0,0) | 0 = sK0 | sK0 = sK0 (4:12) [duplicate literal removal 2149]
[SA] new: 2156. $less(0,i(main_end)) | $less(sK0,sK0) | sK0 = sK0 | 0 = sK0 | i(main_end) = sK0 (6:17) [duplicate literal removal 2148]
[SA] new: 2157. $less(0,X28) | $less(0,sK0) | 0 = sK0 | sK0 = X28 | $less(X28,sK0) (4:15) [duplicate literal removal 2142]
[SA] new: 2158. $less(0,X27) | $less(0,sK0) | 0 = sK0 | sK0 = X27 | $less(X27,alength) (4:15) [duplicate literal removal 2141]
[SA] new: 2159. $less(0,X26) | $less(0,sK0) | 0 = sK0 | sK0 = X26 | alength = 0 | $less(X26,alength) (5:18) [duplicate literal removal 2140]
[SA] new: 2160. $less(0,X24) | $less(0,sK0) | 0 = sK0 | sK0 = X24 | sK0 = X25 | $less(X25,sK0) | $less(X24,X25) (4:21) [duplicate literal removal 2139]
[SA] new: 2161. $less(0,X7) | $less(X8,sK0) | sK0 = X8 | 0 = sK0 | sK0 = X7 | $less(sK0,X8) | $less(X7,alength) (4:21) [duplicate literal removal 2129]
[SA] new: 2162. $less(0,X1) | $less(X2,sK0) | sK0 = X2 | 0 = sK0 | sK0 = X1 | $less(0,X2) | $less(X1,sK0) (4:21) [duplicate literal removal 2126]
[SA] new: 2163. $less(0,X0) | $less(X0,sK0) | sK0 = X0 | 0 = sK0 (4:12) [duplicate literal removal 2125]
[SA] new: 2164. $less(0,sK0) | 0 = sK0 | sK0 = sK0 (4:9) [duplicate literal removal 2124]
[SA] new: 2165. $less(0,X24) | $less(X24,i(main_end)) | i(main_end) = sK0 | 0 = sK0 | sK0 = X24 (6:17) [duplicate literal removal 2123]
[SA] new: 2166. $less(0,X9) | $less(X9,X10) | sK0 = X10 | 0 = sK0 | sK0 = X9 | sK0 = sK0 | $less(0,sK0) | $less(X10,sK0) (4:24) [duplicate literal removal 2116]
[SA] new: 2167. $less(X28,sK0) | $less(0,X28) | sK0 = X28 | 0 = sK0 | $less(0,sK0) (4:15) [duplicate literal removal 2108]
[SA] new: 2168. $less(X27,sK0) | $less(0,X27) | sK0 = X27 | 0 = sK0 | $less(0,alength) (4:15) [duplicate literal removal 2107]
[SA] new: 2169. $less(X26,sK0) | $less(0,X26) | sK0 = X26 | 0 = sK0 | alength = 0 | $less(0,alength) (5:18) [duplicate literal removal 2106]
[SA] new: 2170. $less(X24,sK0) | $less(0,X24) | sK0 = X24 | 0 = sK0 | sK0 = X25 | $less(X25,sK0) | $less(0,X25) (4:21) [duplicate literal removal 2105]
[SA] new: 2171. $less(X22,sK0) | $less(i(main_end),X22) | sK0 = X22 | 0 = sK0 | i(main_end) = sK0 | 'Sub'(sK2(0),nl8) (4:21) [evaluation 2103]
[SA] new: 2172. $less(X21,sK0) | $less(i(main_end),X21) | sK0 = X21 | 0 = sK0 | i(main_end) = sK0 | i(l8(sK2(0))) = 0 (4:23) [evaluation 2102]
[SA] new: 2173. $less(X18,sK0) | $less(alength,X18) | sK0 = X18 | 0 = sK0 | alength = sK0 | alength = i(main_end) | i(l8(sK2(0))) = 0 (4:25) [evaluation 2100]
[SA] new: 2174. $less(X17,sK0) | $less(alength,X17) | sK0 = X17 | 0 = sK0 | alength = sK0 | alength = i(l8(sK2(alength))) | alength = i(main_end) (4:25) [evaluation 2099]
[SA] new: 2175. $less(X7,sK0) | $less(X8,X7) | sK0 = X7 | 0 = sK0 | sK0 = X8 | $less(sK0,X8) | $less(0,alength) (4:21) [duplicate literal removal 2095]
[SA] new: 2176. $less(X1,sK0) | $less(X2,X1) | sK0 = X1 | 0 = sK0 | sK0 = X2 | $less(0,X2) | $less(0,sK0) (4:21) [duplicate literal removal 2092]
[SA] new: 2177. $less(X0,sK0) | $less(0,X0) | sK0 = X0 | 0 = sK0 (4:12) [duplicate literal removal 2091]
[SA] passive: 2097. $less(X12,sK0) | $less(X13,X12) | sK0 = X12 | 0 = sK0 | sK0 = X13 | ~$less(X14,0) | $less(X14,X13) (4:21) [resolution 332,14]
[SA] passive: 2112. $less(0,X0) | $less(X0,X1) | sK0 = X1 | 0 = sK0 | sK0 = X0 | alength = X2 | $less(X2,alength) | $less(X1,X2) (4:24) [resolution 332,247]
[SA] passive: 2113. $less(0,X3) | $less(X3,X4) | sK0 = X4 | 0 = sK0 | sK0 = X3 | sK0 = X5 | $less(X5,alength) | $less(X4,X5) (4:24) [resolution 332,199]
[SA] passive: 2120. $less(0,X17) | $less(X17,X18) | sK0 = X18 | 0 = sK0 | sK0 = X17 | sK0 = X19 | $less(X19,sK0) | $less(X18,X19) (4:24) [resolution 332,99]
[SA] passive: 2121. $less(0,X20) | $less(X20,X21) | sK0 = X21 | 0 = sK0 | sK0 = X20 | ~$less(X22,X21) | $less(X22,sK0) (4:21) [resolution 332,14]
[SA] passive: 2165. $less(0,X24) | $less(X24,i(main_end)) | i(main_end) = sK0 | 0 = sK0 | sK0 = X24 (6:17) [duplicate literal removal 2123]
[SA] passive: 2130. $less(0,X9) | $less(X10,sK0) | sK0 = X10 | 0 = sK0 | sK0 = X9 | X10 = X11 | $less(X11,X10) | $less(X9,X11) (4:24) [resolution 332,99]
[SA] passive: 2131. $less(0,X12) | $less(X13,sK0) | sK0 = X13 | 0 = sK0 | sK0 = X12 | ~$less(X14,X12) | $less(X14,X13) (4:21) [resolution 332,14]
[SA] new: 2178. $less(0,X19) | i(main_end) = sK0 | 0 = sK0 | sK0 = X19 | $less(X20,0) | i(main_end) = X20 | i(l8(sK2(X20))) = X20 | $less(X19,X20) (4:29) [subsumption resolution 2135,722]
[SA] passive: 2178. $less(0,X19) | i(main_end) = sK0 | 0 = sK0 | sK0 = X19 | $less(X20,0) | i(main_end) = X20 | i(l8(sK2(X20))) = X20 | $less(X19,X20) (4:29) [subsumption resolution 2135,722]
[SA] new: 2179. $less(0,X21) | i(main_end) = sK0 | 0 = sK0 | sK0 = X21 | $less(X21,0) | i(l8(sK2(X21))) = X21 (4:22) [subsumption resolution 2136,722]
[SA] passive: 2179. $less(0,X21) | i(main_end) = sK0 | 0 = sK0 | sK0 = X21 | $less(X21,0) | i(l8(sK2(X21))) = X21 (4:22) [subsumption resolution 2136,722]
[SA] new: 2180. $less(0,X22) | i(main_end) = sK0 | 0 = sK0 | sK0 = X22 | $less(X22,0) | 'Sub'(sK2(X22),nl8) (4:20) [subsumption resolution 2137,722]
[SA] passive: 2180. $less(0,X22) | i(main_end) = sK0 | 0 = sK0 | sK0 = X22 | $less(X22,0) | 'Sub'(sK2(X22),nl8) (4:20) [subsumption resolution 2137,722]
[SA] active: 1862. v = a(l11(nl8),alength) | alength = 0 (7:9:1) [superposition 88,1844]
[SA] new: 2181. v = a(l8(nl8),alength) | alength = 0 | alength = i(l8(nl8)) (8:14) [superposition 1862,72]
[SA] new: 2182. v = a(l8(nl8),alength) | alength = i(l8(nl8)) | alength = 0 (8:14) [superposition 72,1862]
[SA] new: 2183. alength = i(main_end) | v = a(l8(nl8),alength) | alength = 0 (8:13) [forward demodulation 2181,54]
[SA] new: 2184. alength = i(main_end) | v = a(l8(nl8),alength) | alength = 0 (8:13) [forward demodulation 2182,54]
[SA] active: 374. 'Sub'(sK2(X1),nl8) | $less(alength,X1) | alength = i(main_end) | $less(X1,0) | i(main_end) = X1 (3:18:1) [resolution 176,83]
[SA] new: 2185. $less(alength,X0) | alength = i(main_end) | $less(X0,0) | i(main_end) = X0 | a(main_end,X1) = a(l11(sK2(X0)),X1) | i(l8(sK1(X1,sK2(X0)))) = X1 (4:31) [resolution 374,150]
[SA] new: 2186. $less(alength,X2) | alength = i(main_end) | $less(X2,0) | i(main_end) = X2 | $less(i(l8(sK2(X2))),alength) (4:20) [resolution 374,56]
[SA] passive: 2185. $less(alength,X0) | alength = i(main_end) | $less(X0,0) | i(main_end) = X0 | a(main_end,X1) = a(l11(sK2(X0)),X1) | i(l8(sK1(X1,sK2(X0)))) = X1 (4:31) [resolution 374,150]
[SA] passive: 2186. $less(alength,X2) | alength = i(main_end) | $less(X2,0) | i(main_end) = X2 | $less(i(l8(sK2(X2))),alength) (4:20) [resolution 374,56]
[SA] active: 1965. v = a(l11(nl8),alength) | 0 = sK0 (7:9:1) [superposition 88,1845]
[SA] new: 2187. v = a(l8(nl8),alength) | 0 = sK0 | alength = i(l8(nl8)) (8:14) [superposition 1965,72]
[SA] new: 2188. v = a(l8(nl8),alength) | alength = i(l8(nl8)) | 0 = sK0 (8:14) [superposition 72,1965]
[SA] new: 2189. alength = i(main_end) | v = a(l8(nl8),alength) | 0 = sK0 (8:13) [forward demodulation 2187,54]
[SA] new: 2190. alength = i(main_end) | v = a(l8(nl8),alength) | 0 = sK0 (8:13) [forward demodulation 2188,54]
[SA] active: 377. ~$less(X5,X4) | $less(alength,X4) | alength = i(main_end) | i(main_end) = X4 | $less(X5,i(main_end)) (3:18:1) [resolution 176,14]
[SA] new: 2191. $less(alength,X0) | alength = i(main_end) | i(main_end) = X0 | $less(X1,i(main_end)) | $less(X0,X1) | X0 = X1 (4:21) [resolution 377,15]
[SA] new: 2192. $less(alength,X2) | alength = i(main_end) | i(main_end) = X2 | $less(X3,i(main_end)) | $less(X2,X3) | X2 = X3 (4:21) [resolution 377,15]
[SA] new: 2193. $less(alength,X4) | alength = i(main_end) | i(main_end) = X4 | $less(X5,i(main_end)) | $less(X6,X5) | $less(X4,X6) | X4 = X5 | X4 = X6 (4:27) [resolution 377,174]
[SA] new: 2194. $less(alength,X7) | alength = i(main_end) | i(main_end) = X7 | $less(X8,i(main_end)) | $less(X7,X9) | $less(X9,X8) | X8 = X9 | X7 = X8 (4:27) [resolution 377,174]
[SA] new: 2195. $less(alength,X10) | alength = i(main_end) | i(main_end) = X10 | $less(X11,i(main_end)) | $less(X12,X11) | $less(X10,X12) | X10 = X12 | X11 = X12 (4:27) [resolution 377,174]
[SA] new: 2196. $less(alength,X13) | alength = i(main_end) | i(main_end) = X13 | $less(X14,i(main_end)) | $less(sK0,X14) | $less(X13,alength) | X13 = X14 | sK0 = X13 (4:27) [resolution 377,198]
[SA] new: 2197. $less(alength,X15) | alength = i(main_end) | i(main_end) = X15 | $less(X16,i(main_end)) | $less(sK0,X16) | $less(X15,alength) | alength = X15 | sK0 = X16 (4:27) [resolution 377,204]
[SA] new: 2198. $less(alength,X17) | alength = i(main_end) | i(main_end) = X17 | $less(X18,i(main_end)) | $less(i(main_end),X18) | $less(X17,0) | X17 = X18 | i(main_end) = X17 | i(l8(sK2(X17))) = X17 (4:35) [resolution 377,214]
[SA] new: 2199. $less(alength,X19) | alength = i(main_end) | i(main_end) = X19 | $less(X20,i(main_end)) | $less(sK0,X20) | $less(X19,alength) | X19 = X20 | alength = X19 (4:27) [resolution 377,246]
[SA] new: 2200. $less(alength,X21) | alength = i(main_end) | i(main_end) = X21 | $less(X22,i(main_end)) | $less(sK0,X22) | $less(X21,alength) | alength = X21 | alength = X22 (4:27) [resolution 377,253]
[SA] new: 2201. $less(alength,X23) | alength = i(main_end) | i(main_end) = X23 | $less(X24,i(main_end)) | $less(0,X24) | $less(X23,sK0) | X23 = X24 | 0 = sK0 | sK0 = X23 (4:30) [resolution 377,323]
[SA] new: 2202. $less(alength,X25) | alength = i(main_end) | i(main_end) = X25 | $less(X26,i(main_end)) | $less(0,X26) | $less(X25,sK0) | sK0 = X25 | 0 = sK0 | sK0 = X26 (4:30) [resolution 377,332]
[SA] new: 2203. $less(alength,X27) | alength = i(main_end) | i(main_end) = X27 | $less(X28,i(main_end)) | $less(sK0,X28) | $less(X27,alength) | X27 = X28 | sK0 = X28 (4:27) [resolution 377,388]
[SA] new: 2204. $less(alength,X29) | alength = i(main_end) | i(main_end) = X29 | $less(X30,i(main_end)) | $less(sK0,X30) | $less(X29,alength) | sK0 = X30 | sK0 = X29 (4:27) [resolution 377,408]
[SA] new: 2205. $less(alength,alength) | alength = i(main_end) | alength = i(main_end) | $less(X31,i(main_end)) | $less(sK0,X31) | sK0 = X31 (4:21) [resolution 377,106]
[SA] new: 2206. $less(alength,alength) | alength = i(main_end) | alength = i(main_end) | $less(X32,i(main_end)) | $less(0,X32) | alength = X32 | 0 = sK0 (4:24) [resolution 377,178]
[SA] new: 2207. $less(alength,alength) | alength = i(main_end) | alength = i(main_end) | $less(X33,i(main_end)) | $less(sK0,X33) | alength = X33 (4:21) [resolution 377,179]
[SA] new: 2208. $less(alength,alength) | alength = i(main_end) | alength = i(main_end) | $less(X34,i(main_end)) | $less(sK0,X35) | $less(X35,X34) | X34 = X35 | sK0 = X34 (4:27) [resolution 377,198]
[SA] new: 2209. $less(alength,alength) | alength = i(main_end) | alength = i(main_end) | $less(X36,i(main_end)) | $less(sK0,X37) | $less(X37,X36) | alength = X36 | sK0 = X37 (4:27) [resolution 377,204]
[SA] new: 2210. $less(alength,alength) | alength = i(main_end) | alength = i(main_end) | $less(X38,i(main_end)) | $less(sK0,X39) | $less(X39,X38) | X38 = X39 | alength = X38 (4:27) [resolution 377,246]
[SA] new: 2211. $less(alength,alength) | alength = i(main_end) | alength = i(main_end) | $less(X40,i(main_end)) | $less(sK0,X41) | $less(X41,X40) | alength = X40 | alength = X41 (4:27) [resolution 377,253]
[SA] new: 2212. $less(alength,alength) | alength = i(main_end) | alength = i(main_end) | $less(X42,i(main_end)) | $less(0,X42) | alength = X42 | alength = 0 (5:24) [resolution 377,261]
[SA] new: 2213. $less(alength,alength) | alength = i(main_end) | alength = i(main_end) | $less(X43,i(main_end)) | $less(0,X43) | 0 = sK0 | 0 = X43 (5:24) [resolution 377,278]
[SA] new: 2214. $less(alength,alength) | alength = i(main_end) | alength = i(main_end) | $less(X44,i(main_end)) | $less(0,X44) | 0 = sK0 | sK0 = X44 (4:24) [resolution 377,330]
[SA] new: 2215. $less(alength,alength) | alength = i(main_end) | alength = i(main_end) | $less(X45,i(main_end)) | $less(sK0,X46) | $less(X46,X45) | X45 = X46 | sK0 = X46 (4:27) [resolution 377,388]
[SA] new: 2216. $less(alength,alength) | alength = i(main_end) | alength = i(main_end) | $less(X47,i(main_end)) | $less(sK0,X48) | $less(X48,X47) | sK0 = X48 | sK0 = X47 (4:27) [resolution 377,408]
[SA] new: 2217. $less(alength,i(main_end)) | alength = i(main_end) | i(main_end) = i(main_end) | $less(X49,i(main_end)) | $less(alength,X49) | alength = i(main_end) | alength = X49 (4:27) [resolution 377,409]
[SA] new: 2218. $less(alength,i(main_end)) | alength = i(main_end) | i(main_end) = i(main_end) | $less(X50,i(main_end)) | $less(alength,X50) | alength = X50 | i(main_end) = X50 (4:27) [resolution 377,389]
[SA] new: 2219. $less(alength,i(main_end)) | alength = i(main_end) | i(main_end) = i(main_end) | $less(X51,i(main_end)) | ~$less(X51,sK0) | alength = i(main_end) (5:24) [resolution 377,268]
[SA] new: 2220. $less(alength,i(main_end)) | alength = i(main_end) | i(main_end) = i(main_end) | $less(X52,i(main_end)) | ~$less(X52,sK0) | i(main_end) = sK0 (5:24) [resolution 377,224]
[SA] new: 2221. $less(alength,i(main_end)) | alength = i(main_end) | i(main_end) = i(main_end) | $less(X53,i(main_end)) | i(main_end) = X53 | $less(alength,X53) | alength = i(main_end) (4:28) [resolution 377,176]
[SA] new: 2222. $less(alength,i(main_end)) | alength = i(main_end) | i(main_end) = i(main_end) | $less(X54,i(main_end)) | ~$less(X54,alength) | alength = i(main_end) (4:24) [resolution 377,144]
[SA] new: 2223. $less(alength,0) | alength = i(main_end) | i(main_end) = 0 | $less(X55,i(main_end)) | $less(i(main_end),X56) | $less(X56,X55) | X55 = X56 | i(main_end) = X55 | i(l8(sK2(X55))) = X55 (4:35) [resolution 377,214]
[SA] new: 2224. $less(alength,sK0) | alength = i(main_end) | i(main_end) = sK0 | $less(X57,i(main_end)) | $less(0,X58) | $less(X58,X57) | sK0 = X57 | 0 = sK0 | sK0 = X58 (4:30) [resolution 377,332]
[SA] new: 2225. $less(alength,sK0) | alength = i(main_end) | i(main_end) = sK0 | $less(X59,i(main_end)) | $less(0,X60) | $less(X60,X59) | X59 = X60 | 0 = sK0 | sK0 = X59 (4:30) [resolution 377,323]
[SA] new: 2226. $less(alength,sK0) | alength = i(main_end) | i(main_end) = sK0 | $less(X61,i(main_end)) | $less(0,X61) | 0 = X61 | sK0 = X61 (4:24) [resolution 377,380]
[SA] new: 2227. $less(alength,sK0) | alength = i(main_end) | i(main_end) = sK0 | $less(X62,i(main_end)) | $less(0,X62) | 0 = sK0 | 0 = X62 (4:24) [resolution 377,229]
[SA] new: 2228. $less(alength,sK0) | alength = i(main_end) | i(main_end) = sK0 | $less(X63,i(main_end)) | $less(0,X63) | sK0 = X63 | 0 = sK0 (4:24) [resolution 377,177]
[SA] new: 2229. $less(alength,i(main_end)) | alength = i(main_end) | i(main_end) = i(main_end) | $less(alength,i(main_end)) | alength = i(main_end) (4:21) [resolution 377,100]
[SA] new: 2230. $less(alength,0) | alength = i(main_end) | i(main_end) = 0 | $less(alength,i(main_end)) | alength = i(main_end) (6:19) [resolution 377,1818]
[SA] new: 2231. $less(alength,0) | alength = i(main_end) | i(main_end) = 0 | $less(alength,i(main_end)) | alength = i(main_end) | alength = i(l8(sK2(alength))) (4:25) [resolution 377,142]
[SA] new: 2232. $less(alength,alength) | alength = i(main_end) | alength = i(main_end) | $less(i(l8(p(nl8))),i(main_end)) | zero = nl8 (4:21) [resolution 377,95]
[SA] new: 2233. $less(alength,alength) | alength = i(main_end) | alength = i(main_end) | $less(i(l8(p(p(nl8)))),i(main_end)) | zero = p(nl8) | zero = nl8 (4:26) [resolution 377,314]
[SA] new: 2234. $less(alength,alength) | alength = i(main_end) | alength = i(main_end) | $less(i(l8(sK2(X64))),i(main_end)) | $less(i(main_end),X64) | i(main_end) = X64 | $less(X64,0) (4:29) [resolution 377,195]
[SA] new: 2235. $less(alength,alength) | alength = i(main_end) | alength = i(main_end) | $less(i(l8(sK2(sK0))),i(main_end)) | alength = i(main_end) (5:22) [resolution 377,272]
[SA] new: 2236. $less(alength,alength) | alength = i(main_end) | alength = i(main_end) | $less(i(l8(sK2(sK0))),i(main_end)) | i(main_end) = sK0 (5:22) [resolution 377,228]
[SA] new: 2237. $less(alength,alength) | alength = i(main_end) | alength = i(main_end) | $less(i(l8(sK2(alength))),i(main_end)) | alength = i(main_end) | $less(alength,0) (4:25) [resolution 377,270]
[SA] new: 2238. $less(alength,X65) | alength = i(main_end) | i(main_end) = X65 | $less(i(main_end),i(main_end)) | i(l8(sK1(X66,sK2(X65)))) = X66 | $less(X65,0) | a(main_end,X66) = a(l11(sK2(X65)),X66) | i(main_end) = X65 (4:40) [resolution 377,225]
[SA] new: 2239. $less(alength,X67) | alength = i(main_end) | i(main_end) = X67 | $less(i(main_end),i(main_end)) | $less(X67,X68) | $less(X68,0) | X67 = X68 | i(main_end) = X68 | i(l8(sK2(X68))) = X68 (4:35) [resolution 377,214]
[SA] new: 2240. $less(alength,X69) | alength = i(main_end) | i(main_end) = X69 | $less(i(main_end),i(main_end)) | i(l8(sK2(X69))) = X69 | $less(X69,0) | i(main_end) = X69 (4:29) [resolution 377,127]
[SA] new: 2241. $less(alength,X70) | alength = i(main_end) | i(main_end) = X70 | $less(i(l8(p(nl8))),i(main_end)) | alength = X70 | $less(X70,alength) | zero = nl8 (4:27) [resolution 377,193]
[SA] new: 2242. $less(alength,sK0) | alength = i(main_end) | i(main_end) = sK0 | $less(i(main_end),i(main_end)) | i(main_end) = sK0 | 0 = sK0 | i(l8(sK2(0))) = 0 (4:29) [resolution 377,337]
[SA] new: 2243. $less(alength,sK0) | alength = i(main_end) | i(main_end) = sK0 | $less(i(main_end),i(main_end)) | 0 = sK0 | i(l8(sK2(0))) = 0 | i(main_end) = 0 (4:29) [resolution 377,234]
[SA] new: 2244. $less(alength,X71) | alength = i(main_end) | i(main_end) = X71 | $less(0,i(main_end)) | $less(X71,sK0) | sK0 = X71 | 0 = sK0 (4:24) [resolution 377,177]
[SA] new: 2245. $less(alength,X72) | alength = i(main_end) | i(main_end) = X72 | $less(0,i(main_end)) | $less(X72,alength) | alength = X72 | 0 = sK0 (4:24) [resolution 377,178]
[SA] new: 2246. $less(alength,X73) | alength = i(main_end) | i(main_end) = X73 | $less(0,i(main_end)) | $less(X73,sK0) | 0 = sK0 | 0 = X73 (4:24) [resolution 377,229]
[SA] new: 2247. $less(alength,X74) | alength = i(main_end) | i(main_end) = X74 | $less(0,i(main_end)) | $less(X74,alength) | alength = X74 | alength = 0 (5:24) [resolution 377,261]
[SA] new: 2248. $less(alength,X75) | alength = i(main_end) | i(main_end) = X75 | $less(0,i(main_end)) | $less(X75,alength) | 0 = sK0 | 0 = X75 (5:24) [resolution 377,278]
[SA] new: 2249. $less(alength,X76) | alength = i(main_end) | i(main_end) = X76 | $less(0,i(main_end)) | $less(X76,X77) | $less(X77,sK0) | X76 = X77 | 0 = sK0 | sK0 = X77 (4:30) [resolution 377,323]
[SA] new: 2250. $less(alength,X78) | alength = i(main_end) | i(main_end) = X78 | $less(0,i(main_end)) | $less(X78,alength) | 0 = sK0 | sK0 = X78 (4:24) [resolution 377,330]
[SA] new: 2251. $less(alength,X79) | alength = i(main_end) | i(main_end) = X79 | $less(0,i(main_end)) | $less(X80,sK0) | $less(X79,X80) | sK0 = X80 | 0 = sK0 | sK0 = X79 (4:30) [resolution 377,332]
[SA] new: 2252. $less(alength,X81) | alength = i(main_end) | i(main_end) = X81 | $less(0,i(main_end)) | $less(X81,sK0) | 0 = X81 | sK0 = X81 (4:24) [resolution 377,380]
[SA] new: 2253. $less(alength,alength) | alength = i(main_end) | alength = i(main_end) | $less(0,i(main_end)) | alength = 0 (4:18) [resolution 377,244]
[SA] new: 2254. $less(alength,alength) | alength = i(main_end) | alength = i(main_end) | $less(0,i(main_end)) | 0 = sK0 (4:18) [resolution 377,113]
[SA] new: 2255. $less(alength,i(main_end)) | alength = i(main_end) | i(main_end) = i(main_end) | $less(0,i(main_end)) | 0 = sK0 | i(main_end) = 0 (6:24) [resolution 377,1619]
[SA] new: 2256. $less(alength,i(main_end)) | alength = i(main_end) | i(main_end) = i(main_end) | $less(0,i(main_end)) | alength = i(main_end) | alength = 0 (6:24) [resolution 377,1429]
[SA] new: 2257. $less(alength,i(main_end)) | alength = i(main_end) | i(main_end) = i(main_end) | $less(0,i(main_end)) | i(main_end) = sK0 | 0 = sK0 (7:24) [resolution 377,761]
[SA] new: 2258. $less(alength,i(main_end)) | alength = i(main_end) | i(main_end) = i(main_end) | $less(0,i(main_end)) | alength = i(main_end) | 0 = sK0 (5:24) [resolution 377,583]
[SA] new: 2259. $less(alength,sK0) | alength = i(main_end) | i(main_end) = sK0 | $less(0,i(main_end)) | 0 = sK0 (4:18) [resolution 377,97]
[SA] new: 2260. $less(alength,alength) | alength = i(main_end) | alength = i(main_end) | $less(sK0,i(main_end)) (4:15) [resolution 377,45]
[SA] new: 2261. $less(alength,X82) | alength = i(main_end) | i(main_end) = X82 | $less(sK0,i(main_end)) | $less(X82,alength) | sK0 = X82 (4:21) [resolution 377,106]
[SA] new: 2262. $less(alength,X83) | alength = i(main_end) | i(main_end) = X83 | $less(sK0,i(main_end)) | $less(X83,alength) | alength = X83 (4:21) [resolution 377,179]
[SA] new: 2263. $less(alength,X84) | alength = i(main_end) | i(main_end) = X84 | $less(sK0,i(main_end)) | $less(X84,X85) | $less(X85,alength) | X84 = X85 | sK0 = X85 (4:27) [resolution 377,198]
[SA] new: 2264. $less(alength,X86) | alength = i(main_end) | i(main_end) = X86 | $less(sK0,i(main_end)) | $less(X87,alength) | $less(X86,X87) | alength = X87 | sK0 = X86 (4:27) [resolution 377,204]
[SA] new: 2265. $less(alength,X88) | alength = i(main_end) | i(main_end) = X88 | $less(sK0,i(main_end)) | $less(X88,X89) | $less(X89,alength) | X88 = X89 | alength = X89 (4:27) [resolution 377,246]
[SA] new: 2266. $less(alength,X90) | alength = i(main_end) | i(main_end) = X90 | $less(sK0,i(main_end)) | $less(X91,alength) | $less(X90,X91) | alength = X91 | alength = X90 (4:27) [resolution 377,253]
[SA] new: 2267. $less(alength,X92) | alength = i(main_end) | i(main_end) = X92 | $less(sK0,i(main_end)) | $less(X93,alength) | $less(X92,X93) | X92 = X93 | sK0 = X92 (4:27) [resolution 377,388]
[SA] new: 2268. $less(alength,X94) | alength = i(main_end) | i(main_end) = X94 | $less(sK0,i(main_end)) | $less(X95,alength) | $less(X94,X95) | sK0 = X94 | sK0 = X95 (4:27) [resolution 377,408]
[SA] new: 2269. $less(alength,i(main_end)) | alength = i(main_end) | i(main_end) = i(main_end) | $less(sK0,i(main_end)) | alength = i(main_end) (4:21) [resolution 377,255]
[SA] new: 2270. $less(alength,i(main_end)) | alength = i(main_end) | i(main_end) = i(main_end) | $less(sK0,i(main_end)) | i(main_end) = sK0 (4:21) [resolution 377,206]
[SA] new: 2271. $less(alength,i(main_end)) | alength = i(main_end) | i(main_end) = i(main_end) | $less(sK0,i(main_end)) (4:17) [duplicate literal removal 2269]
[SA] new: 2272. $less(alength,alength) | alength = i(main_end) | $less(sK0,i(main_end)) (4:11) [duplicate literal removal 2260]
[SA] new: 2273. $less(alength,i(main_end)) | alength = i(main_end) | i(main_end) = i(main_end) | $less(0,i(main_end)) | 0 = sK0 (5:20) [duplicate literal removal 2258]
[SA] new: 2274. $less(alength,i(main_end)) | alength = i(main_end) | i(main_end) = i(main_end) | $less(0,i(main_end)) | alength = 0 (6:20) [duplicate literal removal 2256]
[SA] new: 2275. $less(alength,alength) | alength = i(main_end) | $less(0,i(main_end)) | 0 = sK0 (4:14) [duplicate literal removal 2254]
[SA] new: 2276. $less(alength,alength) | alength = i(main_end) | $less(0,i(main_end)) | alength = 0 (4:14) [duplicate literal removal 2253]
[SA] new: 2277. $less(alength,sK0) | alength = i(main_end) | i(main_end) = sK0 | $less(i(main_end),i(main_end)) | 0 = sK0 | i(l8(sK2(0))) = 0 (4:25) [duplicate literal removal 2242]
[SA] new: 2278. $less(alength,X69) | alength = i(main_end) | i(main_end) = X69 | $less(i(main_end),i(main_end)) | i(l8(sK2(X69))) = X69 | $less(X69,0) (4:25) [duplicate literal removal 2240]
[SA] new: 2279. $less(alength,X65) | alength = i(main_end) | i(main_end) = X65 | $less(i(main_end),i(main_end)) | i(l8(sK1(X66,sK2(X65)))) = X66 | $less(X65,0) | a(main_end,X66) = a(l11(sK2(X65)),X66) (4:36) [duplicate literal removal 2238]
[SA] new: 2280. $less(alength,alength) | alength = i(main_end) | $less(i(l8(sK2(alength))),i(main_end)) | $less(alength,0) (4:17) [duplicate literal removal 2237]
[SA] new: 2281. $less(alength,alength) | alength = i(main_end) | $less(i(l8(sK2(sK0))),i(main_end)) | i(main_end) = sK0 (5:18) [duplicate literal removal 2236]
[SA] new: 2282. $less(alength,alength) | alength = i(main_end) | $less(i(l8(sK2(sK0))),i(main_end)) (5:14) [duplicate literal removal 2235]
[SA] new: 2283. $less(alength,alength) | alength = i(main_end) | $less(i(l8(sK2(X64))),i(main_end)) | $less(i(main_end),X64) | i(main_end) = X64 | $less(X64,0) (4:25) [duplicate literal removal 2234]
[SA] new: 2284. $less(alength,alength) | alength = i(main_end) | $less(i(l8(p(p(nl8)))),i(main_end)) | zero = p(nl8) | zero = nl8 (4:22) [duplicate literal removal 2233]
[SA] new: 2285. $less(alength,alength) | alength = i(main_end) | $less(i(l8(p(nl8))),i(main_end)) | zero = nl8 (4:17) [duplicate literal removal 2232]
[SA] new: 2286. $less(alength,0) | alength = i(main_end) | i(main_end) = 0 | $less(alength,i(main_end)) | alength = i(l8(sK2(alength))) (4:21) [duplicate literal removal 2231]
[SA] new: 2287. $less(alength,0) | alength = i(main_end) | i(main_end) = 0 | $less(alength,i(main_end)) (6:15) [duplicate literal removal 2230]
[SA] new: 2288. $less(alength,i(main_end)) | alength = i(main_end) | i(main_end) = i(main_end) (4:13) [duplicate literal removal 2229]
[SA] new: 2289. $less(alength,i(main_end)) | alength = i(main_end) | i(main_end) = i(main_end) | $less(X54,i(main_end)) | ~$less(X54,alength) (4:20) [duplicate literal removal 2222]
[SA] new: 2290. $less(alength,i(main_end)) | alength = i(main_end) | i(main_end) = i(main_end) | $less(X53,i(main_end)) | i(main_end) = X53 | $less(alength,X53) (4:24) [duplicate literal removal 2221]
[SA] new: 2291. $less(alength,i(main_end)) | alength = i(main_end) | i(main_end) = i(main_end) | $less(X51,i(main_end)) | ~$less(X51,sK0) (5:20) [duplicate literal removal 2219]
[SA] new: 2292. $less(alength,i(main_end)) | alength = i(main_end) | i(main_end) = i(main_end) | $less(X49,i(main_end)) | $less(alength,X49) | alength = X49 (4:23) [duplicate literal removal 2217]
[SA] new: 2293. $less(alength,alength) | alength = i(main_end) | $less(X47,i(main_end)) | $less(sK0,X48) | $less(X48,X47) | sK0 = X48 | sK0 = X47 (4:23) [duplicate literal removal 2216]
[SA] new: 2294. $less(alength,alength) | alength = i(main_end) | $less(X45,i(main_end)) | $less(sK0,X46) | $less(X46,X45) | X45 = X46 | sK0 = X46 (4:23) [duplicate literal removal 2215]
[SA] new: 2295. $less(alength,alength) | alength = i(main_end) | $less(X44,i(main_end)) | $less(0,X44) | 0 = sK0 | sK0 = X44 (4:20) [duplicate literal removal 2214]
[SA] new: 2296. $less(alength,alength) | alength = i(main_end) | $less(X43,i(main_end)) | $less(0,X43) | 0 = sK0 | 0 = X43 (5:20) [duplicate literal removal 2213]
[SA] new: 2297. $less(alength,alength) | alength = i(main_end) | $less(X42,i(main_end)) | $less(0,X42) | alength = X42 | alength = 0 (5:20) [duplicate literal removal 2212]
[SA] new: 2298. $less(alength,alength) | alength = i(main_end) | $less(X40,i(main_end)) | $less(sK0,X41) | $less(X41,X40) | alength = X40 | alength = X41 (4:23) [duplicate literal removal 2211]
[SA] new: 2299. $less(alength,alength) | alength = i(main_end) | $less(X38,i(main_end)) | $less(sK0,X39) | $less(X39,X38) | X38 = X39 | alength = X38 (4:23) [duplicate literal removal 2210]
[SA] new: 2300. $less(alength,alength) | alength = i(main_end) | $less(X36,i(main_end)) | $less(sK0,X37) | $less(X37,X36) | alength = X36 | sK0 = X37 (4:23) [duplicate literal removal 2209]
[SA] new: 2301. $less(alength,alength) | alength = i(main_end) | $less(X34,i(main_end)) | $less(sK0,X35) | $less(X35,X34) | X34 = X35 | sK0 = X34 (4:23) [duplicate literal removal 2208]
[SA] new: 2302. $less(alength,alength) | alength = i(main_end) | $less(X33,i(main_end)) | $less(sK0,X33) | alength = X33 (4:17) [duplicate literal removal 2207]
[SA] new: 2303. $less(alength,alength) | alength = i(main_end) | $less(X32,i(main_end)) | $less(0,X32) | alength = X32 | 0 = sK0 (4:20) [duplicate literal removal 2206]
[SA] new: 2304. $less(alength,alength) | alength = i(main_end) | $less(X31,i(main_end)) | $less(sK0,X31) | sK0 = X31 (4:17) [duplicate literal removal 2205]
[SA] new: 2305. $less(alength,X17) | alength = i(main_end) | i(main_end) = X17 | $less(X18,i(main_end)) | $less(i(main_end),X18) | $less(X17,0) | X17 = X18 | i(l8(sK2(X17))) = X17 (4:31) [duplicate literal removal 2198]
[SA] passive: 2191. $less(alength,X0) | alength = i(main_end) | i(main_end) = X0 | $less(X1,i(main_end)) | $less(X0,X1) | X0 = X1 (4:21) [resolution 377,15]
[SA] passive: 2193. $less(alength,X4) | alength = i(main_end) | i(main_end) = X4 | $less(X5,i(main_end)) | $less(X6,X5) | $less(X4,X6) | X4 = X5 | X4 = X6 (4:27) [resolution 377,174]
[SA] passive: 2194. $less(alength,X7) | alength = i(main_end) | i(main_end) = X7 | $less(X8,i(main_end)) | $less(X7,X9) | $less(X9,X8) | X8 = X9 | X7 = X8 (4:27) [resolution 377,174]
[SA] passive: 2195. $less(alength,X10) | alength = i(main_end) | i(main_end) = X10 | $less(X11,i(main_end)) | $less(X12,X11) | $less(X10,X12) | X10 = X12 | X11 = X12 (4:27) [resolution 377,174]
[SA] passive: 2196. $less(alength,X13) | alength = i(main_end) | i(main_end) = X13 | $less(X14,i(main_end)) | $less(sK0,X14) | $less(X13,alength) | X13 = X14 | sK0 = X13 (4:27) [resolution 377,198]
[SA] new: 2306. alength = i(main_end) | $less(X34,i(main_end)) | $less(sK0,X35) | $less(X35,X34) | X34 = X35 | sK0 = X34 (4:20) [subsumption resolution 2301,13]
[SA] passive: 2306. alength = i(main_end) | $less(X34,i(main_end)) | $less(sK0,X35) | $less(X35,X34) | X34 = X35 | sK0 = X34 (4:20) [subsumption resolution 2301,13]
[SA] new: 2307. alength = i(main_end) | $less(X36,i(main_end)) | $less(sK0,X37) | $less(X37,X36) | alength = X36 | sK0 = X37 (4:20) [subsumption resolution 2300,13]
[SA] passive: 2307. alength = i(main_end) | $less(X36,i(main_end)) | $less(sK0,X37) | $less(X37,X36) | alength = X36 | sK0 = X37 (4:20) [subsumption resolution 2300,13]
[SA] new: 2308. alength = i(main_end) | $less(X38,i(main_end)) | $less(sK0,X39) | $less(X39,X38) | X38 = X39 | alength = X38 (4:20) [subsumption resolution 2299,13]
[SA] passive: 2308. alength = i(main_end) | $less(X38,i(main_end)) | $less(sK0,X39) | $less(X39,X38) | X38 = X39 | alength = X38 (4:20) [subsumption resolution 2299,13]
[SA] new: 2309. alength = i(main_end) | $less(X40,i(main_end)) | $less(sK0,X41) | $less(X41,X40) | alength = X40 | alength = X41 (4:20) [subsumption resolution 2298,13]
[SA] passive: 2309. alength = i(main_end) | $less(X40,i(main_end)) | $less(sK0,X41) | $less(X41,X40) | alength = X40 | alength = X41 (4:20) [subsumption resolution 2298,13]
[SA] new: 2310. alength = i(main_end) | $less(X45,i(main_end)) | $less(sK0,X46) | $less(X46,X45) | X45 = X46 | sK0 = X46 (4:20) [subsumption resolution 2294,13]
[SA] passive: 2310. alength = i(main_end) | $less(X45,i(main_end)) | $less(sK0,X46) | $less(X46,X45) | X45 = X46 | sK0 = X46 (4:20) [subsumption resolution 2294,13]
[SA] new: 2311. alength = i(main_end) | $less(X47,i(main_end)) | $less(sK0,X48) | $less(X48,X47) | sK0 = X48 | sK0 = X47 (4:20) [subsumption resolution 2293,13]
[SA] passive: 2311. alength = i(main_end) | $less(X47,i(main_end)) | $less(sK0,X48) | $less(X48,X47) | sK0 = X48 | sK0 = X47 (4:20) [subsumption resolution 2293,13]
[SA] new: 2312. alength = i(main_end) | $less(i(l8(p(nl8))),i(main_end)) | zero = nl8 (4:14) [subsumption resolution 2285,13]
[SA] passive: 2312. alength = i(main_end) | $less(i(l8(p(nl8))),i(main_end)) | zero = nl8 (4:14) [subsumption resolution 2285,13]
[SA] new: 2313. alength = i(main_end) | $less(i(l8(p(p(nl8)))),i(main_end)) | zero = p(nl8) | zero = nl8 (4:19) [subsumption resolution 2284,13]
[SA] passive: 2313. alength = i(main_end) | $less(i(l8(p(p(nl8)))),i(main_end)) | zero = p(nl8) | zero = nl8 (4:19) [subsumption resolution 2284,13]
[SA] new: 2314. alength = i(main_end) | $less(i(l8(sK2(X64))),i(main_end)) | $less(i(main_end),X64) | i(main_end) = X64 | $less(X64,0) (4:22) [subsumption resolution 2283,13]
[SA] passive: 2314. alength = i(main_end) | $less(i(l8(sK2(X64))),i(main_end)) | $less(i(main_end),X64) | i(main_end) = X64 | $less(X64,0) (4:22) [subsumption resolution 2283,13]
[SA] new: 2315. alength = i(main_end) | $less(i(l8(sK2(sK0))),i(main_end)) (5:11) [subsumption resolution 2282,13]
[SA] passive: 2315. alength = i(main_end) | $less(i(l8(sK2(sK0))),i(main_end)) (5:11) [subsumption resolution 2282,13]
[SA] passive: 2252. $less(alength,X81) | alength = i(main_end) | i(main_end) = X81 | $less(0,i(main_end)) | $less(X81,sK0) | 0 = X81 | sK0 = X81 (4:24) [resolution 377,380]
[SA] active: 1969. alength != 0 | zero = nl8 | 0 = sK0 (7:9:1) [superposition 146,1845]
[SA] active: 387. $less(0,X17) | $less(X18,sK0) | $less(X17,X18) | X17 = X18 | 0 = sK0 | 0 = X17 (3:18:3) [resolution 174,114]
[SA] new: 2316. $less(X0,sK0) | $less(0,X0) | 0 = X0 | 0 = sK0 | 0 = 0 (4:15) [resolution 387,13]
[SA] new: 2317. $less(X1,sK0) | $less(X2,X1) | X1 = X2 | 0 = sK0 | 0 = X2 | ~$less(X3,0) | $less(X3,X2) (4:21) [resolution 387,14]
[SA] new: 2318. $less(X4,sK0) | $less(X5,X4) | X4 = X5 | 0 = sK0 | 0 = X5 | X5 = X6 | $less(X6,X5) | $less(0,X6) (4:24) [resolution 387,99]
[SA] new: 2319. $less(X7,sK0) | $less(X8,X7) | X7 = X8 | 0 = sK0 | 0 = X8 | sK0 = X8 | $less(sK0,X8) | $less(0,alength) (4:24) [resolution 387,205]
[SA] new: 2320. $less(X9,sK0) | $less(X10,X9) | X9 = X10 | 0 = sK0 | 0 = X10 | alength = X10 | $less(sK0,X10) | $less(0,alength) (4:24) [resolution 387,254]
[SA] new: 2321. $less(X11,sK0) | $less(X12,X11) | X11 = X12 | 0 = sK0 | 0 = X12 | alength = i(main_end) | ~$less(X12,alength) | $less(0,i(main_end)) (4:26) [resolution 387,277]
[SA] new: 2322. $less(X13,sK0) | $less(X14,X13) | X13 = X14 | 0 = sK0 | 0 = X14 | sK0 = X14 | 0 = sK0 | $less(0,X14) | $less(0,sK0) (4:27) [resolution 387,333]
[SA] new: 2323. $less(X15,sK0) | $less(X16,X15) | X15 = X16 | 0 = sK0 | 0 = X16 | $less(alength,X16) | alength = i(main_end) | i(main_end) = X16 | $less(0,i(main_end)) (4:30) [resolution 387,377]
[SA] new: 2324. $less(X17,sK0) | $less(alength,X17) | alength = X17 | 0 = sK0 | alength = 0 | alength = i(main_end) | i(main_end) = X18 | $less(X18,i(main_end)) | $less(0,X18) (4:30) [resolution 387,276]
[SA] new: 2325. $less(X19,sK0) | $less(alength,X19) | alength = X19 | 0 = sK0 | alength = 0 | alength = i(l8(sK2(alength))) | alength = i(main_end) | $less(0,0) (4:28) [resolution 387,365]
[SA] new: 2326. $less(X20,sK0) | $less(alength,X20) | alength = X20 | 0 = sK0 | alength = 0 | alength = i(main_end) | $less(0,0) | i(l8(sK2(0))) = 0 (4:28) [resolution 387,273]
[SA] new: 2327. $less(X21,sK0) | $less(i(main_end),X21) | i(main_end) = X21 | 0 = sK0 | i(main_end) = 0 | $less(X22,0) | i(main_end) = X22 | i(l8(sK2(X22))) = X22 | $less(0,X22) (4:34) [resolution 387,215]
[SA] new: 2328. $less(X23,sK0) | $less(i(main_end),X23) | i(main_end) = X23 | 0 = sK0 | i(main_end) = 0 | $less(0,0) | i(l8(sK2(0))) = 0 (4:27) [resolution 387,84]
[SA] new: 2329. $less(X24,sK0) | $less(i(main_end),X24) | i(main_end) = X24 | 0 = sK0 | i(main_end) = 0 | $less(0,0) | 'Sub'(sK2(0),nl8) (4:25) [resolution 387,83]
[SA] new: 2330. $less(X25,sK0) | $less(i(l8(p(nl8))),X25) | i(l8(p(nl8))) = X25 | 0 = sK0 | i(l8(p(nl8))) = 0 | zero = nl8 | $less(0,alength) (4:30) [resolution 387,194]
[SA] new: 2331. $less(X26,sK0) | $less(0,X26) | 0 = X26 | 0 = sK0 | 0 = 0 | sK0 = X27 | 0 = sK0 | $less(X27,sK0) | $less(0,X27) (4:27) [resolution 387,324]
[SA] new: 2332. $less(X28,sK0) | $less(0,X28) | 0 = X28 | 0 = sK0 | 0 = 0 | alength = 0 | $less(0,alength) (5:21) [resolution 387,262]
[SA] new: 2333. $less(X29,sK0) | $less(0,X29) | 0 = X29 | 0 = sK0 | 0 = 0 | 0 = sK0 | $less(0,alength) (4:21) [resolution 387,116]
[SA] new: 2334. $less(X30,sK0) | $less(0,X30) | 0 = X30 | 0 = sK0 | 0 = 0 | 0 = sK0 | $less(0,sK0) (4:21) [resolution 387,114]
[SA] new: 2335. $less(X31,sK0) | $less(sK0,X31) | sK0 = X31 | 0 = sK0 | 0 = sK0 | alength = X32 | $less(X32,alength) | $less(0,X32) (4:24) [resolution 387,247]
[SA] new: 2336. $less(X33,sK0) | $less(sK0,X33) | sK0 = X33 | 0 = sK0 | 0 = sK0 | sK0 = X34 | $less(X34,alength) | $less(0,X34) (4:24) [resolution 387,199]
[SA] new: 2337. $less(X35,sK0) | $less(sK0,X35) | sK0 = X35 | 0 = sK0 | 0 = sK0 | $less(0,alength) (4:18) [resolution 387,96]
[SA] new: 2338. $less(0,X0) | $less(X0,X1) | X0 = X1 | 0 = sK0 | 0 = X0 | alength = X2 | $less(X2,alength) | $less(X1,X2) (4:24) [resolution 387,247]
[SA] new: 2339. $less(0,X3) | $less(X3,X4) | X3 = X4 | 0 = sK0 | 0 = X3 | sK0 = X5 | $less(X5,alength) | $less(X4,X5) (4:24) [resolution 387,199]
[SA] new: 2340. $less(0,X6) | $less(X6,X7) | X6 = X7 | 0 = sK0 | 0 = X6 | $less(X7,alength) (4:18) [resolution 387,96]
[SA] new: 2341. $less(0,X8) | $less(X8,sK0) | sK0 = X8 | 0 = sK0 | 0 = X8 (4:15) [resolution 387,13]
[SA] new: 2342. $less(0,X9) | $less(X9,X10) | X9 = X10 | 0 = sK0 | 0 = X9 | ~$less(X11,X10) | $less(X11,sK0) (4:21) [resolution 387,14]
[SA] new: 2343. $less(0,X12) | $less(X12,X13) | X12 = X13 | 0 = sK0 | 0 = X12 | sK0 = X14 | $less(X14,sK0) | $less(X13,X14) (4:24) [resolution 387,99]
[SA] new: 2344. $less(0,X15) | $less(X15,X16) | X15 = X16 | 0 = sK0 | 0 = X15 | sK0 = sK0 | $less(sK0,sK0) | $less(X16,alength) (4:24) [resolution 387,205]
[SA] new: 2345. $less(0,X17) | $less(X17,X18) | X17 = X18 | 0 = sK0 | 0 = X17 | alength = sK0 | $less(sK0,sK0) | $less(X18,alength) (4:24) [resolution 387,254]
[SA] new: 2346. $less(0,X19) | $less(X19,X20) | X19 = X20 | 0 = sK0 | 0 = X19 | alength = i(main_end) | ~$less(sK0,alength) | $less(X20,i(main_end)) (4:26) [resolution 387,277]
[SA] new: 2347. $less(0,X21) | $less(X21,X22) | X21 = X22 | 0 = sK0 | 0 = X21 | sK0 = sK0 | 0 = sK0 | $less(0,sK0) | $less(X22,sK0) (4:27) [resolution 387,333]
[SA] new: 2348. $less(0,X23) | $less(X23,X24) | X23 = X24 | 0 = sK0 | 0 = X23 | $less(alength,sK0) | alength = i(main_end) | i(main_end) = sK0 | $less(X24,i(main_end)) (4:30) [resolution 387,377]
[SA] new: 2349. $less(0,X25) | $less(X25,i(main_end)) | i(main_end) = X25 | 0 = sK0 | 0 = X25 | alength = i(main_end) (6:21) [resolution 387,812]
[SA] new: 2350. $less(0,X26) | $less(X26,i(main_end)) | i(main_end) = X26 | 0 = sK0 | 0 = X26 | i(main_end) = sK0 (6:21) [resolution 387,722]
[SA] new: 2351. $less(0,sK0) | $less(0,sK0) | 0 = sK0 | 0 = sK0 | 0 = sK0 (4:15) [resolution 387,46]
[SA] new: 2352. $less(0,X0) | $less(X0,sK0) | X0 = X0 | 0 = sK0 | 0 = X0 (4:15) [resolution 387,13]
[SA] new: 2353. $less(0,X1) | $less(X2,sK0) | X1 = X2 | 0 = sK0 | 0 = X1 | ~$less(X3,X1) | $less(X3,X2) (4:21) [resolution 387,14]
[SA] new: 2354. $less(0,X4) | $less(X5,sK0) | X4 = X5 | 0 = sK0 | 0 = X4 | X5 = X6 | $less(X6,X5) | $less(X4,X6) (4:24) [resolution 387,99]
[SA] new: 2355. $less(0,X7) | $less(X8,sK0) | X7 = X8 | 0 = sK0 | 0 = X7 | sK0 = X8 | $less(sK0,X8) | $less(X7,alength) (4:24) [resolution 387,205]
[SA] new: 2356. $less(0,X9) | $less(X10,sK0) | X9 = X10 | 0 = sK0 | 0 = X9 | alength = X10 | $less(sK0,X10) | $less(X9,alength) (4:24) [resolution 387,254]
[SA] new: 2357. $less(0,X11) | $less(X12,sK0) | X11 = X12 | 0 = sK0 | 0 = X11 | alength = i(main_end) | ~$less(X12,alength) | $less(X11,i(main_end)) (4:26) [resolution 387,277]
[SA] new: 2358. $less(0,X13) | $less(X14,sK0) | X13 = X14 | 0 = sK0 | 0 = X13 | sK0 = X14 | 0 = sK0 | $less(0,X14) | $less(X13,sK0) (4:27) [resolution 387,333]
[SA] new: 2359. $less(0,X15) | $less(X16,sK0) | X15 = X16 | 0 = sK0 | 0 = X15 | $less(alength,X16) | alength = i(main_end) | i(main_end) = X16 | $less(X15,i(main_end)) (4:30) [resolution 387,377]
[SA] new: 2360. $less(0,X17) | $less(alength,sK0) | alength = X17 | 0 = sK0 | 0 = X17 | alength = i(main_end) | i(main_end) = X18 | $less(X18,i(main_end)) | $less(X17,X18) (4:30) [resolution 387,276]
[SA] new: 2361. $less(0,X19) | $less(alength,sK0) | alength = X19 | 0 = sK0 | 0 = X19 | alength = i(l8(sK2(alength))) | alength = i(main_end) | $less(X19,0) (4:28) [resolution 387,365]
[SA] new: 2362. $less(0,X20) | $less(alength,sK0) | alength = X20 | 0 = sK0 | 0 = X20 | alength = i(main_end) | $less(X20,0) | i(l8(sK2(X20))) = X20 (4:28) [resolution 387,273]
[SA] new: 2363. $less(0,X21) | $less(i(main_end),sK0) | i(main_end) = X21 | 0 = sK0 | 0 = X21 | $less(X22,0) | i(main_end) = X22 | i(l8(sK2(X22))) = X22 | $less(X21,X22) (4:33) [resolution 387,215]
[SA] new: 2364. $less(0,X23) | $less(i(main_end),sK0) | i(main_end) = X23 | 0 = sK0 | 0 = X23 | $less(X23,0) | i(l8(sK2(X23))) = X23 (4:26) [resolution 387,84]
[SA] new: 2365. $less(0,X24) | $less(i(main_end),sK0) | i(main_end) = X24 | 0 = sK0 | 0 = X24 | $less(X24,0) | 'Sub'(sK2(X24),nl8) (4:24) [resolution 387,83]
[SA] new: 2366. $less(0,X25) | $less(i(l8(p(nl8))),sK0) | i(l8(p(nl8))) = X25 | 0 = sK0 | 0 = X25 | zero = nl8 | $less(X25,alength) (4:27) [resolution 387,194]
[SA] new: 2367. $less(0,X26) | $less(0,sK0) | 0 = X26 | 0 = sK0 | 0 = X26 | sK0 = X27 | 0 = sK0 | $less(X27,sK0) | $less(X26,X27) (4:27) [resolution 387,324]
[SA] new: 2368. $less(0,X28) | $less(0,sK0) | 0 = X28 | 0 = sK0 | 0 = X28 | alength = 0 | $less(X28,alength) (5:21) [resolution 387,262]
[SA] new: 2369. $less(0,X29) | $less(0,sK0) | 0 = X29 | 0 = sK0 | 0 = X29 | 0 = sK0 | $less(X29,alength) (4:21) [resolution 387,116]
[SA] new: 2370. $less(0,X30) | $less(0,sK0) | 0 = X30 | 0 = sK0 | 0 = X30 | 0 = sK0 | $less(X30,sK0) (4:21) [resolution 387,114]
[SA] new: 2371. $less(0,X31) | $less(sK0,sK0) | sK0 = X31 | 0 = sK0 | 0 = X31 | alength = X32 | $less(X32,alength) | $less(X31,X32) (4:24) [resolution 387,247]
[SA] new: 2372. $less(0,X33) | $less(sK0,sK0) | sK0 = X33 | 0 = sK0 | 0 = X33 | sK0 = X34 | $less(X34,alength) | $less(X33,X34) (4:24) [resolution 387,199]
[SA] new: 2373. $less(0,X35) | $less(sK0,sK0) | sK0 = X35 | 0 = sK0 | 0 = X35 | $less(X35,alength) (4:18) [resolution 387,96]
[SA] new: 2374. $less(0,i(main_end)) | $less(alength,sK0) | alength = i(main_end) | 0 = sK0 | i(main_end) = 0 (4:18) [resolution 387,82]
[SA] new: 2375. $less(0,i(main_end)) | $less(sK0,sK0) | i(main_end) = sK0 | 0 = sK0 | i(main_end) = 0 | alength = i(main_end) (6:22) [resolution 387,812]
[SA] new: 2376. $less(0,i(main_end)) | $less(sK0,sK0) | i(main_end) = sK0 | 0 = sK0 | i(main_end) = 0 | i(main_end) = sK0 (6:22) [resolution 387,722]
[SA] new: 2377. $less(0,sK0) | $less(sK0,0) | 0 = sK0 | 0 = sK0 | 0 = sK0 (4:15) [factoring 387]
[SA] new: 2378. $less(0,0) | $less(0,sK0) | 0 = 0 | 0 = sK0 | 0 = 0 (4:15) [factoring 387]
[SA] new: 2379. $less(0,sK0) | $less(sK0,sK0) | sK0 = sK0 | 0 = sK0 | 0 = sK0 (4:15) [factoring 387]
[SA] new: 2380. $less(0,sK0) | $less(sK0,sK0) | sK0 = sK0 | 0 = sK0 (4:12) [duplicate literal removal 2379]
[SA] new: 2381. $less(0,0) | $less(0,sK0) | 0 = 0 | 0 = sK0 (4:12) [duplicate literal removal 2378]
[SA] new: 2382. $less(0,sK0) | $less(sK0,0) | 0 = sK0 (4:9) [duplicate literal removal 2377]
[SA] new: 2383. $less(0,i(main_end)) | $less(sK0,sK0) | i(main_end) = sK0 | 0 = sK0 | i(main_end) = 0 (6:18) [duplicate literal removal 2376]
[SA] new: 2384. $less(0,X30) | $less(0,sK0) | 0 = X30 | 0 = sK0 | $less(X30,sK0) (4:15) [duplicate literal removal 2370]
[SA] new: 2385. $less(0,X29) | $less(0,sK0) | 0 = X29 | 0 = sK0 | $less(X29,alength) (4:15) [duplicate literal removal 2369]
[SA] new: 2386. $less(0,X28) | $less(0,sK0) | 0 = X28 | 0 = sK0 | alength = 0 | $less(X28,alength) (5:18) [duplicate literal removal 2368]
[SA] new: 2387. $less(0,X26) | $less(0,sK0) | 0 = X26 | 0 = sK0 | sK0 = X27 | $less(X27,sK0) | $less(X26,X27) (4:21) [duplicate literal removal 2367]
[SA] new: 2388. $less(0,X13) | $less(X14,sK0) | X13 = X14 | 0 = sK0 | 0 = X13 | sK0 = X14 | $less(0,X14) | $less(X13,sK0) (4:24) [duplicate literal removal 2358]
[SA] new: 2389. $less(0,sK0) | 0 = sK0 (4:6) [duplicate literal removal 2351]
[SA] new: 2390. $less(0,X21) | $less(X21,X22) | X21 = X22 | 0 = sK0 | 0 = X21 | sK0 = sK0 | $less(0,sK0) | $less(X22,sK0) (4:24) [duplicate literal removal 2347]
[SA] new: 2391. $less(X35,sK0) | $less(sK0,X35) | sK0 = X35 | 0 = sK0 | $less(0,alength) (4:15) [duplicate literal removal 2337]
[SA] new: 2392. $less(X33,sK0) | $less(sK0,X33) | sK0 = X33 | 0 = sK0 | sK0 = X34 | $less(X34,alength) | $less(0,X34) (4:21) [duplicate literal removal 2336]
[SA] new: 2393. $less(X31,sK0) | $less(sK0,X31) | sK0 = X31 | 0 = sK0 | alength = X32 | $less(X32,alength) | $less(0,X32) (4:21) [duplicate literal removal 2335]
[SA] new: 2394. $less(X30,sK0) | $less(0,X30) | 0 = X30 | 0 = sK0 | 0 = 0 | $less(0,sK0) (4:18) [duplicate literal removal 2334]
[SA] new: 2395. $less(X29,sK0) | $less(0,X29) | 0 = X29 | 0 = sK0 | 0 = 0 | $less(0,alength) (4:18) [duplicate literal removal 2333]
[SA] new: 2396. $less(X26,sK0) | $less(0,X26) | 0 = X26 | 0 = sK0 | 0 = 0 | sK0 = X27 | $less(X27,sK0) | $less(0,X27) (4:24) [duplicate literal removal 2331]
[SA] new: 2397. $less(X24,sK0) | $less(i(main_end),X24) | i(main_end) = X24 | 0 = sK0 | i(main_end) = 0 | 'Sub'(sK2(0),nl8) (4:22) [evaluation 2329]
[SA] new: 2398. $less(X23,sK0) | $less(i(main_end),X23) | i(main_end) = X23 | 0 = sK0 | i(main_end) = 0 | i(l8(sK2(0))) = 0 (4:24) [evaluation 2328]
[SA] new: 2399. $less(X20,sK0) | $less(alength,X20) | alength = X20 | 0 = sK0 | alength = 0 | alength = i(main_end) | i(l8(sK2(0))) = 0 (4:25) [evaluation 2326]
[SA] new: 2400. $less(X19,sK0) | $less(alength,X19) | alength = X19 | 0 = sK0 | alength = 0 | alength = i(l8(sK2(alength))) | alength = i(main_end) (4:25) [evaluation 2325]
[SA] new: 2401. $less(X13,sK0) | $less(X14,X13) | X13 = X14 | 0 = sK0 | 0 = X14 | sK0 = X14 | $less(0,X14) | $less(0,sK0) (4:24) [duplicate literal removal 2322]
[SA] passive: 2317. $less(X1,sK0) | $less(X2,X1) | X1 = X2 | 0 = sK0 | 0 = X2 | ~$less(X3,0) | $less(X3,X2) (4:21) [resolution 387,14]
[SA] passive: 2318. $less(X4,sK0) | $less(X5,X4) | X4 = X5 | 0 = sK0 | 0 = X5 | X5 = X6 | $less(X6,X5) | $less(0,X6) (4:24) [resolution 387,99]
[SA] passive: 2338. $less(0,X0) | $less(X0,X1) | X0 = X1 | 0 = sK0 | 0 = X0 | alength = X2 | $less(X2,alength) | $less(X1,X2) (4:24) [resolution 387,247]
[SA] passive: 2339. $less(0,X3) | $less(X3,X4) | X3 = X4 | 0 = sK0 | 0 = X3 | sK0 = X5 | $less(X5,alength) | $less(X4,X5) (4:24) [resolution 387,199]
[SA] passive: 2342. $less(0,X9) | $less(X9,X10) | X9 = X10 | 0 = sK0 | 0 = X9 | ~$less(X11,X10) | $less(X11,sK0) (4:21) [resolution 387,14]
[SA] passive: 2353. $less(0,X1) | $less(X2,sK0) | X1 = X2 | 0 = sK0 | 0 = X1 | ~$less(X3,X1) | $less(X3,X2) (4:21) [resolution 387,14]
[SA] passive: 2354. $less(0,X4) | $less(X5,sK0) | X4 = X5 | 0 = sK0 | 0 = X4 | X5 = X6 | $less(X6,X5) | $less(X4,X6) (4:24) [resolution 387,99]
[SA] passive: 2363. $less(0,X21) | $less(i(main_end),sK0) | i(main_end) = X21 | 0 = sK0 | 0 = X21 | $less(X22,0) | i(main_end) = X22 | i(l8(sK2(X22))) = X22 | $less(X21,X22) (4:33) [resolution 387,215]
[SA] active: 1840. ~$less(X2,alength) | alength = i(main_end) | $less(X2,0) (6:10:1) [resolution 1818,14]
[SA] new: 2402. alength = i(main_end) | $less(X0,0) | $less(sK0,X0) | sK0 = X0 (7:13) [resolution 1840,106]
[SA] new: 2403. alength = i(main_end) | $less(X1,0) | $less(0,X1) | alength = X1 | 0 = sK0 (7:16) [resolution 1840,178]
[SA] new: 2404. alength = i(main_end) | $less(X2,0) | $less(sK0,X2) | alength = X2 (7:13) [resolution 1840,179]
[SA] new: 2405. alength = i(main_end) | $less(X3,0) | $less(sK0,X4) | $less(X4,X3) | X3 = X4 | sK0 = X3 (7:19) [resolution 1840,198]
[SA] new: 2406. alength = i(main_end) | $less(X5,0) | $less(sK0,X6) | $less(X6,X5) | alength = X5 | sK0 = X6 (7:19) [resolution 1840,204]
[SA] new: 2407. alength = i(main_end) | $less(X7,0) | $less(sK0,X8) | $less(X8,X7) | X7 = X8 | alength = X7 (7:19) [resolution 1840,246]
[SA] new: 2408. alength = i(main_end) | $less(X9,0) | $less(sK0,X10) | $less(X10,X9) | alength = X9 | alength = X10 (7:19) [resolution 1840,253]
[SA] new: 2409. alength = i(main_end) | $less(X11,0) | $less(0,X11) | alength = X11 | alength = 0 (7:16) [resolution 1840,261]
[SA] new: 2410. alength = i(main_end) | $less(X12,0) | $less(0,X12) | 0 = sK0 | 0 = X12 (7:16) [resolution 1840,278]
[SA] new: 2411. alength = i(main_end) | $less(X13,0) | $less(0,X13) | 0 = sK0 | sK0 = X13 (7:16) [resolution 1840,330]
[SA] new: 2412. alength = i(main_end) | $less(X14,0) | $less(sK0,X15) | $less(X15,X14) | X14 = X15 | sK0 = X15 (7:19) [resolution 1840,388]
[SA] new: 2413. alength = i(main_end) | $less(X16,0) | $less(sK0,X17) | $less(X17,X16) | sK0 = X17 | sK0 = X16 (7:19) [resolution 1840,408]
[SA] new: 2414. alength = i(main_end) | $less(X18,0) | $less(alength,X18) | alength = X18 (7:13) [resolution 1840,15]
[SA] new: 2415. alength = i(main_end) | $less(X19,0) | $less(alength,X19) | alength = X19 (7:13) [resolution 1840,15]
[SA] new: 2416. alength = i(main_end) | $less(X20,0) | $less(X21,X20) | $less(alength,X21) | alength = X20 | alength = X21 (7:19) [resolution 1840,174]
[SA] new: 2417. alength = i(main_end) | $less(X22,0) | $less(alength,X23) | $less(X23,X22) | X22 = X23 | alength = X22 (7:19) [resolution 1840,174]
[SA] new: 2418. alength = i(main_end) | $less(X24,0) | $less(X25,X24) | $less(alength,X25) | alength = X25 | X24 = X25 (7:19) [resolution 1840,174]
[SA] new: 2419. alength = i(main_end) | $less(X26,0) | $less(sK0,X26) | $less(alength,alength) | alength = X26 | alength = sK0 (7:19) [resolution 1840,198]
[SA] new: 2420. alength = i(main_end) | $less(X27,0) | $less(sK0,X27) | $less(alength,alength) | alength = alength | sK0 = X27 (7:19) [resolution 1840,204]
[SA] new: 2421. alength = i(main_end) | $less(X28,0) | $less(i(main_end),X28) | $less(alength,0) | alength = X28 | alength = i(main_end) | alength = i(l8(sK2(alength))) (7:27) [resolution 1840,214]
[SA] new: 2422. alength = i(main_end) | $less(X29,0) | $less(sK0,X29) | $less(alength,alength) | alength = X29 | alength = alength (7:19) [resolution 1840,246]
[SA] new: 2423. alength = i(main_end) | $less(X30,0) | $less(sK0,X30) | $less(alength,alength) | alength = alength | alength = X30 (7:19) [resolution 1840,253]
[SA] new: 2424. alength = i(main_end) | $less(X31,0) | $less(0,X31) | $less(alength,sK0) | alength = X31 | 0 = sK0 | alength = sK0 (7:22) [resolution 1840,323]
[SA] new: 2425. alength = i(main_end) | $less(X32,0) | $less(0,X32) | $less(alength,sK0) | alength = sK0 | 0 = sK0 | sK0 = X32 (7:22) [resolution 1840,332]
[SA] new: 2426. alength = i(main_end) | $less(X33,0) | $less(0,X33) | $less(alength,sK0) | alength = X33 | 0 = sK0 | 0 = X33 (7:22) [resolution 1840,387]
[SA] new: 2427. alength = i(main_end) | $less(X34,0) | $less(sK0,X34) | $less(alength,alength) | alength = X34 | sK0 = X34 (7:19) [resolution 1840,388]
[SA] new: 2428. alength = i(main_end) | $less(X35,0) | $less(sK0,X35) | $less(alength,alength) | sK0 = X35 | alength = sK0 (7:19) [resolution 1840,408]
[SA] new: 2429. alength = i(main_end) | $less(i(l8(p(nl8))),0) | zero = nl8 (7:13) [resolution 1840,95]
[SA] new: 2430. alength = i(main_end) | $less(i(l8(p(p(nl8)))),0) | zero = p(nl8) | zero = nl8 (7:18) [resolution 1840,314]
[SA] new: 2431. alength = i(main_end) | $less(i(l8(sK2(X36))),0) | $less(i(main_end),X36) | i(main_end) = X36 | $less(X36,0) (7:21) [resolution 1840,195]
[SA] new: 2432. alength = i(main_end) | $less(i(l8(sK2(sK0))),0) | alength = i(main_end) (7:14) [resolution 1840,272]
[SA] new: 2433. alength = i(main_end) | $less(i(l8(sK2(sK0))),0) | i(main_end) = sK0 (7:14) [resolution 1840,228]
[SA] new: 2434. alength = i(main_end) | $less(i(l8(sK2(alength))),0) | alength = i(main_end) | $less(alength,0) (7:17) [resolution 1840,270]
[SA] new: 2435. alength = i(main_end) | $less(i(main_end),0) | i(l8(sK1(X37,sK2(alength)))) = X37 | $less(alength,0) | a(main_end,X37) = a(l11(sK2(alength)),X37) | alength = i(main_end) (7:32) [resolution 1840,225]
[SA] new: 2436. alength = i(main_end) | $less(i(main_end),0) | $less(alength,X38) | $less(X38,0) | alength = X38 | i(main_end) = X38 | i(l8(sK2(X38))) = X38 (7:27) [resolution 1840,214]
[SA] new: 2437. alength = i(main_end) | $less(i(main_end),0) | alength = i(l8(sK2(alength))) | $less(alength,0) | alength = i(main_end) (7:21) [resolution 1840,127]
[SA] new: 2438. alength = i(main_end) | $less(i(l8(p(nl8))),0) | alength = alength | $less(alength,alength) | zero = nl8 (7:19) [resolution 1840,193]
[SA] new: 2439. alength = i(main_end) | $less(0,0) | alength = 0 (7:10) [resolution 1840,244]
[SA] new: 2440. alength = i(main_end) | $less(0,0) | 0 = sK0 (7:10) [resolution 1840,113]
[SA] new: 2441. alength = i(main_end) | $less(0,0) | $less(alength,sK0) | alength = sK0 | 0 = sK0 (7:16) [resolution 1840,177]
[SA] new: 2442. alength = i(main_end) | $less(0,0) | $less(alength,alength) | alength = alength | 0 = sK0 (7:16) [resolution 1840,178]
[SA] new: 2443. alength = i(main_end) | $less(0,0) | $less(alength,sK0) | 0 = sK0 | alength = 0 (7:16) [resolution 1840,229]
[SA] new: 2444. alength = i(main_end) | $less(0,0) | $less(alength,alength) | alength = alength | alength = 0 (7:16) [resolution 1840,261]
[SA] new: 2445. alength = i(main_end) | $less(0,0) | $less(alength,alength) | 0 = sK0 | alength = 0 (7:16) [resolution 1840,278]
[SA] new: 2446. alength = i(main_end) | $less(0,0) | $less(alength,X39) | $less(X39,sK0) | alength = X39 | 0 = sK0 | sK0 = X39 (7:22) [resolution 1840,323]
[SA] new: 2447. alength = i(main_end) | $less(0,0) | $less(alength,alength) | 0 = sK0 | alength = sK0 (7:16) [resolution 1840,330]
[SA] new: 2448. alength = i(main_end) | $less(0,0) | $less(X40,sK0) | $less(alength,X40) | sK0 = X40 | 0 = sK0 | alength = sK0 (7:22) [resolution 1840,332]
[SA] new: 2449. alength = i(main_end) | $less(0,0) | $less(alength,sK0) | alength = 0 | alength = sK0 (7:16) [resolution 1840,380]
[SA] new: 2450. alength = i(main_end) | $less(0,0) | $less(X41,sK0) | $less(alength,X41) | alength = X41 | 0 = sK0 | alength = 0 (7:22) [resolution 1840,387]
[SA] new: 2451. alength = i(main_end) | $less(sK0,0) (7:7) [resolution 1840,45]
[SA] new: 2452. alength = i(main_end) | $less(sK0,0) | $less(alength,alength) | alength = sK0 (7:13) [resolution 1840,106]
[SA] new: 2453. alength = i(main_end) | $less(sK0,0) | $less(alength,alength) | alength = alength (7:13) [resolution 1840,179]
[SA] new: 2454. alength = i(main_end) | $less(sK0,0) | $less(alength,X42) | $less(X42,alength) | alength = X42 | sK0 = X42 (7:19) [resolution 1840,198]
[SA] new: 2455. alength = i(main_end) | $less(sK0,0) | $less(X43,alength) | $less(alength,X43) | alength = X43 | alength = sK0 (7:19) [resolution 1840,204]
[SA] new: 2456. alength = i(main_end) | $less(sK0,0) | $less(alength,X44) | $less(X44,alength) | alength = X44 | alength = X44 (7:19) [resolution 1840,246]
[SA] new: 2457. alength = i(main_end) | $less(sK0,0) | $less(X45,alength) | $less(alength,X45) | alength = X45 | alength = alength (7:19) [resolution 1840,253]
[SA] new: 2458. alength = i(main_end) | $less(sK0,0) | $less(X46,alength) | $less(alength,X46) | alength = X46 | alength = sK0 (7:19) [resolution 1840,388]
[SA] new: 2459. alength = i(main_end) | $less(sK0,0) | $less(X47,alength) | $less(alength,X47) | alength = sK0 | sK0 = X47 (7:19) [resolution 1840,408]
[SA] new: 2460. alength = i(main_end) | $less(sK0,0) | $less(alength,X44) | $less(X44,alength) | alength = X44 (7:16) [duplicate literal removal 2456]
[SA] new: 2461. alength = i(main_end) | $less(X41,sK0) | $less(alength,X41) | alength = X41 | 0 = sK0 | alength = 0 (7:19) [evaluation 2450]
[SA] new: 2462. alength = i(main_end) | $less(alength,sK0) | alength = 0 | alength = sK0 (7:13) [evaluation 2449]
[SA] new: 2463. alength = i(main_end) | $less(X40,sK0) | $less(alength,X40) | sK0 = X40 | 0 = sK0 | alength = sK0 (7:19) [evaluation 2448]
[SA] new: 2464. alength = i(main_end) | $less(alength,alength) | 0 = sK0 | alength = sK0 (7:13) [evaluation 2447]
[SA] new: 2465. alength = i(main_end) | $less(alength,X39) | $less(X39,sK0) | alength = X39 | 0 = sK0 | sK0 = X39 (7:19) [evaluation 2446]
[SA] new: 2466. alength = i(main_end) | $less(alength,alength) | 0 = sK0 | alength = 0 (7:13) [evaluation 2445]
[SA] new: 2467. alength = i(main_end) | $less(alength,sK0) | 0 = sK0 | alength = 0 (7:13) [evaluation 2443]
[SA] new: 2468. alength = i(main_end) | $less(alength,sK0) | alength = sK0 | 0 = sK0 (7:13) [evaluation 2441]
[SA] new: 2469. alength = i(main_end) | 0 = sK0 (7:7) [evaluation 2440]
[SA] new: 2470. alength = i(main_end) | alength = 0 (7:7) [evaluation 2439]
[SA] new: 2471. alength = i(main_end) | $less(i(main_end),0) | alength = i(l8(sK2(alength))) | $less(alength,0) (7:17) [duplicate literal removal 2437]
[SA] new: 2472. alength = i(main_end) | $less(i(main_end),0) | i(l8(sK1(X37,sK2(alength)))) = X37 | $less(alength,0) | a(main_end,X37) = a(l11(sK2(alength)),X37) (7:28) [duplicate literal removal 2435]
[SA] new: 2473. alength = i(main_end) | $less(i(l8(sK2(alength))),0) | $less(alength,0) (7:13) [duplicate literal removal 2434]
[SA] new: 2474. alength = i(main_end) | $less(i(l8(sK2(sK0))),0) (7:10) [duplicate literal removal 2432]
[SA] new: 2475. alength = i(main_end) | $less(X28,0) | $less(i(main_end),X28) | $less(alength,0) | alength = X28 | alength = i(l8(sK2(alength))) (7:23) [duplicate literal removal 2421]
[SA] passive: 2402. alength = i(main_end) | $less(X0,0) | $less(sK0,X0) | sK0 = X0 (7:13) [resolution 1840,106]
[SA] passive: 2404. alength = i(main_end) | $less(X2,0) | $less(sK0,X2) | alength = X2 (7:13) [resolution 1840,179]
[SA] passive: 2405. alength = i(main_end) | $less(X3,0) | $less(sK0,X4) | $less(X4,X3) | X3 = X4 | sK0 = X3 (7:19) [resolution 1840,198]
[SA] passive: 2406. alength = i(main_end) | $less(X5,0) | $less(sK0,X6) | $less(X6,X5) | alength = X5 | sK0 = X6 (7:19) [resolution 1840,204]
[SA] passive: 2407. alength = i(main_end) | $less(X7,0) | $less(sK0,X8) | $less(X8,X7) | X7 = X8 | alength = X7 (7:19) [resolution 1840,246]
[SA] passive: 2408. alength = i(main_end) | $less(X9,0) | $less(sK0,X10) | $less(X10,X9) | alength = X9 | alength = X10 (7:19) [resolution 1840,253]
[SA] passive: 2412. alength = i(main_end) | $less(X14,0) | $less(sK0,X15) | $less(X15,X14) | X14 = X15 | sK0 = X15 (7:19) [resolution 1840,388]
[SA] passive: 2413. alength = i(main_end) | $less(X16,0) | $less(sK0,X17) | $less(X17,X16) | sK0 = X17 | sK0 = X16 (7:19) [resolution 1840,408]
[SA] passive: 2414. alength = i(main_end) | $less(X18,0) | $less(alength,X18) | alength = X18 (7:13) [resolution 1840,15]
[SA] passive: 2416. alength = i(main_end) | $less(X20,0) | $less(X21,X20) | $less(alength,X21) | alength = X20 | alength = X21 (7:19) [resolution 1840,174]
[SA] passive: 2417. alength = i(main_end) | $less(X22,0) | $less(alength,X23) | $less(X23,X22) | X22 = X23 | alength = X22 (7:19) [resolution 1840,174]
[SA] passive: 2418. alength = i(main_end) | $less(X24,0) | $less(X25,X24) | $less(alength,X25) | alength = X25 | X24 = X25 (7:19) [resolution 1840,174]
[SA] passive: 2429. alength = i(main_end) | $less(i(l8(p(nl8))),0) | zero = nl8 (7:13) [resolution 1840,95]
[SA] passive: 2430. alength = i(main_end) | $less(i(l8(p(p(nl8)))),0) | zero = p(nl8) | zero = nl8 (7:18) [resolution 1840,314]
[SA] passive: 2431. alength = i(main_end) | $less(i(l8(sK2(X36))),0) | $less(i(main_end),X36) | i(main_end) = X36 | $less(X36,0) (7:21) [resolution 1840,195]
[SA] passive: 2474. alength = i(main_end) | $less(i(l8(sK2(sK0))),0) (7:10) [duplicate literal removal 2432]
[SA] new: 2476. alength = i(main_end) (7:4) [subsumption resolution 2451,46]
[SA] new: 2477. alength = i(l8(nl8)) (0:5) [backward demodulation 2476,54]
[SA] new: 2478. ~$less(alength,alength) (0:3) [backward demodulation 2476,82]
[SA] new: 2479. ~$less(X1,alength) | $less(X1,0) | 'Sub'(sK2(X1),nl8) (0:10) [backward demodulation 2476,83]
[SA] new: 2480. ~$less(X1,alength) | $less(X1,0) | i(l8(sK2(X1))) = X1 (0:12) [backward demodulation 2476,84]
[SA] new: 2481. v = a(l11(nl8),alength) (1:6) [backward demodulation 2476,88]
[SA] new: 2482. $less(alength,alength) | alength = i(main_end) (1:7) [backward demodulation 2476,100]
[SA] new: 2483. alength != i(l8(X1)) | nl8 = X1 (1:8) [backward demodulation 2476,109]
[SA] new: 2484. i(l8(s(nl8))) = $sum(alength,1) (1:8) [backward demodulation 2476,118]
[SA] new: 2485. $less(alength,X0) | 'Sub'(sK2(X0),nl8) | $less(X0,0) | i(main_end) = X0 (1:14) [backward demodulation 2476,125]
[SA] new: 2486. $less(alength,X0) | i(l8(sK2(X0))) = X0 | $less(X0,0) | i(main_end) = X0 (1:16) [backward demodulation 2476,127]
[SA] new: 2487. alength != 2 | s(s(zero)) = nl8 (3:8) [backward demodulation 2476,154]
[SA] new: 2488. alength != $sum(alength,1) | s(nl8) = nl8 (2:9) [backward demodulation 2476,155]
[SA] new: 2489. i(l8(s(s(nl8)))) = $sum($sum(alength,1),1) (2:11) [backward demodulation 2476,157]
[SA] new: 2490. i(l8(X0)) != $sum(alength,1) | s(nl8) = X0 (2:11) [backward demodulation 2476,158]
[SA] new: 2491. v = a(l11(s(nl8)),$sum(alength,1)) (2:9) [backward demodulation 2476,160]
[SA] new: 2492. alength != 3 | s(s(s(zero))) = nl8 (4:9) [backward demodulation 2476,180]
[SA] new: 2493. $less(alength,X0) | $less(i(l8(sK2(X0))),alength) | i(main_end) = X0 | $less(X0,0) (2:16) [backward demodulation 2476,195]
[SA] new: 2494. alength = sK0 | $less(sK0,i(main_end)) (3:7) [backward demodulation 2476,206]
[SA] new: 2495. alength = sK0 | i(l8(sK2(sK0))) = sK0 (3:9) [backward demodulation 2476,209]
[SA] new: 2496. alength = sK0 | 'Sub'(sK2(sK0),nl8) (3:7) [backward demodulation 2476,211]
[SA] new: 2497. alength = X0 | $less(i(main_end),X1) | $less(X1,X0) | $less(X0,0) | X0 = X1 | i(l8(sK2(X0))) = X0 (2:22) [backward demodulation 2476,214]
[SA] new: 2498. alength = X2 | ~$less(X3,i(main_end)) | $less(X2,0) | i(l8(sK2(X2))) = X2 | $less(X3,X2) (2:19) [backward demodulation 2476,215]
[SA] new: 2499. alength = X0 | i(main_end) = sK0 | $less(X0,i(main_end)) | $less(sK0,X0) (4:14) [backward demodulation 2476,223]
[SA] new: 2500. alength = sK0 | $less(X1,i(main_end)) | ~$less(X1,sK0) (4:10) [backward demodulation 2476,224]
[SA] new: 2501. $less(alength,X1) | i(l8(sK1(X0,sK2(X1)))) = X0 | $less(X1,0) | a(main_end,X0) = a(l11(sK2(X1)),X0) | i(main_end) = X1 (2:27) [backward demodulation 2476,225]
[SA] new: 2502. alength = sK0 | a(main_end,X0) = a(l11(sK2(sK0)),X0) | i(l8(sK1(X0,sK2(sK0)))) = X0 (4:20) [backward demodulation 2476,227]
[SA] new: 2503. alength = sK0 | $less(i(l8(sK2(sK0))),alength) (4:9) [backward demodulation 2476,228]
[SA] new: 2504. $less(alength,sK0) | 0 = sK0 | i(l8(sK2(0))) = 0 | i(main_end) = 0 (3:16) [backward demodulation 2476,234]
[SA] new: 2505. alength = 0 | 0 = sK0 | i(l8(sK2(0))) = 0 (4:12) [backward demodulation 2476,285]
[SA] new: 2506. alength != 4 | s(s(s(s(zero)))) = nl8 (5:10) [backward demodulation 2476,306]
[SA] new: 2507. alength = sK0 | 'Sub'(sK2(0),nl8) | 0 = sK0 | $less(i(main_end),sK0) (3:14) [backward demodulation 2476,336]
[SA] new: 2508. alength = sK0 | $less(i(main_end),sK0) | 0 = sK0 | i(l8(sK2(0))) = 0 (3:16) [backward demodulation 2476,337]
[SA] new: 2509. v = a(l8(s(nl8)),alength) (4:7) [backward demodulation 2476,354]
[SA] new: 2510. v = a(l8(s(s(nl8))),$sum(alength,1)) (4:10) [backward demodulation 2476,359]
[SA] new: 2511. $less(alength,X11) | $less(X11,X12) | i(main_end) = X11 | X11 = X12 | $less(X12,0) | i(l8(sK2(X12))) = X12 (3:22) [backward demodulation 2476,384]
[SA] new: 2512. $less(alength,X13) | $less(X13,X14) | i(main_end) = X13 | X13 = X14 | $less(X14,0) | 'Sub'(sK2(X14),nl8) (3:20) [backward demodulation 2476,385]
[SA] new: 2513. $less(X21,alength) | $less(alength,X21) | alength = X21 | i(main_end) = X21 (3:13) [backward demodulation 2476,389]
[SA] new: 2514. $less(alength,X13) | $less(X13,X14) | X13 = X14 | i(main_end) = X14 | $less(X14,0) | 'Sub'(sK2(X14),nl8) (3:20) [backward demodulation 2476,395]
[SA] new: 2515. $less(alength,X11) | $less(X11,X12) | i(main_end) = X12 | i(main_end) = X11 | $less(X12,0) | i(l8(sK2(X12))) = X12 (3:23) [backward demodulation 2476,404]
[SA] new: 2516. $less(alength,X13) | $less(X13,X14) | i(main_end) = X14 | i(main_end) = X13 | $less(X14,0) | 'Sub'(sK2(X14),nl8) (3:21) [backward demodulation 2476,405]
[SA] new: 2517. alength = 0 | alength = 0 | i(l8(sK2(0))) = 0 (5:12) [backward demodulation 2476,431]
[SA] new: 2518. $less(alength,X0) | i(main_end) = X0 | $less(X0,0) | alength = X1 | $less(X1,alength) | $less(i(l8(sK2(X0))),X1) (3:22) [backward demodulation 2476,432]
[SA] new: 2519. alength = X2 | $less(i(main_end),X2) | $less(X2,0) | ~$less(X3,i(l8(sK2(X2)))) | $less(X3,alength) (3:19) [backward demodulation 2476,433]
[SA] new: 2520. alength = X6 | $less(X4,0) | i(main_end) = X4 | i(l8(sK2(X4))) = X4 | $less(X5,X4) | $less(X6,X5) | $less(i(main_end),X6) | i(main_end) = X5 (3:30) [backward demodulation 2476,453]
[SA] new: 2521. alength = X7 | $less(X7,0) | i(l8(sK2(X7))) = X7 | $less(X8,X7) | $less(i(main_end),X9) | $less(X9,X8) | X8 = X9 | i(main_end) = X8 (3:29) [backward demodulation 2476,454]
[SA] new: 2522. alength = X12 | $less(X10,0) | i(main_end) = X10 | i(l8(sK2(X10))) = X10 | $less(X11,X10) | $less(X12,X11) | $less(i(main_end),X12) | X11 = X12 (3:29) [backward demodulation 2476,455]
[SA] new: 2523. alength = sK0 | $less(X19,0) | i(main_end) = X19 | i(l8(sK2(X19))) = X19 | $less(0,X19) | $less(i(main_end),sK0) | 0 = sK0 (3:26) [backward demodulation 2476,460]
[SA] new: 2524. alength = sK0 | $less(X21,0) | i(main_end) = X21 | i(l8(sK2(X21))) = X21 | $less(sK0,X21) (4:19) [backward demodulation 2476,462]
[SA] new: 2525. alength = sK0 | i(l8(X0)) != sK0 | sK2(sK0) = X0 (4:12) [backward demodulation 2476,466]
[SA] new: 2526. alength = sK0 | v = a(l11(sK2(sK0)),sK0) (4:10) [backward demodulation 2476,468]
[SA] new: 2527. alength = sK0 | i(l8(s(sK2(sK0)))) = $sum(sK0,1) (4:12) [backward demodulation 2476,469]
[SA] new: 2528. alength = sK0 | 0 != sK0 | zero = sK2(sK0) (4:10) [backward demodulation 2476,470]
[SA] new: 2529. alength = sK0 | 1 != sK0 | s(zero) = sK2(sK0) (4:11) [backward demodulation 2476,472]
[SA] new: 2530. alength = sK0 | sK0 != 2 | s(s(zero)) = sK2(sK0) (4:12) [backward demodulation 2476,473]
[SA] new: 2531. $sum(alength,1) != sK0 | s(nl8) = sK2(sK0) | i(main_end) = sK0 (4:14) [backward demodulation 2476,474]
[SA] new: 2532. alength = sK0 | v = a(l8(s(sK2(sK0))),sK0) (4:11) [backward demodulation 2476,475]
[SA] new: 2533. alength != X7 | nl8 = sK1(X7,p(nl8)) | a(main_end,X7) = a(l11(p(nl8)),X7) | zero = nl8 (3:21) [backward demodulation 2476,482]
[SA] new: 2534. $sum(alength,1) != X10 | s(nl8) = sK1(X10,p(nl8)) | a(main_end,X10) = a(l11(p(nl8)),X10) | zero = nl8 (3:24) [backward demodulation 2476,485]
[SA] new: 2535. $sum(alength,1) != sK0 | s(nl8) = sK2(sK0) | alength = i(main_end) (4:14) [backward demodulation 2476,496]
[SA] new: 2536. alength = X0 | $less(alength,X0) | $less(X0,0) | alength = X0 | i(l8(sK2(X0))) = X0 (3:18) [backward demodulation 2476,498]
[SA] new: 2537. alength = X3 | $less(X2,X3) | $less(X3,0) | X2 = X3 | i(l8(sK2(X3))) = X3 | X2 = X4 | $less(X4,X2) | $less(i(main_end),X4) (3:28) [backward demodulation 2476,500]
[SA] new: 2538. alength = X6 | $less(X5,X6) | $less(X6,0) | X5 = X6 | i(l8(sK2(X6))) = X6 | ~$less(X7,i(main_end)) | $less(X7,X5) (3:25) [backward demodulation 2476,501]
[SA] new: 2539. $less(alength,X1) | $less(X2,0) | X1 = X2 | i(main_end) = X2 | i(l8(sK2(X2))) = X2 | X2 = X3 | $less(X3,X2) | $less(X1,X3) (3:28) [backward demodulation 2476,511]
[SA] new: 2540. alength = X5 | $less(i(main_end),X4) | $less(X5,0) | X4 = X5 | i(l8(sK2(X5))) = X5 | ~$less(X6,X4) | $less(X6,X5) (3:25) [backward demodulation 2476,512]
[SA] new: 2541. $less(alength,X1) | $less(X1,X2) | X1 = X2 | i(main_end) = X2 | i(l8(sK2(X2))) = X2 | alength = 0 | $less(X2,alength) (5:25) [backward demodulation 2476,522]
[SA] new: 2542. alength = X4 | $less(i(main_end),X3) | $less(X3,X4) | X3 = X4 | i(l8(sK2(X4))) = X4 | 0 = sK0 | $less(X4,alength) (4:25) [backward demodulation 2476,523]
[SA] new: 2543. alength = X6 | $less(i(main_end),X5) | $less(X5,X6) | X5 = X6 | i(l8(sK2(X6))) = X6 | 0 = sK0 | $less(X6,sK0) (3:25) [backward demodulation 2476,524]
[SA] new: 2544. alength = 0 | $less(i(main_end),X7) | $less(X7,0) | 0 = X7 | i(l8(sK2(0))) = 0 (3:19) [backward demodulation 2476,525]
[SA] new: 2545. $less(alength,X8) | $less(X8,X9) | X8 = X9 | i(main_end) = X9 | i(l8(sK2(X9))) = X9 | 0 = X10 | $less(X10,0) | $less(X9,X10) (3:28) [backward demodulation 2476,526]
[SA] new: 2546. $less(alength,X11) | $less(X11,X12) | X11 = X12 | i(main_end) = X12 | i(l8(sK2(X12))) = X12 | ~$less(X13,X12) | $less(X13,0) (3:25) [backward demodulation 2476,527]
[SA] new: 2547. alength = X15 | $less(sK0,X15) | $less(X15,0) | sK0 = X15 | i(l8(sK2(X15))) = X15 (3:18) [backward demodulation 2476,541]
[SA] new: 2548. alength = X3 | i(l8(sK1(X2,sK2(X3)))) = X2 | $less(X3,0) | a(main_end,X2) = a(l11(sK2(X3)),X2) | X3 = X4 | $less(X4,X3) | $less(i(main_end),X4) (3:33) [backward demodulation 2476,544]
[SA] new: 2549. alength = X6 | i(l8(sK1(X5,sK2(X6)))) = X5 | $less(X6,0) | a(main_end,X5) = a(l11(sK2(X6)),X5) | ~$less(X7,i(main_end)) | $less(X7,X6) (3:30) [backward demodulation 2476,545]
[SA] new: 2550. alength = 0 | i(l8(sK1(X12,sK2(0)))) = X12 | a(main_end,X12) = a(l11(sK2(0)),X12) | alength = 0 (5:23) [backward demodulation 2476,558]
[SA] new: 2551. alength = 0 | i(l8(sK1(X13,sK2(0)))) = X13 | a(main_end,X13) = a(l11(sK2(0)),X13) | 0 = sK0 (4:23) [backward demodulation 2476,559]
[SA] new: 2552. alength = sK0 | alength = X0 | $less(X0,alength) | $less(i(l8(sK2(sK0))),X0) (5:15) [backward demodulation 2476,560]
[SA] new: 2553. alength = sK0 | ~$less(X1,i(l8(sK2(sK0)))) | $less(X1,alength) (5:12) [backward demodulation 2476,561]
[SA] new: 2554. alength = i(l8(p(nl8))) | zero = nl8 | $less(X0,alength) | $less(i(main_end),X0) | $less(i(l8(p(nl8))),0) | i(l8(p(nl8))) = X0 | i(l8(p(nl8))) = i(l8(sK2(i(l8(p(nl8)))))) (4:40) [backward demodulation 2476,594]
[SA] new: 2555. alength = i(l8(p(nl8))) | zero = nl8 | i(l8(sK1(X9,sK2(i(l8(p(nl8))))))) = X9 | $less(i(l8(p(nl8))),0) | a(main_end,X9) = a(l11(sK2(i(l8(p(nl8))))),X9) (4:38) [backward demodulation 2476,607]
[SA] new: 2556. alength = X10 | zero = nl8 | $less(i(l8(p(nl8))),X10) | $less(X10,0) | i(l8(p(nl8))) = X10 | i(l8(sK2(X10))) = X10 (4:27) [backward demodulation 2476,608]
[SA] new: 2557. alength = i(l8(p(nl8))) | zero = nl8 | i(l8(p(nl8))) = i(l8(sK2(i(l8(p(nl8)))))) | $less(i(l8(p(nl8))),0) (4:27) [backward demodulation 2476,609]
[SA] new: 2558. alength != 5 | s(s(s(s(s(zero))))) = nl8 (6:11) [backward demodulation 2476,632]
[SA] new: 2559. alength = X24 | sK0 = X24 | $less(sK0,X24) | i(l8(sK1(X25,sK2(X24)))) = X25 | $less(X24,0) | a(main_end,X25) = a(l11(sK2(X24)),X25) (4:29) [backward demodulation 2476,677]
[SA] new: 2560. alength = X27 | sK0 = X26 | $less(sK0,X26) | $less(X26,X27) | $less(X27,0) | X26 = X27 | i(l8(sK2(X27))) = X27 (4:24) [backward demodulation 2476,678]
[SA] new: 2561. alength = sK0 | sK0 != 3 | s(s(s(zero))) = sK2(sK0) (5:13) [backward demodulation 2476,689]
[SA] new: 2562. $less(alength,sK0) | 0 = sK0 | i(main_end) = 0 | $less(X5,0) | i(main_end) = X5 | i(l8(sK2(X5))) = X5 | $less(0,X5) (4:26) [backward demodulation 2476,697]
[SA] new: 2563. $less(alength,sK0) | 0 = sK0 | i(main_end) = 0 | 'Sub'(sK2(0),nl8) (4:14) [backward demodulation 2476,717]
[SA] new: 2564. alength = X1 | ~$less(X0,sK0) | i(main_end) = sK0 | $less(X1,0) | i(l8(sK2(X1))) = X1 | $less(X0,X1) (5:22) [backward demodulation 2476,719]
[SA] new: 2565. alength = sK0 | ~$less(X2,sK0) | $less(X2,0) | i(l8(sK2(X2))) = X2 (5:15) [backward demodulation 2476,720]
[SA] new: 2566. alength = sK0 | ~$less(X3,sK0) | $less(X3,0) | 'Sub'(sK2(X3),nl8) (5:13) [backward demodulation 2476,721]
[SA] new: 2567. alength = sK0 | ~$less(i(main_end),sK0) (5:7) [backward demodulation 2476,722]
[SA] new: 2568. alength = X6 | ~$less(X5,sK0) | i(main_end) = sK0 | $less(X6,i(main_end)) | $less(X5,X6) (5:17) [backward demodulation 2476,724]
[SA] new: 2569. alength = sK0 | ~$less(X7,sK0) | ~$less(X8,X7) | $less(X8,i(main_end)) (5:13) [backward demodulation 2476,725]
[SA] new: 2570. alength = sK0 | $less(0,i(main_end)) | 0 = sK0 | i(main_end) = 0 (6:14) [backward demodulation 2476,748]
[SA] new: 2571. alength = sK0 | $less(X2,i(main_end)) | $less(sK0,X2) | sK0 = X2 (6:13) [backward demodulation 2476,759]
[SA] new: 2572. alength = sK0 | $less(0,i(main_end)) | 0 = sK0 (6:10) [backward demodulation 2476,761]
[SA] new: 2573. alength = X0 | alength = X0 | $less(sK0,X0) | $less(X1,alength) | $less(i(main_end),X1) | $less(X0,0) | X0 = X1 | i(l8(sK2(X0))) = X0 (4:28) [backward demodulation 2476,764]
[SA] new: 2574. alength = X26 | alength = X26 | $less(sK0,X26) | i(l8(sK1(X27,sK2(X26)))) = X27 | $less(X26,0) | a(main_end,X27) = a(l11(sK2(X26)),X27) (4:29) [backward demodulation 2476,806]
[SA] new: 2575. alength = X29 | alength = X28 | $less(sK0,X28) | $less(X28,X29) | $less(X29,0) | X28 = X29 | i(l8(sK2(X29))) = X29 (4:24) [backward demodulation 2476,807]
[SA] new: 2576. alength = X30 | alength = X30 | $less(sK0,X30) | i(l8(sK2(X30))) = X30 | $less(X30,0) (4:18) [backward demodulation 2476,808]
[SA] new: 2577. alength = X6 | 0 = sK0 | i(main_end) = sK0 | $less(X6,0) | i(l8(sK2(X6))) = X6 | $less(0,X6) (4:22) [backward demodulation 2476,846]
[SA] new: 2578. alength = sK0 | 0 = sK0 | i(l8(sK2(0))) = 0 (4:12) [backward demodulation 2476,847]
[SA] new: 2579. alength = sK0 | 'Sub'(sK2(0),nl8) | 0 = sK0 (4:10) [backward demodulation 2476,848]
[SA] new: 2580. alength = sK0 | $less(0,i(main_end)) | i(main_end) = 0 (6:11) [backward demodulation 2476,886]
[SA] new: 2581. alength = X6 | i(main_end) = 0 | i(main_end) = sK0 | $less(X6,0) | i(l8(sK2(X6))) = X6 | $less(0,X6) (4:23) [backward demodulation 2476,892]
[SA] new: 2582. alength = sK0 | i(main_end) = 0 | i(l8(sK2(0))) = 0 (4:13) [backward demodulation 2476,893]
[SA] new: 2583. alength = sK0 | i(main_end) = 0 | 'Sub'(sK2(0),nl8) (4:11) [backward demodulation 2476,894]
[SA] new: 2584. alength = X0 | $less(alength,X0) | alength = X0 | $less(X1,0) | i(main_end) = X1 | i(l8(sK2(X1))) = X1 | $less(X0,X1) (4:25) [backward demodulation 2476,964]
[SA] new: 2585. alength = X3 | $less(alength,X3) | alength = X3 | $less(X3,0) | 'Sub'(sK2(X3),nl8) (4:16) [backward demodulation 2476,966]
[SA] new: 2586. alength = X7 | $less(alength,X7) | alength = X7 | i(main_end) = X8 | $less(X8,i(main_end)) | $less(X7,X8) (4:20) [backward demodulation 2476,971]
[SA] new: 2587. alength = X9 | $less(alength,X9) | alength = X9 | ~$less(X10,X9) | $less(X10,i(main_end)) (4:16) [backward demodulation 2476,972]
[SA] new: 2588. alength = sK0 | 0 = sK0 | a(main_end,X0) = a(l11(sK2(0)),X0) | i(l8(sK1(X0,sK2(0)))) = X0 (5:23) [backward demodulation 2476,1021]
[SA] new: 2589. alength = sK0 | 0 = sK0 | $less(i(l8(sK2(0))),alength) (5:12) [backward demodulation 2476,1022]
[SA] new: 2590. alength = X14 | $less(i(main_end),X14) | $less(X14,alength) | sK0 = X14 | $less(X15,0) | i(main_end) = X15 | i(l8(sK2(X15))) = X15 | $less(sK0,X15) (4:29) [backward demodulation 2476,1030]
[SA] new: 2591. $less(alength,X16) | $less(X16,alength) | i(main_end) = X16 | sK0 = X16 | i(l8(sK2(sK0))) = sK0 (4:19) [backward demodulation 2476,1080]
[SA] new: 2592. alength = X17 | $less(i(main_end),X17) | $less(X17,alength) | sK0 = X17 | 'Sub'(sK2(sK0),nl8) (4:17) [backward demodulation 2476,1081]
[SA] new: 2593. alength = sK0 | $less(sK0,X13) | i(main_end) = X13 | $less(X14,0) | i(main_end) = X14 | i(l8(sK2(X14))) = X14 | $less(X13,X14) (4:26) [backward demodulation 2476,1082]
[SA] new: 2594. alength = sK0 | $less(sK0,X16) | i(main_end) = X16 | $less(X16,0) | 'Sub'(sK2(X16),nl8) (4:17) [backward demodulation 2476,1083]
[SA] new: 2595. alength = X14 | $less(i(main_end),X14) | $less(X14,alength) | alength = X14 | $less(X15,0) | i(main_end) = X15 | i(l8(sK2(X15))) = X15 | $less(sK0,X15) (4:29) [backward demodulation 2476,1150]
[SA] new: 2596. $less(alength,X16) | $less(X16,alength) | i(main_end) = X16 | alength = X16 | i(l8(sK2(sK0))) = sK0 (4:19) [backward demodulation 2476,1195]
[SA] new: 2597. alength = X17 | $less(i(main_end),X17) | $less(X17,alength) | alength = X17 | 'Sub'(sK2(sK0),nl8) (4:17) [backward demodulation 2476,1196]
[SA] new: 2598. alength != 6 | s(s(s(s(s(s(zero)))))) = nl8 (7:12) [backward demodulation 2476,1264]
[SA] new: 2599. alength = X1 | sK0 = X0 | 0 = sK0 | $less(X0,sK0) | $less(X1,X0) | $less(i(main_end),X2) | $less(X2,X1) | X1 = X2 | i(l8(sK2(X1))) = X1 (4:31) [backward demodulation 2476,1274]
[SA] new: 2600. alength = X1 | i(main_end) = sK0 | 0 = sK0 | $less(X1,i(main_end)) | $less(0,X1) (7:17) [backward demodulation 2476,1318]
[SA] new: 2601. alength = sK0 | 0 = sK0 | ~$less(X2,0) | $less(X2,i(main_end)) (7:13) [backward demodulation 2476,1319]
[SA] new: 2602. alength = X17 | sK0 = X17 | 0 = sK0 | $less(0,X17) | $less(X18,sK0) | $less(i(main_end),X18) | $less(X17,0) | X17 = X18 | i(l8(sK2(X17))) = X17 (4:31) [backward demodulation 2476,1330]
[SA] new: 2603. $less(alength,sK0) | sK0 = X47 | 0 = sK0 | $less(0,X47) | i(l8(sK1(X48,sK2(X47)))) = X48 | $less(X47,0) | a(main_end,X48) = a(l11(sK2(X47)),X48) | i(main_end) = X47 (4:36) [backward demodulation 2476,1359]
[SA] new: 2604. $less(alength,sK0) | sK0 = X49 | 0 = sK0 | $less(0,X49) | $less(X49,X50) | $less(X50,0) | X49 = X50 | i(main_end) = X50 | i(l8(sK2(X50))) = X50 (4:31) [backward demodulation 2476,1360]
[SA] new: 2605. $less(alength,sK0) | sK0 = X51 | 0 = sK0 | $less(0,X51) | i(l8(sK2(X51))) = X51 | $less(X51,0) | i(main_end) = X51 (4:25) [backward demodulation 2476,1361]
[SA] new: 2606. alength = X15 | $less(sK0,X15) | $less(X15,i(main_end)) | sK0 = X15 (4:13) [backward demodulation 2476,1482]
[SA] new: 2607. alength = X15 | $less(sK0,X15) | sK0 = X15 | $less(X16,0) | i(main_end) = X16 | i(l8(sK2(X16))) = X16 | $less(X15,X16) (4:25) [backward demodulation 2476,1521]
[SA] new: 2608. alength = X18 | $less(sK0,X18) | sK0 = X18 | $less(X18,0) | 'Sub'(sK2(X18),nl8) (4:16) [backward demodulation 2476,1522]
[SA] new: 2609. alength = sK0 | $less(sK0,X15) | sK0 = X15 | $less(X16,0) | i(main_end) = X16 | i(l8(sK2(X16))) = X16 | $less(X15,X16) (4:25) [backward demodulation 2476,1591]
[SA] new: 2610. alength = sK0 | $less(sK0,X17) | sK0 = X17 | $less(X17,0) | i(l8(sK2(X17))) = X17 (4:18) [backward demodulation 2476,1592]
[SA] new: 2611. alength = sK0 | $less(sK0,X18) | sK0 = X18 | $less(X18,0) | 'Sub'(sK2(X18),nl8) (4:16) [backward demodulation 2476,1593]
[SA] new: 2612. alength = 0 | $less(0,i(main_end)) | 0 = sK0 (5:10) [backward demodulation 2476,1619]
[SA] new: 2613. alength = 0 | 0 = sK0 | $less(X8,0) | i(main_end) = X8 | i(l8(sK2(X8))) = X8 | $less(0,X8) (5:22) [backward demodulation 2476,1631]
[SA] new: 2614. alength = 0 | 'Sub'(sK2(0),nl8) | 0 = sK0 (5:10) [backward demodulation 2476,1632]
[SA] new: 2615. alength = X1 | 0 = sK0 | i(main_end) = 0 | $less(X1,i(main_end)) | $less(0,X1) (6:17) [backward demodulation 2476,1665]
[SA] new: 2616. alength = 0 | 0 = sK0 | ~$less(X2,0) | $less(X2,i(main_end)) (6:13) [backward demodulation 2476,1666]
[SA] new: 2617. alength = 0 | 0 = sK0 | $less(i(l8(sK2(0))),alength) (6:12) [backward demodulation 2476,1736]
[SA] new: 2618. alength != $sum(alength,1) | s(nl8) = sK2(alength) | alength = i(main_end) (5:14) [backward demodulation 2476,1813]
[SA] new: 2619. alength = X3 | $less(alength,X3) | i(l8(sK2(X3))) = X3 | $less(X3,0) | alength = 0 (7:18) [backward demodulation 2476,1864]
[SA] new: 2620. alength = X8 | $less(alength,X7) | $less(X7,X8) | $less(X8,0) | X7 = X8 | i(l8(sK2(X8))) = X8 | alength = 0 (7:24) [backward demodulation 2476,1877]
[SA] new: 2621. alength = X10 | ~$less(X9,alength) | $less(X10,0) | i(l8(sK2(X10))) = X10 | $less(X9,X10) | alength = 0 (7:21) [backward demodulation 2476,1878]
[SA] new: 2622. alength = X12 | $less(alength,X12) | i(l8(sK1(X13,sK2(X12)))) = X13 | $less(X12,0) | a(main_end,X13) = a(l11(sK2(X12)),X13) | alength = 0 (7:29) [backward demodulation 2476,1880]
[SA] new: 2623. alength = sK0 | ~$less(alength,sK0) | alength = 0 (7:9) [backward demodulation 2476,1891]
[SA] new: 2624. alength = X3 | $less(alength,X3) | i(l8(sK2(X3))) = X3 | $less(X3,0) | 0 = sK0 (7:18) [backward demodulation 2476,1967]
[SA] new: 2625. alength = X8 | $less(alength,X7) | $less(X7,X8) | $less(X8,0) | X7 = X8 | i(l8(sK2(X8))) = X8 | 0 = sK0 (7:24) [backward demodulation 2476,1980]
[SA] new: 2626. alength = X10 | ~$less(X9,alength) | $less(X10,0) | i(l8(sK2(X10))) = X10 | $less(X9,X10) | 0 = sK0 (7:21) [backward demodulation 2476,1981]
[SA] new: 2627. alength = X12 | $less(alength,X12) | i(l8(sK1(X13,sK2(X12)))) = X13 | $less(X12,0) | a(main_end,X13) = a(l11(sK2(X12)),X13) | 0 = sK0 (7:29) [backward demodulation 2476,1983]
[SA] new: 2628. alength = sK0 | ~$less(alength,sK0) | 0 = sK0 (7:9) [backward demodulation 2476,1994]
[SA] new: 2629. alength = X19 | $less(i(main_end),X19) | $less(X19,sK0) | 0 = sK0 | sK0 = X19 | $less(X20,0) | i(main_end) = X20 | i(l8(sK2(X20))) = X20 | $less(0,X20) (4:32) [backward demodulation 2476,2014]
[SA] new: 2630. alength = X22 | $less(i(main_end),X22) | $less(X22,sK0) | 0 = sK0 | sK0 = X22 | 'Sub'(sK2(0),nl8) (4:20) [backward demodulation 2476,2084]
[SA] new: 2631. alength = X21 | $less(i(main_end),X21) | $less(X21,sK0) | 0 = sK0 | sK0 = X21 | i(l8(sK2(0))) = 0 (4:22) [backward demodulation 2476,2085]
[SA] new: 2632. alength = sK0 | $less(0,X19) | i(main_end) = X19 | 0 = sK0 | $less(X20,0) | i(main_end) = X20 | i(l8(sK2(X20))) = X20 | $less(X19,X20) (4:29) [backward demodulation 2476,2089]
[SA] new: 2633. alength = sK0 | $less(0,X22) | i(main_end) = X22 | 0 = sK0 | $less(X22,0) | 'Sub'(sK2(X22),nl8) (4:20) [backward demodulation 2476,2090]
[SA] new: 2634. alength = sK0 | $less(0,X24) | $less(X24,i(main_end)) | 0 = sK0 | sK0 = X24 (6:16) [backward demodulation 2476,2165]
[SA] new: 2635. alength = sK0 | $less(0,X19) | 0 = sK0 | sK0 = X19 | $less(X20,0) | i(main_end) = X20 | i(l8(sK2(X20))) = X20 | $less(X19,X20) (4:28) [backward demodulation 2476,2178]
[SA] new: 2636. alength = sK0 | $less(0,X21) | 0 = sK0 | sK0 = X21 | $less(X21,0) | i(l8(sK2(X21))) = X21 (4:21) [backward demodulation 2476,2179]
[SA] new: 2637. alength = sK0 | $less(0,X22) | 0 = sK0 | sK0 = X22 | $less(X22,0) | 'Sub'(sK2(X22),nl8) (4:19) [backward demodulation 2476,2180]
[SA] new: 2638. $less(i(l8(p(nl8))),alength) | alength = i(main_end) | zero = nl8 (4:13) [backward demodulation 2476,2312]
[SA] new: 2639. $less(i(l8(p(p(nl8)))),alength) | alength = i(main_end) | zero = p(nl8) | zero = nl8 (4:18) [backward demodulation 2476,2313]
[SA] new: 2640. $less(i(l8(sK2(X64))),alength) | alength = i(main_end) | $less(i(main_end),X64) | i(main_end) = X64 | $less(X64,0) (4:21) [backward demodulation 2476,2314]
[SA] new: 2641. $less(i(l8(sK2(sK0))),alength) | alength = i(main_end) (5:10) [backward demodulation 2476,2315]
[SA] new: 2642. $less(alength,sK0) | $less(0,X21) | i(main_end) = X21 | 0 = sK0 | 0 = X21 | $less(X22,0) | i(main_end) = X22 | i(l8(sK2(X22))) = X22 | $less(X21,X22) (4:32) [backward demodulation 2476,2363]
[SA] passive: 2476. alength = i(main_end) (7:4) [subsumption resolution 2451,46]
[SA] new: 2643. alength = X17 | $less(i(main_end),X17) | $less(X17,alength) | 'Sub'(sK2(sK0),nl8) (4:14) [duplicate literal removal 2597]
[SA] new: 2644. alength = X14 | $less(i(main_end),X14) | $less(X14,alength) | $less(X15,0) | i(main_end) = X15 | i(l8(sK2(X15))) = X15 | $less(sK0,X15) (4:26) [duplicate literal removal 2595]
[SA] new: 2645. alength = X9 | $less(alength,X9) | ~$less(X10,X9) | $less(X10,i(main_end)) (4:13) [duplicate literal removal 2587]
[SA] new: 2646. alength = X7 | $less(alength,X7) | i(main_end) = X8 | $less(X8,i(main_end)) | $less(X7,X8) (4:17) [duplicate literal removal 2586]
[SA] new: 2647. alength = X3 | $less(alength,X3) | $less(X3,0) | 'Sub'(sK2(X3),nl8) (4:13) [duplicate literal removal 2585]
[SA] new: 2648. alength = X0 | $less(alength,X0) | $less(X1,0) | i(main_end) = X1 | i(l8(sK2(X1))) = X1 | $less(X0,X1) (4:22) [duplicate literal removal 2584]
[SA] new: 2649. alength = X30 | $less(sK0,X30) | i(l8(sK2(X30))) = X30 | $less(X30,0) (4:15) [duplicate literal removal 2576]
[SA] new: 2650. alength = X26 | $less(sK0,X26) | i(l8(sK1(X27,sK2(X26)))) = X27 | $less(X26,0) | a(main_end,X27) = a(l11(sK2(X26)),X27) (4:26) [duplicate literal removal 2574]
[SA] new: 2651. alength = X0 | $less(sK0,X0) | $less(X1,alength) | $less(i(main_end),X1) | $less(X0,0) | X0 = X1 | i(l8(sK2(X0))) = X0 (4:25) [duplicate literal removal 2573]
[SA] new: 2652. alength = 0 | i(l8(sK1(X12,sK2(0)))) = X12 | a(main_end,X12) = a(l11(sK2(0)),X12) (5:20) [duplicate literal removal 2550]
[SA] new: 2653. alength = X0 | $less(alength,X0) | $less(X0,0) | i(l8(sK2(X0))) = X0 (3:15) [duplicate literal removal 2536]
[SA] new: 2654. alength = 0 | i(l8(sK2(0))) = 0 (5:9) [duplicate literal removal 2517]
[SA] passive: 2477. alength = i(l8(nl8)) (0:5) [backward demodulation 2476,54]
[SA] passive: 2479. ~$less(X1,alength) | $less(X1,0) | 'Sub'(sK2(X1),nl8) (0:10) [backward demodulation 2476,83]
[SA] passive: 2480. ~$less(X1,alength) | $less(X1,0) | i(l8(sK2(X1))) = X1 (0:12) [backward demodulation 2476,84]
[SA] passive: 2481. v = a(l11(nl8),alength) (1:6) [backward demodulation 2476,88]
[SA] passive: 2483. alength != i(l8(X1)) | nl8 = X1 (1:8) [backward demodulation 2476,109]
[SA] passive: 2484. i(l8(s(nl8))) = $sum(alength,1) (1:8) [backward demodulation 2476,118]
[SA] new: 2655. alength = X0 | $less(alength,X0) | 'Sub'(sK2(X0),nl8) | $less(X0,0) (1:13) [forward demodulation 2485,2476]
[SA] passive: 2655. alength = X0 | $less(alength,X0) | 'Sub'(sK2(X0),nl8) | $less(X0,0) (1:13) [forward demodulation 2485,2476]
[SA] new: 2656. alength = X0 | $less(alength,X0) | i(l8(sK2(X0))) = X0 | $less(X0,0) (1:15) [forward demodulation 2486,2476]
[SA] passive: 2656. alength = X0 | $less(alength,X0) | i(l8(sK2(X0))) = X0 | $less(X0,0) (1:15) [forward demodulation 2486,2476]
[SA] passive: 2487. alength != 2 | s(s(zero)) = nl8 (3:8) [backward demodulation 2476,154]
[SA] passive: 2488. alength != $sum(alength,1) | s(nl8) = nl8 (2:9) [backward demodulation 2476,155]
[SA] passive: 2489. i(l8(s(s(nl8)))) = $sum($sum(alength,1),1) (2:11) [backward demodulation 2476,157]
[SA] passive: 2490. i(l8(X0)) != $sum(alength,1) | s(nl8) = X0 (2:11) [backward demodulation 2476,158]
[SA] passive: 2491. v = a(l11(s(nl8)),$sum(alength,1)) (2:9) [backward demodulation 2476,160]
[SA] passive: 2492. alength != 3 | s(s(s(zero))) = nl8 (4:9) [backward demodulation 2476,180]
[SA] new: 2657. alength = X0 | $less(alength,X0) | $less(i(l8(sK2(X0))),alength) | $less(X0,0) (2:15) [forward demodulation 2493,2476]
[SA] passive: 2657. alength = X0 | $less(alength,X0) | $less(i(l8(sK2(X0))),alength) | $less(X0,0) (2:15) [forward demodulation 2493,2476]
[SA] new: 2658. $less(sK0,alength) | alength = sK0 (3:6) [forward demodulation 2494,2476]
[SA] passive: 2495. alength = sK0 | i(l8(sK2(sK0))) = sK0 (3:9) [backward demodulation 2476,209]
[SA] passive: 2496. alength = sK0 | 'Sub'(sK2(sK0),nl8) (3:7) [backward demodulation 2476,211]
[SA] new: 2659. $less(alength,X1) | alength = X0 | $less(X1,X0) | $less(X0,0) | X0 = X1 | i(l8(sK2(X0))) = X0 (2:21) [forward demodulation 2497,2476]
[SA] passive: 2659. $less(alength,X1) | alength = X0 | $less(X1,X0) | $less(X0,0) | X0 = X1 | i(l8(sK2(X0))) = X0 (2:21) [forward demodulation 2497,2476]
[SA] new: 2660. ~$less(X3,alength) | alength = X2 | $less(X2,0) | i(l8(sK2(X2))) = X2 | $less(X3,X2) (2:18) [forward demodulation 2498,2476]
[SA] passive: 2660. ~$less(X3,alength) | alength = X2 | $less(X2,0) | i(l8(sK2(X2))) = X2 | $less(X3,X2) (2:18) [forward demodulation 2498,2476]
[SA] new: 2661. alength = sK0 | alength = X0 | $less(X0,i(main_end)) | $less(sK0,X0) (4:13) [forward demodulation 2499,2476]
[SA] new: 2662. $less(X0,alength) | alength = sK0 | alength = X0 | $less(sK0,X0) (4:12) [forward demodulation 2661,2476]
[SA] new: 2663. $less(X1,alength) | alength = sK0 | ~$less(X1,sK0) (4:9) [forward demodulation 2500,2476]
[SA] new: 2664. alength = X1 | $less(alength,X1) | i(l8(sK1(X0,sK2(X1)))) = X0 | $less(X1,0) | a(main_end,X0) = a(l11(sK2(X1)),X0) (2:26) [forward demodulation 2501,2476]
[SA] passive: 2664. alength = X1 | $less(alength,X1) | i(l8(sK1(X0,sK2(X1)))) = X0 | $less(X1,0) | a(main_end,X0) = a(l11(sK2(X1)),X0) (2:26) [forward demodulation 2501,2476]
[SA] passive: 2502. alength = sK0 | a(main_end,X0) = a(l11(sK2(sK0)),X0) | i(l8(sK1(X0,sK2(sK0)))) = X0 (4:20) [backward demodulation 2476,227]
[SA] passive: 2503. alength = sK0 | $less(i(l8(sK2(sK0))),alength) (4:9) [backward demodulation 2476,228]
[SA] new: 2665. alength = 0 | $less(alength,sK0) | 0 = sK0 | i(l8(sK2(0))) = 0 (3:15) [forward demodulation 2504,2476]
[SA] passive: 2665. alength = 0 | $less(alength,sK0) | 0 = sK0 | i(l8(sK2(0))) = 0 (3:15) [forward demodulation 2504,2476]
[SA] passive: 2505. alength = 0 | 0 = sK0 | i(l8(sK2(0))) = 0 (4:12) [backward demodulation 2476,285]
[SA] passive: 2506. alength != 4 | s(s(s(s(zero)))) = nl8 (5:10) [backward demodulation 2476,306]
[SA] new: 2666. $less(alength,sK0) | alength = sK0 | 'Sub'(sK2(0),nl8) | 0 = sK0 (3:13) [forward demodulation 2507,2476]
[SA] passive: 2666. $less(alength,sK0) | alength = sK0 | 'Sub'(sK2(0),nl8) | 0 = sK0 (3:13) [forward demodulation 2507,2476]
[SA] new: 2667. $less(alength,sK0) | alength = sK0 | 0 = sK0 | i(l8(sK2(0))) = 0 (3:15) [forward demodulation 2508,2476]
[SA] passive: 2667. $less(alength,sK0) | alength = sK0 | 0 = sK0 | i(l8(sK2(0))) = 0 (3:15) [forward demodulation 2508,2476]
[SA] passive: 2509. v = a(l8(s(nl8)),alength) (4:7) [backward demodulation 2476,354]
[SA] passive: 2510. v = a(l8(s(s(nl8))),$sum(alength,1)) (4:10) [backward demodulation 2476,359]
[SA] new: 2668. alength = X11 | $less(alength,X11) | $less(X11,X12) | X11 = X12 | $less(X12,0) | i(l8(sK2(X12))) = X12 (3:21) [forward demodulation 2511,2476]
[SA] passive: 2668. alength = X11 | $less(alength,X11) | $less(X11,X12) | X11 = X12 | $less(X12,0) | i(l8(sK2(X12))) = X12 (3:21) [forward demodulation 2511,2476]
[SA] new: 2669. alength = X13 | $less(alength,X13) | $less(X13,X14) | X13 = X14 | $less(X14,0) | 'Sub'(sK2(X14),nl8) (3:19) [forward demodulation 2512,2476]
[SA] passive: 2669. alength = X13 | $less(alength,X13) | $less(X13,X14) | X13 = X14 | $less(X14,0) | 'Sub'(sK2(X14),nl8) (3:19) [forward demodulation 2512,2476]
[SA] new: 2670. alength = X14 | $less(alength,X13) | $less(X13,X14) | X13 = X14 | $less(X14,0) | 'Sub'(sK2(X14),nl8) (3:19) [forward demodulation 2514,2476]
[SA] passive: 2670. alength = X14 | $less(alength,X13) | $less(X13,X14) | X13 = X14 | $less(X14,0) | 'Sub'(sK2(X14),nl8) (3:19) [forward demodulation 2514,2476]
[SA] new: 2671. alength = X12 | $less(alength,X11) | $less(X11,X12) | i(main_end) = X11 | $less(X12,0) | i(l8(sK2(X12))) = X12 (3:22) [forward demodulation 2515,2476]
[SA] new: 2672. alength = X11 | alength = X12 | $less(alength,X11) | $less(X11,X12) | $less(X12,0) | i(l8(sK2(X12))) = X12 (3:21) [forward demodulation 2671,2476]
[SA] passive: 2672. alength = X11 | alength = X12 | $less(alength,X11) | $less(X11,X12) | $less(X12,0) | i(l8(sK2(X12))) = X12 (3:21) [forward demodulation 2671,2476]
[SA] new: 2673. alength = X14 | $less(alength,X13) | $less(X13,X14) | i(main_end) = X13 | $less(X14,0) | 'Sub'(sK2(X14),nl8) (3:20) [forward demodulation 2516,2476]
[SA] new: 2674. alength = X13 | alength = X14 | $less(alength,X13) | $less(X13,X14) | $less(X14,0) | 'Sub'(sK2(X14),nl8) (3:19) [forward demodulation 2673,2476]
[SA] passive: 2674. alength = X13 | alength = X14 | $less(alength,X13) | $less(X13,X14) | $less(X14,0) | 'Sub'(sK2(X14),nl8) (3:19) [forward demodulation 2673,2476]
[SA] passive: 2654. alength = 0 | i(l8(sK2(0))) = 0 (5:9) [duplicate literal removal 2517]
[SA] new: 2675. alength = X0 | $less(alength,X0) | $less(X0,0) | alength = X1 | $less(X1,alength) | $less(i(l8(sK2(X0))),X1) (3:21) [forward demodulation 2518,2476]
[SA] passive: 2675. alength = X0 | $less(alength,X0) | $less(X0,0) | alength = X1 | $less(X1,alength) | $less(i(l8(sK2(X0))),X1) (3:21) [forward demodulation 2518,2476]
[SA] new: 2676. $less(alength,X2) | alength = X2 | $less(X2,0) | ~$less(X3,i(l8(sK2(X2)))) | $less(X3,alength) (3:18) [forward demodulation 2519,2476]
[SA] passive: 2676. $less(alength,X2) | alength = X2 | $less(X2,0) | ~$less(X3,i(l8(sK2(X2)))) | $less(X3,alength) (3:18) [forward demodulation 2519,2476]
[SA] new: 2677. alength = X4 | alength = X6 | $less(X4,0) | i(l8(sK2(X4))) = X4 | $less(X5,X4) | $less(X6,X5) | $less(i(main_end),X6) | i(main_end) = X5 (3:29) [forward demodulation 2520,2476]
[SA] new: 2678. $less(alength,X6) | alength = X4 | alength = X6 | $less(X4,0) | i(l8(sK2(X4))) = X4 | $less(X5,X4) | $less(X6,X5) | i(main_end) = X5 (3:28) [forward demodulation 2677,2476]
[SA] new: 2679. alength = X5 | $less(alength,X6) | alength = X4 | alength = X6 | $less(X4,0) | i(l8(sK2(X4))) = X4 | $less(X5,X4) | $less(X6,X5) (3:27) [forward demodulation 2678,2476]
[SA] passive: 2679. alength = X5 | $less(alength,X6) | alength = X4 | alength = X6 | $less(X4,0) | i(l8(sK2(X4))) = X4 | $less(X5,X4) | $less(X6,X5) (3:27) [forward demodulation 2678,2476]
[SA] new: 2680. $less(alength,X9) | alength = X7 | $less(X7,0) | i(l8(sK2(X7))) = X7 | $less(X8,X7) | $less(X9,X8) | X8 = X9 | i(main_end) = X8 (3:28) [forward demodulation 2521,2476]
[SA] new: 2681. alength = X8 | $less(alength,X9) | alength = X7 | $less(X7,0) | i(l8(sK2(X7))) = X7 | $less(X8,X7) | $less(X9,X8) | X8 = X9 (3:27) [forward demodulation 2680,2476]
[SA] passive: 2681. alength = X8 | $less(alength,X9) | alength = X7 | $less(X7,0) | i(l8(sK2(X7))) = X7 | $less(X8,X7) | $less(X9,X8) | X8 = X9 (3:27) [forward demodulation 2680,2476]
[SA] new: 2682. alength = X10 | alength = X12 | $less(X10,0) | i(l8(sK2(X10))) = X10 | $less(X11,X10) | $less(X12,X11) | $less(i(main_end),X12) | X11 = X12 (3:28) [forward demodulation 2522,2476]
[SA] new: 2683. $less(alength,X12) | alength = X10 | alength = X12 | $less(X10,0) | i(l8(sK2(X10))) = X10 | $less(X11,X10) | $less(X12,X11) | X11 = X12 (3:27) [forward demodulation 2682,2476]
[SA] passive: 2683. $less(alength,X12) | alength = X10 | alength = X12 | $less(X10,0) | i(l8(sK2(X10))) = X10 | $less(X11,X10) | $less(X12,X11) | X11 = X12 (3:27) [forward demodulation 2682,2476]
[SA] new: 2684. alength = X19 | alength = sK0 | $less(X19,0) | i(l8(sK2(X19))) = X19 | $less(0,X19) | $less(i(main_end),sK0) | 0 = sK0 (3:25) [forward demodulation 2523,2476]
[SA] new: 2685. $less(alength,sK0) | alength = X19 | alength = sK0 | $less(X19,0) | i(l8(sK2(X19))) = X19 | $less(0,X19) | 0 = sK0 (3:24) [forward demodulation 2684,2476]
[SA] passive: 2685. $less(alength,sK0) | alength = X19 | alength = sK0 | $less(X19,0) | i(l8(sK2(X19))) = X19 | $less(0,X19) | 0 = sK0 (3:24) [forward demodulation 2684,2476]
[SA] new: 2686. alength = X21 | alength = sK0 | $less(X21,0) | i(l8(sK2(X21))) = X21 | $less(sK0,X21) (4:18) [forward demodulation 2524,2476]
[SA] passive: 2686. alength = X21 | alength = sK0 | $less(X21,0) | i(l8(sK2(X21))) = X21 | $less(sK0,X21) (4:18) [forward demodulation 2524,2476]
[SA] passive: 2525. alength = sK0 | i(l8(X0)) != sK0 | sK2(sK0) = X0 (4:12) [backward demodulation 2476,466]
[SA] passive: 2526. alength = sK0 | v = a(l11(sK2(sK0)),sK0) (4:10) [backward demodulation 2476,468]
[SA] passive: 2527. alength = sK0 | i(l8(s(sK2(sK0)))) = $sum(sK0,1) (4:12) [backward demodulation 2476,469]
[SA] passive: 2528. alength = sK0 | 0 != sK0 | zero = sK2(sK0) (4:10) [backward demodulation 2476,470]
[SA] passive: 2529. alength = sK0 | 1 != sK0 | s(zero) = sK2(sK0) (4:11) [backward demodulation 2476,472]
[SA] passive: 2530. alength = sK0 | sK0 != 2 | s(s(zero)) = sK2(sK0) (4:12) [backward demodulation 2476,473]
[SA] new: 2687. alength = sK0 | $sum(alength,1) != sK0 | s(nl8) = sK2(sK0) (4:13) [forward demodulation 2531,2476]
[SA] passive: 2687. alength = sK0 | $sum(alength,1) != sK0 | s(nl8) = sK2(sK0) (4:13) [forward demodulation 2531,2476]
[SA] passive: 2532. alength = sK0 | v = a(l8(s(sK2(sK0))),sK0) (4:11) [backward demodulation 2476,475]
[SA] passive: 2533. alength != X7 | nl8 = sK1(X7,p(nl8)) | a(main_end,X7) = a(l11(p(nl8)),X7) | zero = nl8 (3:21) [backward demodulation 2476,482]
[SA] passive: 2534. $sum(alength,1) != X10 | s(nl8) = sK1(X10,p(nl8)) | a(main_end,X10) = a(l11(p(nl8)),X10) | zero = nl8 (3:24) [backward demodulation 2476,485]
[SA] new: 2688. $less(alength,X4) | alength = X3 | $less(X2,X3) | $less(X3,0) | X2 = X3 | i(l8(sK2(X3))) = X3 | X2 = X4 | $less(X4,X2) (3:27) [forward demodulation 2537,2476]
[SA] passive: 2688. $less(alength,X4) | alength = X3 | $less(X2,X3) | $less(X3,0) | X2 = X3 | i(l8(sK2(X3))) = X3 | X2 = X4 | $less(X4,X2) (3:27) [forward demodulation 2537,2476]
[SA] new: 2689. ~$less(X7,alength) | alength = X6 | $less(X5,X6) | $less(X6,0) | X5 = X6 | i(l8(sK2(X6))) = X6 | $less(X7,X5) (3:24) [forward demodulation 2538,2476]
[SA] passive: 2689. ~$less(X7,alength) | alength = X6 | $less(X5,X6) | $less(X6,0) | X5 = X6 | i(l8(sK2(X6))) = X6 | $less(X7,X5) (3:24) [forward demodulation 2538,2476]
[SA] new: 2690. alength = X2 | $less(alength,X1) | $less(X2,0) | X1 = X2 | i(l8(sK2(X2))) = X2 | X2 = X3 | $less(X3,X2) | $less(X1,X3) (3:27) [forward demodulation 2539,2476]
[SA] passive: 2690. alength = X2 | $less(alength,X1) | $less(X2,0) | X1 = X2 | i(l8(sK2(X2))) = X2 | X2 = X3 | $less(X3,X2) | $less(X1,X3) (3:27) [forward demodulation 2539,2476]
[SA] new: 2691. $less(alength,X4) | alength = X5 | $less(X5,0) | X4 = X5 | i(l8(sK2(X5))) = X5 | ~$less(X6,X4) | $less(X6,X5) (3:24) [forward demodulation 2540,2476]
[SA] passive: 2691. $less(alength,X4) | alength = X5 | $less(X5,0) | X4 = X5 | i(l8(sK2(X5))) = X5 | ~$less(X6,X4) | $less(X6,X5) (3:24) [forward demodulation 2540,2476]
[SA] new: 2692. alength = X2 | $less(alength,X1) | $less(X1,X2) | X1 = X2 | i(l8(sK2(X2))) = X2 | alength = 0 | $less(X2,alength) (5:24) [forward demodulation 2541,2476]
[SA] new: 2693. $less(alength,X3) | alength = X4 | $less(X3,X4) | X3 = X4 | i(l8(sK2(X4))) = X4 | 0 = sK0 | $less(X4,alength) (4:24) [forward demodulation 2542,2476]
[SA] new: 2694. $less(alength,X5) | alength = X6 | $less(X5,X6) | X5 = X6 | i(l8(sK2(X6))) = X6 | 0 = sK0 | $less(X6,sK0) (3:24) [forward demodulation 2543,2476]
[SA] passive: 2694. $less(alength,X5) | alength = X6 | $less(X5,X6) | X5 = X6 | i(l8(sK2(X6))) = X6 | 0 = sK0 | $less(X6,sK0) (3:24) [forward demodulation 2543,2476]
[SA] new: 2695. alength = X9 | $less(alength,X8) | $less(X8,X9) | X8 = X9 | i(l8(sK2(X9))) = X9 | 0 = X10 | $less(X10,0) | $less(X9,X10) (3:27) [forward demodulation 2545,2476]
[SA] passive: 2695. alength = X9 | $less(alength,X8) | $less(X8,X9) | X8 = X9 | i(l8(sK2(X9))) = X9 | 0 = X10 | $less(X10,0) | $less(X9,X10) (3:27) [forward demodulation 2545,2476]
[SA] new: 2696. alength = X12 | $less(alength,X11) | $less(X11,X12) | X11 = X12 | i(l8(sK2(X12))) = X12 | ~$less(X13,X12) | $less(X13,0) (3:24) [forward demodulation 2546,2476]
[SA] passive: 2696. alength = X12 | $less(alength,X11) | $less(X11,X12) | X11 = X12 | i(l8(sK2(X12))) = X12 | ~$less(X13,X12) | $less(X13,0) (3:24) [forward demodulation 2546,2476]
[SA] passive: 2547. alength = X15 | $less(sK0,X15) | $less(X15,0) | sK0 = X15 | i(l8(sK2(X15))) = X15 (3:18) [backward demodulation 2476,541]
[SA] new: 2697. $less(alength,X4) | alength = X3 | i(l8(sK1(X2,sK2(X3)))) = X2 | $less(X3,0) | a(main_end,X2) = a(l11(sK2(X3)),X2) | X3 = X4 | $less(X4,X3) (3:32) [forward demodulation 2548,2476]
[SA] passive: 2697. $less(alength,X4) | alength = X3 | i(l8(sK1(X2,sK2(X3)))) = X2 | $less(X3,0) | a(main_end,X2) = a(l11(sK2(X3)),X2) | X3 = X4 | $less(X4,X3) (3:32) [forward demodulation 2548,2476]
[SA] new: 2698. ~$less(X7,alength) | alength = X6 | i(l8(sK1(X5,sK2(X6)))) = X5 | $less(X6,0) | a(main_end,X5) = a(l11(sK2(X6)),X5) | $less(X7,X6) (3:29) [forward demodulation 2549,2476]
[SA] passive: 2698. ~$less(X7,alength) | alength = X6 | i(l8(sK1(X5,sK2(X6)))) = X5 | $less(X6,0) | a(main_end,X5) = a(l11(sK2(X6)),X5) | $less(X7,X6) (3:29) [forward demodulation 2549,2476]
[SA] passive: 2652. alength = 0 | i(l8(sK1(X12,sK2(0)))) = X12 | a(main_end,X12) = a(l11(sK2(0)),X12) (5:20) [duplicate literal removal 2550]
[SA] passive: 2552. alength = sK0 | alength = X0 | $less(X0,alength) | $less(i(l8(sK2(sK0))),X0) (5:15) [backward demodulation 2476,560]
[SA] passive: 2553. alength = sK0 | ~$less(X1,i(l8(sK2(sK0)))) | $less(X1,alength) (5:12) [backward demodulation 2476,561]
[SA] new: 2699. $less(alength,X0) | alength = i(l8(p(nl8))) | zero = nl8 | $less(X0,alength) | $less(i(l8(p(nl8))),0) | i(l8(p(nl8))) = X0 | i(l8(p(nl8))) = i(l8(sK2(i(l8(p(nl8)))))) (4:39) [forward demodulation 2554,2476]
[SA] passive: 2699. $less(alength,X0) | alength = i(l8(p(nl8))) | zero = nl8 | $less(X0,alength) | $less(i(l8(p(nl8))),0) | i(l8(p(nl8))) = X0 | i(l8(p(nl8))) = i(l8(sK2(i(l8(p(nl8)))))) (4:39) [forward demodulation 2554,2476]
[SA] passive: 2555. alength = i(l8(p(nl8))) | zero = nl8 | i(l8(sK1(X9,sK2(i(l8(p(nl8))))))) = X9 | $less(i(l8(p(nl8))),0) | a(main_end,X9) = a(l11(sK2(i(l8(p(nl8))))),X9) (4:38) [backward demodulation 2476,607]
[SA] passive: 2556. alength = X10 | zero = nl8 | $less(i(l8(p(nl8))),X10) | $less(X10,0) | i(l8(p(nl8))) = X10 | i(l8(sK2(X10))) = X10 (4:27) [backward demodulation 2476,608]
[SA] passive: 2557. alength = i(l8(p(nl8))) | zero = nl8 | i(l8(p(nl8))) = i(l8(sK2(i(l8(p(nl8)))))) | $less(i(l8(p(nl8))),0) (4:27) [backward demodulation 2476,609]
[SA] passive: 2558. alength != 5 | s(s(s(s(s(zero))))) = nl8 (6:11) [backward demodulation 2476,632]
[SA] passive: 2559. alength = X24 | sK0 = X24 | $less(sK0,X24) | i(l8(sK1(X25,sK2(X24)))) = X25 | $less(X24,0) | a(main_end,X25) = a(l11(sK2(X24)),X25) (4:29) [backward demodulation 2476,677]
[SA] passive: 2560. alength = X27 | sK0 = X26 | $less(sK0,X26) | $less(X26,X27) | $less(X27,0) | X26 = X27 | i(l8(sK2(X27))) = X27 (4:24) [backward demodulation 2476,678]
[SA] passive: 2561. alength = sK0 | sK0 != 3 | s(s(s(zero))) = sK2(sK0) (5:13) [backward demodulation 2476,689]
[SA] new: 2700. alength = 0 | $less(alength,sK0) | 0 = sK0 | $less(X5,0) | i(main_end) = X5 | i(l8(sK2(X5))) = X5 | $less(0,X5) (4:25) [forward demodulation 2562,2476]
[SA] new: 2701. alength = X5 | alength = 0 | $less(alength,sK0) | 0 = sK0 | $less(X5,0) | i(l8(sK2(X5))) = X5 | $less(0,X5) (4:24) [forward demodulation 2700,2476]
[SA] passive: 2701. alength = X5 | alength = 0 | $less(alength,sK0) | 0 = sK0 | $less(X5,0) | i(l8(sK2(X5))) = X5 | $less(0,X5) (4:24) [forward demodulation 2700,2476]
[SA] new: 2702. alength = 0 | $less(alength,sK0) | 0 = sK0 | 'Sub'(sK2(0),nl8) (4:13) [forward demodulation 2563,2476]
[SA] passive: 2702. alength = 0 | $less(alength,sK0) | 0 = sK0 | 'Sub'(sK2(0),nl8) (4:13) [forward demodulation 2563,2476]
[SA] new: 2703. alength = sK0 | alength = X1 | ~$less(X0,sK0) | $less(X1,0) | i(l8(sK2(X1))) = X1 | $less(X0,X1) (5:21) [forward demodulation 2564,2476]
[SA] passive: 2703. alength = sK0 | alength = X1 | ~$less(X0,sK0) | $less(X1,0) | i(l8(sK2(X1))) = X1 | $less(X0,X1) (5:21) [forward demodulation 2564,2476]
[SA] passive: 2565. alength = sK0 | ~$less(X2,sK0) | $less(X2,0) | i(l8(sK2(X2))) = X2 (5:15) [backward demodulation 2476,720]
[SA] passive: 2566. alength = sK0 | ~$less(X3,sK0) | $less(X3,0) | 'Sub'(sK2(X3),nl8) (5:13) [backward demodulation 2476,721]
[SA] new: 2704. ~$less(alength,sK0) | alength = sK0 (5:6) [forward demodulation 2567,2476]
[SA] passive: 2704. ~$less(alength,sK0) | alength = sK0 (5:6) [forward demodulation 2567,2476]
[SA] new: 2705. alength = sK0 | alength = X6 | ~$less(X5,sK0) | $less(X6,i(main_end)) | $less(X5,X6) (5:16) [forward demodulation 2568,2476]
[SA] new: 2706. $less(X6,alength) | alength = sK0 | alength = X6 | ~$less(X5,sK0) | $less(X5,X6) (5:15) [forward demodulation 2705,2476]
[SA] new: 2707. $less(X8,alength) | alength = sK0 | ~$less(X7,sK0) | ~$less(X8,X7) (5:12) [forward demodulation 2569,2476]
[SA] passive: 2707. $less(X8,alength) | alength = sK0 | ~$less(X7,sK0) | ~$less(X8,X7) (5:12) [forward demodulation 2569,2476]
[SA] new: 2708. $less(0,alength) | alength = sK0 | 0 = sK0 | i(main_end) = 0 (6:13) [forward demodulation 2570,2476]
[SA] new: 2709. $less(X2,alength) | alength = sK0 | $less(sK0,X2) | sK0 = X2 (6:12) [forward demodulation 2571,2476]
[SA] new: 2710. $less(0,alength) | alength = sK0 | 0 = sK0 (6:9) [forward demodulation 2572,2476]
[SA] new: 2711. $less(alength,X1) | alength = X0 | $less(sK0,X0) | $less(X1,alength) | $less(X0,0) | X0 = X1 | i(l8(sK2(X0))) = X0 (4:24) [forward demodulation 2651,2476]
[SA] passive: 2711. $less(alength,X1) | alength = X0 | $less(sK0,X0) | $less(X1,alength) | $less(X0,0) | X0 = X1 | i(l8(sK2(X0))) = X0 (4:24) [forward demodulation 2651,2476]
[SA] passive: 2650. alength = X26 | $less(sK0,X26) | i(l8(sK1(X27,sK2(X26)))) = X27 | $less(X26,0) | a(main_end,X27) = a(l11(sK2(X26)),X27) (4:26) [duplicate literal removal 2574]
[SA] passive: 2575. alength = X29 | alength = X28 | $less(sK0,X28) | $less(X28,X29) | $less(X29,0) | X28 = X29 | i(l8(sK2(X29))) = X29 (4:24) [backward demodulation 2476,807]
[SA] passive: 2649. alength = X30 | $less(sK0,X30) | i(l8(sK2(X30))) = X30 | $less(X30,0) (4:15) [duplicate literal removal 2576]
[SA] new: 2712. alength = sK0 | alength = X6 | 0 = sK0 | $less(X6,0) | i(l8(sK2(X6))) = X6 | $less(0,X6) (4:21) [forward demodulation 2577,2476]
[SA] passive: 2712. alength = sK0 | alength = X6 | 0 = sK0 | $less(X6,0) | i(l8(sK2(X6))) = X6 | $less(0,X6) (4:21) [forward demodulation 2577,2476]
[SA] passive: 2578. alength = sK0 | 0 = sK0 | i(l8(sK2(0))) = 0 (4:12) [backward demodulation 2476,847]
[SA] passive: 2579. alength = sK0 | 'Sub'(sK2(0),nl8) | 0 = sK0 (4:10) [backward demodulation 2476,848]
[SA] new: 2713. $less(0,alength) | alength = sK0 | i(main_end) = 0 (6:10) [forward demodulation 2580,2476]
[SA] new: 2714. alength = 0 | $less(0,alength) | alength = sK0 (6:9) [forward demodulation 2713,2476]
[SA] new: 2715. alength = 0 | alength = X6 | i(main_end) = sK0 | $less(X6,0) | i(l8(sK2(X6))) = X6 | $less(0,X6) (4:22) [forward demodulation 2581,2476]
[SA] new: 2716. alength = sK0 | alength = 0 | alength = X6 | $less(X6,0) | i(l8(sK2(X6))) = X6 | $less(0,X6) (4:21) [forward demodulation 2715,2476]
[SA] passive: 2716. alength = sK0 | alength = 0 | alength = X6 | $less(X6,0) | i(l8(sK2(X6))) = X6 | $less(0,X6) (4:21) [forward demodulation 2715,2476]
[SA] new: 2717. alength = 0 | alength = sK0 | i(l8(sK2(0))) = 0 (4:12) [forward demodulation 2582,2476]
[SA] new: 2718. alength = 0 | alength = sK0 | 'Sub'(sK2(0),nl8) (4:10) [forward demodulation 2583,2476]
[SA] passive: 2718. alength = 0 | alength = sK0 | 'Sub'(sK2(0),nl8) (4:10) [forward demodulation 2583,2476]
[SA] new: 2719. alength = X1 | alength = X0 | $less(alength,X0) | $less(X1,0) | i(l8(sK2(X1))) = X1 | $less(X0,X1) (4:21) [forward demodulation 2648,2476]
[SA] new: 2720. alength = X8 | alength = X7 | $less(alength,X7) | $less(X8,i(main_end)) | $less(X7,X8) (4:16) [forward demodulation 2646,2476]
[SA] new: 2721. $less(X8,alength) | alength = X8 | alength = X7 | $less(alength,X7) | $less(X7,X8) (4:15) [forward demodulation 2720,2476]
[SA] new: 2722. $less(X10,alength) | alength = X9 | $less(alength,X9) | ~$less(X10,X9) (4:12) [forward demodulation 2645,2476]
[SA] passive: 2588. alength = sK0 | 0 = sK0 | a(main_end,X0) = a(l11(sK2(0)),X0) | i(l8(sK1(X0,sK2(0)))) = X0 (5:23) [backward demodulation 2476,1021]
[SA] passive: 2589. alength = sK0 | 0 = sK0 | $less(i(l8(sK2(0))),alength) (5:12) [backward demodulation 2476,1022]
[SA] new: 2723. $less(alength,X14) | alength = X14 | $less(X14,alength) | sK0 = X14 | $less(X15,0) | i(main_end) = X15 | i(l8(sK2(X15))) = X15 | $less(sK0,X15) (4:28) [forward demodulation 2590,2476]
[SA] new: 2724. alength = X16 | $less(alength,X16) | $less(X16,alength) | sK0 = X16 | i(l8(sK2(sK0))) = sK0 (4:18) [forward demodulation 2591,2476]
[SA] new: 2725. $less(alength,X17) | alength = X17 | $less(X17,alength) | sK0 = X17 | 'Sub'(sK2(sK0),nl8) (4:16) [forward demodulation 2592,2476]
[SA] new: 2726. alength = X13 | alength = sK0 | $less(sK0,X13) | $less(X14,0) | i(main_end) = X14 | i(l8(sK2(X14))) = X14 | $less(X13,X14) (4:25) [forward demodulation 2593,2476]
[SA] new: 2727. alength = X14 | alength = X13 | alength = sK0 | $less(sK0,X13) | $less(X14,0) | i(l8(sK2(X14))) = X14 | $less(X13,X14) (4:24) [forward demodulation 2726,2476]
[SA] passive: 2727. alength = X14 | alength = X13 | alength = sK0 | $less(sK0,X13) | $less(X14,0) | i(l8(sK2(X14))) = X14 | $less(X13,X14) (4:24) [forward demodulation 2726,2476]
[SA] new: 2728. alength = X16 | alength = sK0 | $less(sK0,X16) | $less(X16,0) | 'Sub'(sK2(X16),nl8) (4:16) [forward demodulation 2594,2476]
[SA] passive: 2728. alength = X16 | alength = sK0 | $less(sK0,X16) | $less(X16,0) | 'Sub'(sK2(X16),nl8) (4:16) [forward demodulation 2594,2476]
[SA] new: 2729. $less(alength,X14) | alength = X14 | $less(X14,alength) | $less(X15,0) | i(main_end) = X15 | i(l8(sK2(X15))) = X15 | $less(sK0,X15) (4:25) [forward demodulation 2644,2476]
[SA] new: 2730. $less(alength,X17) | alength = X17 | $less(X17,alength) | 'Sub'(sK2(sK0),nl8) (4:13) [forward demodulation 2643,2476]
[SA] passive: 2598. alength != 6 | s(s(s(s(s(s(zero)))))) = nl8 (7:12) [backward demodulation 2476,1264]
[SA] new: 2731. $less(alength,X2) | alength = X1 | sK0 = X0 | 0 = sK0 | $less(X0,sK0) | $less(X1,X0) | $less(X2,X1) | X1 = X2 | i(l8(sK2(X1))) = X1 (4:30) [forward demodulation 2599,2476]
[SA] passive: 2731. $less(alength,X2) | alength = X1 | sK0 = X0 | 0 = sK0 | $less(X0,sK0) | $less(X1,X0) | $less(X2,X1) | X1 = X2 | i(l8(sK2(X1))) = X1 (4:30) [forward demodulation 2599,2476]
[SA] new: 2732. alength = sK0 | alength = X1 | 0 = sK0 | $less(X1,i(main_end)) | $less(0,X1) (7:16) [forward demodulation 2600,2476]
[SA] new: 2733. $less(X1,alength) | alength = sK0 | alength = X1 | 0 = sK0 | $less(0,X1) (7:15) [forward demodulation 2732,2476]
[SA] new: 2734. $less(X2,alength) | alength = sK0 | 0 = sK0 | ~$less(X2,0) (7:12) [forward demodulation 2601,2476]
[SA] new: 2735. $less(alength,X18) | alength = X17 | sK0 = X17 | 0 = sK0 | $less(0,X17) | $less(X18,sK0) | $less(X17,0) | X17 = X18 | i(l8(sK2(X17))) = X17 (4:30) [forward demodulation 2602,2476]
[SA] passive: 2735. $less(alength,X18) | alength = X17 | sK0 = X17 | 0 = sK0 | $less(0,X17) | $less(X18,sK0) | $less(X17,0) | X17 = X18 | i(l8(sK2(X17))) = X17 (4:30) [forward demodulation 2602,2476]
[SA] new: 2736. alength = X47 | $less(alength,sK0) | sK0 = X47 | 0 = sK0 | $less(0,X47) | i(l8(sK1(X48,sK2(X47)))) = X48 | $less(X47,0) | a(main_end,X48) = a(l11(sK2(X47)),X48) (4:35) [forward demodulation 2603,2476]
[SA] passive: 2736. alength = X47 | $less(alength,sK0) | sK0 = X47 | 0 = sK0 | $less(0,X47) | i(l8(sK1(X48,sK2(X47)))) = X48 | $less(X47,0) | a(main_end,X48) = a(l11(sK2(X47)),X48) (4:35) [forward demodulation 2603,2476]
[SA] new: 2737. alength = X50 | $less(alength,sK0) | sK0 = X49 | 0 = sK0 | $less(0,X49) | $less(X49,X50) | $less(X50,0) | X49 = X50 | i(l8(sK2(X50))) = X50 (4:30) [forward demodulation 2604,2476]
[SA] passive: 2737. alength = X50 | $less(alength,sK0) | sK0 = X49 | 0 = sK0 | $less(0,X49) | $less(X49,X50) | $less(X50,0) | X49 = X50 | i(l8(sK2(X50))) = X50 (4:30) [forward demodulation 2604,2476]
[SA] new: 2738. alength = X51 | $less(alength,sK0) | sK0 = X51 | 0 = sK0 | $less(0,X51) | i(l8(sK2(X51))) = X51 | $less(X51,0) (4:24) [forward demodulation 2605,2476]
[SA] passive: 2738. alength = X51 | $less(alength,sK0) | sK0 = X51 | 0 = sK0 | $less(0,X51) | i(l8(sK2(X51))) = X51 | $less(X51,0) (4:24) [forward demodulation 2605,2476]
[SA] new: 2739. $less(X15,alength) | alength = X15 | $less(sK0,X15) | sK0 = X15 (4:12) [forward demodulation 2606,2476]
[SA] new: 2740. alength = X16 | alength = X15 | $less(sK0,X15) | sK0 = X15 | $less(X16,0) | i(l8(sK2(X16))) = X16 | $less(X15,X16) (4:24) [forward demodulation 2607,2476]
[SA] passive: 2740. alength = X16 | alength = X15 | $less(sK0,X15) | sK0 = X15 | $less(X16,0) | i(l8(sK2(X16))) = X16 | $less(X15,X16) (4:24) [forward demodulation 2607,2476]
[SA] passive: 2608. alength = X18 | $less(sK0,X18) | sK0 = X18 | $less(X18,0) | 'Sub'(sK2(X18),nl8) (4:16) [backward demodulation 2476,1522]
[SA] new: 2741. alength = X16 | alength = sK0 | $less(sK0,X15) | sK0 = X15 | $less(X16,0) | i(l8(sK2(X16))) = X16 | $less(X15,X16) (4:24) [forward demodulation 2609,2476]
[SA] passive: 2741. alength = X16 | alength = sK0 | $less(sK0,X15) | sK0 = X15 | $less(X16,0) | i(l8(sK2(X16))) = X16 | $less(X15,X16) (4:24) [forward demodulation 2609,2476]
[SA] passive: 2610. alength = sK0 | $less(sK0,X17) | sK0 = X17 | $less(X17,0) | i(l8(sK2(X17))) = X17 (4:18) [backward demodulation 2476,1592]
[SA] passive: 2611. alength = sK0 | $less(sK0,X18) | sK0 = X18 | $less(X18,0) | 'Sub'(sK2(X18),nl8) (4:16) [backward demodulation 2476,1593]
[SA] new: 2742. $less(0,alength) | alength = 0 | 0 = sK0 (5:9) [forward demodulation 2612,2476]
[SA] new: 2743. alength = X8 | alength = 0 | 0 = sK0 | $less(X8,0) | i(l8(sK2(X8))) = X8 | $less(0,X8) (5:21) [forward demodulation 2613,2476]
[SA] passive: 2743. alength = X8 | alength = 0 | 0 = sK0 | $less(X8,0) | i(l8(sK2(X8))) = X8 | $less(0,X8) (5:21) [forward demodulation 2613,2476]
[SA] passive: 2614. alength = 0 | 'Sub'(sK2(0),nl8) | 0 = sK0 (5:10) [backward demodulation 2476,1632]
[SA] new: 2744. alength = 0 | alength = X1 | 0 = sK0 | $less(X1,i(main_end)) | $less(0,X1) (6:16) [forward demodulation 2615,2476]
[SA] new: 2745. $less(X1,alength) | alength = 0 | alength = X1 | 0 = sK0 | $less(0,X1) (6:15) [forward demodulation 2744,2476]
[SA] new: 2746. $less(X2,alength) | alength = 0 | 0 = sK0 | ~$less(X2,0) (6:12) [forward demodulation 2616,2476]
[SA] passive: 2617. alength = 0 | 0 = sK0 | $less(i(l8(sK2(0))),alength) (6:12) [backward demodulation 2476,1736]
[SA] new: 2747. $less(alength,X19) | alength = X19 | $less(X19,sK0) | 0 = sK0 | sK0 = X19 | $less(X20,0) | i(main_end) = X20 | i(l8(sK2(X20))) = X20 | $less(0,X20) (4:31) [forward demodulation 2629,2476]
[SA] new: 2748. alength = X20 | $less(alength,X19) | alength = X19 | $less(X19,sK0) | 0 = sK0 | sK0 = X19 | $less(X20,0) | i(l8(sK2(X20))) = X20 | $less(0,X20) (4:30) [forward demodulation 2747,2476]
[SA] passive: 2748. alength = X20 | $less(alength,X19) | alength = X19 | $less(X19,sK0) | 0 = sK0 | sK0 = X19 | $less(X20,0) | i(l8(sK2(X20))) = X20 | $less(0,X20) (4:30) [forward demodulation 2747,2476]
[SA] new: 2749. $less(alength,X22) | alength = X22 | $less(X22,sK0) | 0 = sK0 | sK0 = X22 | 'Sub'(sK2(0),nl8) (4:19) [forward demodulation 2630,2476]
[SA] passive: 2749. $less(alength,X22) | alength = X22 | $less(X22,sK0) | 0 = sK0 | sK0 = X22 | 'Sub'(sK2(0),nl8) (4:19) [forward demodulation 2630,2476]
[SA] new: 2750. $less(alength,X21) | alength = X21 | $less(X21,sK0) | 0 = sK0 | sK0 = X21 | i(l8(sK2(0))) = 0 (4:21) [forward demodulation 2631,2476]
[SA] passive: 2750. $less(alength,X21) | alength = X21 | $less(X21,sK0) | 0 = sK0 | sK0 = X21 | i(l8(sK2(0))) = 0 (4:21) [forward demodulation 2631,2476]
[SA] new: 2751. alength = X19 | alength = sK0 | $less(0,X19) | 0 = sK0 | $less(X20,0) | i(main_end) = X20 | i(l8(sK2(X20))) = X20 | $less(X19,X20) (4:28) [forward demodulation 2632,2476]
[SA] new: 2752. alength = X20 | alength = X19 | alength = sK0 | $less(0,X19) | 0 = sK0 | $less(X20,0) | i(l8(sK2(X20))) = X20 | $less(X19,X20) (4:27) [forward demodulation 2751,2476]
[SA] passive: 2752. alength = X20 | alength = X19 | alength = sK0 | $less(0,X19) | 0 = sK0 | $less(X20,0) | i(l8(sK2(X20))) = X20 | $less(X19,X20) (4:27) [forward demodulation 2751,2476]
[SA] new: 2753. alength = X22 | alength = sK0 | $less(0,X22) | 0 = sK0 | $less(X22,0) | 'Sub'(sK2(X22),nl8) (4:19) [forward demodulation 2633,2476]
[SA] passive: 2753. alength = X22 | alength = sK0 | $less(0,X22) | 0 = sK0 | $less(X22,0) | 'Sub'(sK2(X22),nl8) (4:19) [forward demodulation 2633,2476]
[SA] new: 2754. $less(X24,alength) | alength = sK0 | $less(0,X24) | 0 = sK0 | sK0 = X24 (6:15) [forward demodulation 2634,2476]
[SA] new: 2755. alength = X20 | alength = sK0 | $less(0,X19) | 0 = sK0 | sK0 = X19 | $less(X20,0) | i(l8(sK2(X20))) = X20 | $less(X19,X20) (4:27) [forward demodulation 2635,2476]
[SA] passive: 2755. alength = X20 | alength = sK0 | $less(0,X19) | 0 = sK0 | sK0 = X19 | $less(X20,0) | i(l8(sK2(X20))) = X20 | $less(X19,X20) (4:27) [forward demodulation 2635,2476]
[SA] passive: 2636. alength = sK0 | $less(0,X21) | 0 = sK0 | sK0 = X21 | $less(X21,0) | i(l8(sK2(X21))) = X21 (4:21) [backward demodulation 2476,2179]
[SA] passive: 2637. alength = sK0 | $less(0,X22) | 0 = sK0 | sK0 = X22 | $less(X22,0) | 'Sub'(sK2(X22),nl8) (4:19) [backward demodulation 2476,2180]
[SA] new: 2756. alength = X21 | $less(alength,sK0) | $less(0,X21) | 0 = sK0 | 0 = X21 | $less(X22,0) | i(main_end) = X22 | i(l8(sK2(X22))) = X22 | $less(X21,X22) (4:31) [forward demodulation 2642,2476]
[SA] new: 2757. alength = X22 | alength = X21 | $less(alength,sK0) | $less(0,X21) | 0 = sK0 | 0 = X21 | $less(X22,0) | i(l8(sK2(X22))) = X22 | $less(X21,X22) (4:30) [forward demodulation 2756,2476]
[SA] passive: 2757. alength = X22 | alength = X21 | $less(alength,sK0) | $less(0,X21) | 0 = sK0 | 0 = X21 | $less(X22,0) | i(l8(sK2(X22))) = X22 | $less(X21,X22) (4:30) [forward demodulation 2756,2476]
[SA] active: 2477. alength = i(l8(nl8)) (0:5:1) [backward demodulation 2476,54]
[SA] new: 2758. alength != i(l8(X0)) | nl8 = X0 (1:8) [superposition 48,2477]
[SA] new: 2759. alength != i(l8(X1)) | nl8 = X1 (1:8) [superposition 48,2477]
[SA] new: 2760. v = a(l11(nl8),alength) (1:6) [superposition 70,2477]
[SA] new: 2761. i(l8(s(nl8))) = $sum(alength,1) (1:8) [superposition 74,2477]
[SA] new: 2762. alength != 0 | zero = nl8 (2:6) [superposition 108,2477]
[SA] new: 2763. alength != 1 | s(zero) = nl8 (3:7) [superposition 121,2477]
[SA] new: 2764. alength != 2 | s(s(zero)) = nl8 (4:8) [superposition 130,2477]
[SA] new: 2765. alength != 3 | s(s(s(zero))) = nl8 (5:9) [superposition 183,2477]
[SA] new: 2766. v = a(l8(s(nl8)),alength) (4:7) [superposition 349,2477]
[SA] passive: 2762. alength != 0 | zero = nl8 (2:6) [superposition 108,2477]
[SA] passive: 2763. alength != 1 | s(zero) = nl8 (3:7) [superposition 121,2477]
[SA] active: 2476. alength = i(main_end) (7:4:1) [subsumption resolution 2451,46]
[SA] new: 2767. alength != 0 | zero = nl8 (8:6) [superposition 146,2476]
[SA] new: 2768. alength != 1 | s(zero) = nl8 (8:7) [superposition 153,2476]
[SA] new: 2769. alength != -1 (8:3) [superposition 163,2476]
[SA] passive: 2769. alength != -1 (8:3) [superposition 163,2476]
[SA] active: 2479. 'Sub'(sK2(X1),nl8) | $less(X1,0) | ~$less(X1,alength) (0:10:1) [backward demodulation 2476,83]
[SA] new: 2770. $less(X0,0) | ~$less(X0,alength) | a(main_end,X1) = a(l11(sK2(X0)),X1) | i(l8(sK1(X1,sK2(X0)))) = X1 (2:23) [resolution 2479,150]
[SA] new: 2771. $less(X2,0) | ~$less(X2,alength) | $less(i(l8(sK2(X2))),alength) (1:12) [resolution 2479,56]
[SA] passive: 2770. $less(X0,0) | ~$less(X0,alength) | a(main_end,X1) = a(l11(sK2(X0)),X1) | i(l8(sK1(X1,sK2(X0)))) = X1 (2:23) [resolution 2479,150]
[SA] passive: 2771. $less(X2,0) | ~$less(X2,alength) | $less(i(l8(sK2(X2))),alength) (1:12) [resolution 2479,56]
[SA] active: 2769. alength != -1 (8:3:1) [superposition 163,2476]
[SA] active: 2480. ~$less(X1,alength) | $less(X1,0) | i(l8(sK2(X1))) = X1 (0:12:1) [backward demodulation 2476,84]
[SA] new: 2772. $less(X0,0) | i(l8(sK2(X0))) = X0 | $less(sK0,X0) | sK0 = X0 (3:15) [resolution 2480,106]
[SA] new: 2773. $less(X1,0) | i(l8(sK2(X1))) = X1 | $less(0,X1) | alength = X1 | 0 = sK0 (4:18) [resolution 2480,178]
[SA] new: 2774. $less(X2,0) | i(l8(sK2(X2))) = X2 | $less(sK0,X2) | alength = X2 (3:15) [resolution 2480,179]
[SA] new: 2775. $less(X3,0) | i(l8(sK2(X3))) = X3 | $less(sK0,X4) | $less(X4,X3) | X3 = X4 | sK0 = X3 (4:21) [resolution 2480,198]
[SA] new: 2776. $less(X5,0) | i(l8(sK2(X5))) = X5 | $less(sK0,X6) | $less(X6,X5) | alength = X5 | sK0 = X6 (4:21) [resolution 2480,204]
[SA] new: 2777. $less(X7,0) | i(l8(sK2(X7))) = X7 | $less(sK0,X8) | $less(X8,X7) | X7 = X8 | alength = X7 (4:21) [resolution 2480,246]
[SA] new: 2778. $less(X9,0) | i(l8(sK2(X9))) = X9 | $less(sK0,X10) | $less(X10,X9) | alength = X9 | alength = X10 (4:21) [resolution 2480,253]
[SA] new: 2779. $less(X11,0) | i(l8(sK2(X11))) = X11 | $less(0,X11) | alength = X11 | alength = 0 (5:18) [resolution 2480,261]
[SA] new: 2780. $less(X12,0) | i(l8(sK2(X12))) = X12 | $less(0,X12) | 0 = sK0 | 0 = X12 (5:18) [resolution 2480,278]
[SA] new: 2781. $less(X13,0) | i(l8(sK2(X13))) = X13 | $less(0,X13) | 0 = sK0 | sK0 = X13 (4:18) [resolution 2480,330]
[SA] new: 2782. $less(X14,0) | i(l8(sK2(X14))) = X14 | $less(sK0,X15) | $less(X15,X14) | X14 = X15 | sK0 = X15 (4:21) [resolution 2480,388]
[SA] new: 2783. $less(X16,0) | i(l8(sK2(X16))) = X16 | $less(sK0,X17) | $less(X17,X16) | sK0 = X17 | sK0 = X16 (4:21) [resolution 2480,408]
[SA] new: 2784. $less(X18,0) | i(l8(sK2(X18))) = X18 | $less(alength,X18) | alength = X18 (1:15) [resolution 2480,15]
[SA] new: 2785. $less(X19,0) | i(l8(sK2(X19))) = X19 | $less(alength,X19) | alength = X19 (1:15) [resolution 2480,15]
[SA] new: 2786. $less(X20,0) | i(l8(sK2(X20))) = X20 | $less(X21,X20) | $less(alength,X21) | alength = X20 | alength = X21 (3:21) [resolution 2480,174]
[SA] new: 2787. $less(X22,0) | i(l8(sK2(X22))) = X22 | $less(alength,X23) | $less(X23,X22) | X22 = X23 | alength = X22 (3:21) [resolution 2480,174]
[SA] new: 2788. $less(X24,0) | i(l8(sK2(X24))) = X24 | $less(X25,X24) | $less(alength,X25) | alength = X25 | X24 = X25 (3:21) [resolution 2480,174]
[SA] new: 2789. $less(X26,0) | i(l8(sK2(X26))) = X26 | $less(sK0,X26) | $less(alength,alength) | alength = X26 | alength = sK0 (4:21) [resolution 2480,198]
[SA] new: 2790. $less(X27,0) | i(l8(sK2(X27))) = X27 | $less(sK0,X27) | $less(alength,alength) | alength = alength | sK0 = X27 (4:21) [resolution 2480,204]
[SA] new: 2791. $less(X28,0) | i(l8(sK2(X28))) = X28 | $less(sK0,X28) | $less(alength,alength) | alength = X28 | alength = alength (4:21) [resolution 2480,246]
[SA] new: 2792. $less(X29,0) | i(l8(sK2(X29))) = X29 | $less(sK0,X29) | $less(alength,alength) | alength = alength | alength = X29 (4:21) [resolution 2480,253]
[SA] new: 2793. $less(X30,0) | i(l8(sK2(X30))) = X30 | $less(0,X30) | $less(alength,sK0) | alength = X30 | 0 = sK0 | alength = sK0 (4:24) [resolution 2480,323]
[SA] new: 2794. $less(X31,0) | i(l8(sK2(X31))) = X31 | $less(0,X31) | $less(alength,sK0) | alength = sK0 | 0 = sK0 | sK0 = X31 (4:24) [resolution 2480,332]
[SA] new: 2795. $less(X32,0) | i(l8(sK2(X32))) = X32 | $less(0,X32) | $less(alength,sK0) | alength = X32 | 0 = sK0 | 0 = X32 (4:24) [resolution 2480,387]
[SA] new: 2796. $less(X33,0) | i(l8(sK2(X33))) = X33 | $less(sK0,X33) | $less(alength,alength) | alength = X33 | sK0 = X33 (4:21) [resolution 2480,388]
[SA] new: 2797. $less(X34,0) | i(l8(sK2(X34))) = X34 | $less(sK0,X34) | $less(alength,alength) | sK0 = X34 | alength = sK0 (4:21) [resolution 2480,408]
[SA] new: 2798. $less(i(l8(p(nl8))),0) | i(l8(p(nl8))) = i(l8(sK2(i(l8(p(nl8)))))) | zero = nl8 (3:21) [resolution 2480,95]
[SA] new: 2799. $less(i(l8(p(p(nl8)))),0) | i(l8(p(p(nl8)))) = i(l8(sK2(i(l8(p(p(nl8))))))) | zero = p(nl8) | zero = nl8 (4:28) [resolution 2480,314]
[SA] new: 2800. $less(i(l8(p(nl8))),0) | i(l8(p(nl8))) = i(l8(sK2(i(l8(p(nl8)))))) | alength = alength | $less(alength,alength) | zero = nl8 (4:27) [resolution 2480,193]
[SA] new: 2801. $less(0,0) | i(l8(sK2(0))) = 0 | alength = 0 (4:12) [resolution 2480,244]
[SA] new: 2802. $less(0,0) | i(l8(sK2(0))) = 0 | 0 = sK0 (3:12) [resolution 2480,113]
[SA] new: 2803. $less(0,0) | i(l8(sK2(0))) = 0 | $less(alength,sK0) | alength = sK0 | 0 = sK0 (3:18) [resolution 2480,177]
[SA] new: 2804. $less(0,0) | i(l8(sK2(0))) = 0 | $less(alength,alength) | alength = alength | 0 = sK0 (4:18) [resolution 2480,178]
[SA] new: 2805. $less(0,0) | i(l8(sK2(0))) = 0 | $less(alength,sK0) | 0 = sK0 | alength = 0 (4:18) [resolution 2480,229]
[SA] new: 2806. $less(0,0) | i(l8(sK2(0))) = 0 | $less(alength,alength) | alength = alength | alength = 0 (5:18) [resolution 2480,261]
[SA] new: 2807. $less(0,0) | i(l8(sK2(0))) = 0 | $less(alength,alength) | 0 = sK0 | alength = 0 (5:18) [resolution 2480,278]
[SA] new: 2808. $less(0,0) | i(l8(sK2(0))) = 0 | $less(alength,X35) | $less(X35,sK0) | alength = X35 | 0 = sK0 | sK0 = X35 (4:24) [resolution 2480,323]
[SA] new: 2809. $less(0,0) | i(l8(sK2(0))) = 0 | $less(alength,alength) | 0 = sK0 | alength = sK0 (4:18) [resolution 2480,330]
[SA] new: 2810. $less(0,0) | i(l8(sK2(0))) = 0 | $less(X36,sK0) | $less(alength,X36) | sK0 = X36 | 0 = sK0 | alength = sK0 (4:24) [resolution 2480,332]
[SA] new: 2811. $less(0,0) | i(l8(sK2(0))) = 0 | $less(alength,sK0) | alength = 0 | alength = sK0 (4:18) [resolution 2480,380]
[SA] new: 2812. $less(0,0) | i(l8(sK2(0))) = 0 | $less(X37,sK0) | $less(alength,X37) | alength = X37 | 0 = sK0 | alength = 0 (4:24) [resolution 2480,387]
[SA] new: 2813. $less(sK0,0) | i(l8(sK2(sK0))) = sK0 (1:9) [resolution 2480,45]
[SA] new: 2814. $less(sK0,0) | i(l8(sK2(sK0))) = sK0 | $less(alength,alength) | alength = sK0 (3:15) [resolution 2480,106]
[SA] new: 2815. $less(sK0,0) | i(l8(sK2(sK0))) = sK0 | $less(alength,alength) | alength = alength (3:15) [resolution 2480,179]
[SA] new: 2816. $less(sK0,0) | i(l8(sK2(sK0))) = sK0 | $less(alength,X38) | $less(X38,alength) | alength = X38 | sK0 = X38 (4:21) [resolution 2480,198]
[SA] new: 2817. $less(sK0,0) | i(l8(sK2(sK0))) = sK0 | $less(X39,alength) | $less(alength,X39) | alength = X39 | alength = sK0 (4:21) [resolution 2480,204]
[SA] new: 2818. $less(sK0,0) | i(l8(sK2(sK0))) = sK0 | $less(alength,X40) | $less(X40,alength) | alength = X40 | alength = X40 (4:21) [resolution 2480,246]
[SA] new: 2819. $less(sK0,0) | i(l8(sK2(sK0))) = sK0 | $less(X41,alength) | $less(alength,X41) | alength = X41 | alength = alength (4:21) [resolution 2480,253]
[SA] new: 2820. $less(sK0,0) | i(l8(sK2(sK0))) = sK0 | $less(X42,alength) | $less(alength,X42) | alength = X42 | alength = sK0 (4:21) [resolution 2480,388]
[SA] new: 2821. $less(sK0,0) | i(l8(sK2(sK0))) = sK0 | $less(X43,alength) | $less(alength,X43) | alength = sK0 | sK0 = X43 (4:21) [resolution 2480,408]
[SA] new: 2822. $less(sK0,0) | i(l8(sK2(sK0))) = sK0 | $less(alength,X40) | $less(X40,alength) | alength = X40 (4:18) [duplicate literal removal 2818]
[SA] new: 2823. i(l8(sK2(0))) = 0 | $less(X37,sK0) | $less(alength,X37) | alength = X37 | 0 = sK0 | alength = 0 (4:21) [evaluation 2812]
[SA] new: 2824. i(l8(sK2(0))) = 0 | $less(alength,sK0) | alength = 0 | alength = sK0 (4:15) [evaluation 2811]
[SA] new: 2825. i(l8(sK2(0))) = 0 | $less(X36,sK0) | $less(alength,X36) | sK0 = X36 | 0 = sK0 | alength = sK0 (4:21) [evaluation 2810]
[SA] new: 2826. i(l8(sK2(0))) = 0 | $less(alength,alength) | 0 = sK0 | alength = sK0 (4:15) [evaluation 2809]
[SA] new: 2827. i(l8(sK2(0))) = 0 | $less(alength,X35) | $less(X35,sK0) | alength = X35 | 0 = sK0 | sK0 = X35 (4:21) [evaluation 2808]
[SA] new: 2828. i(l8(sK2(0))) = 0 | $less(alength,alength) | 0 = sK0 | alength = 0 (5:15) [evaluation 2807]
[SA] new: 2829. i(l8(sK2(0))) = 0 | $less(alength,sK0) | 0 = sK0 | alength = 0 (4:15) [evaluation 2805]
[SA] new: 2830. i(l8(sK2(0))) = 0 | $less(alength,sK0) | alength = sK0 | 0 = sK0 (3:15) [evaluation 2803]
[SA] new: 2831. i(l8(sK2(0))) = 0 | 0 = sK0 (3:9) [evaluation 2802]
[SA] new: 2832. i(l8(sK2(0))) = 0 | alength = 0 (4:9) [evaluation 2801]
[SA] passive: 2772. $less(X0,0) | i(l8(sK2(X0))) = X0 | $less(sK0,X0) | sK0 = X0 (3:15) [resolution 2480,106]
[SA] passive: 2773. $less(X1,0) | i(l8(sK2(X1))) = X1 | $less(0,X1) | alength = X1 | 0 = sK0 (4:18) [resolution 2480,178]
[SA] passive: 2775. $less(X3,0) | i(l8(sK2(X3))) = X3 | $less(sK0,X4) | $less(X4,X3) | X3 = X4 | sK0 = X3 (4:21) [resolution 2480,198]
[SA] passive: 2776. $less(X5,0) | i(l8(sK2(X5))) = X5 | $less(sK0,X6) | $less(X6,X5) | alength = X5 | sK0 = X6 (4:21) [resolution 2480,204]
[SA] passive: 2777. $less(X7,0) | i(l8(sK2(X7))) = X7 | $less(sK0,X8) | $less(X8,X7) | X7 = X8 | alength = X7 (4:21) [resolution 2480,246]
[SA] passive: 2778. $less(X9,0) | i(l8(sK2(X9))) = X9 | $less(sK0,X10) | $less(X10,X9) | alength = X9 | alength = X10 (4:21) [resolution 2480,253]
[SA] passive: 2779. $less(X11,0) | i(l8(sK2(X11))) = X11 | $less(0,X11) | alength = X11 | alength = 0 (5:18) [resolution 2480,261]
[SA] passive: 2781. $less(X13,0) | i(l8(sK2(X13))) = X13 | $less(0,X13) | 0 = sK0 | sK0 = X13 (4:18) [resolution 2480,330]
[SA] passive: 2782. $less(X14,0) | i(l8(sK2(X14))) = X14 | $less(sK0,X15) | $less(X15,X14) | X14 = X15 | sK0 = X15 (4:21) [resolution 2480,388]
[SA] passive: 2783. $less(X16,0) | i(l8(sK2(X16))) = X16 | $less(sK0,X17) | $less(X17,X16) | sK0 = X17 | sK0 = X16 (4:21) [resolution 2480,408]
[SA] passive: 2798. $less(i(l8(p(nl8))),0) | i(l8(p(nl8))) = i(l8(sK2(i(l8(p(nl8)))))) | zero = nl8 (3:21) [resolution 2480,95]
[SA] passive: 2799. $less(i(l8(p(p(nl8)))),0) | i(l8(p(p(nl8)))) = i(l8(sK2(i(l8(p(p(nl8))))))) | zero = p(nl8) | zero = nl8 (4:28) [resolution 2480,314]
[SA] passive: 2831. i(l8(sK2(0))) = 0 | 0 = sK0 (3:9) [evaluation 2802]
[SA] new: 2833. i(l8(sK2(sK0))) = sK0 (1:6) [subsumption resolution 2813,46]
[SA] new: 2834. ~$less(X1,sK0) | alength = sK0 | $less(X1,alength) (5:9) [backward demodulation 2833,2553]
[SA] new: 2835. $less(sK0,X0) | alength = sK0 | alength = X0 | $less(X0,alength) (5:12) [backward demodulation 2833,2552]
[SA] new: 2836. $less(sK0,alength) | alength = sK0 (4:6) [backward demodulation 2833,2503]
[SA] passive: 2833. i(l8(sK2(sK0))) = sK0 (1:6) [subsumption resolution 2813,46]
[SA] active: 2481. v = a(l11(nl8),alength) (1:6:1) [backward demodulation 2476,88]
[SA] new: 2837. v = a(l8(nl8),alength) | alength = i(l8(nl8)) (2:11) [superposition 2481,72]
[SA] new: 2838. v = a(l8(nl8),alength) | alength = i(l8(nl8)) (2:11) [superposition 72,2481]
[SA] active: 2833. i(l8(sK2(sK0))) = sK0 (1:6:1) [subsumption resolution 2813,46]
[SA] new: 2839. i(l8(X0)) != sK0 | sK2(sK0) = X0 (2:9) [superposition 48,2833]
[SA] new: 2840. i(l8(X1)) != sK0 | sK2(sK0) = X1 (2:9) [superposition 48,2833]
[SA] new: 2841. v = a(l11(sK2(sK0)),sK0) (2:7) [superposition 70,2833]
[SA] new: 2842. i(l8(s(sK2(sK0)))) = $sum(sK0,1) (2:9) [superposition 74,2833]
[SA] new: 2843. 0 != sK0 | zero = sK2(sK0) (2:7) [superposition 108,2833]
[SA] new: 2844. 1 != sK0 | s(zero) = sK2(sK0) (3:8) [superposition 121,2833]
[SA] new: 2845. sK0 != 2 | s(s(zero)) = sK2(sK0) (4:9) [superposition 130,2833]
[SA] new: 2846. sK0 != 3 | s(s(s(zero))) = sK2(sK0) (5:10) [superposition 183,2833]
[SA] new: 2847. v = a(l8(s(sK2(sK0))),sK0) (4:8) [superposition 349,2833]
[SA] passive: 2839. i(l8(X0)) != sK0 | sK2(sK0) = X0 (2:9) [superposition 48,2833]
[SA] passive: 2841. v = a(l11(sK2(sK0)),sK0) (2:7) [superposition 70,2833]
[SA] passive: 2842. i(l8(s(sK2(sK0)))) = $sum(sK0,1) (2:9) [superposition 74,2833]
[SA] passive: 2843. 0 != sK0 | zero = sK2(sK0) (2:7) [superposition 108,2833]
[SA] passive: 2844. 1 != sK0 | s(zero) = sK2(sK0) (3:8) [superposition 121,2833]
[SA] passive: 2845. sK0 != 2 | s(s(zero)) = sK2(sK0) (4:9) [superposition 130,2833]
[SA] passive: 2846. sK0 != 3 | s(s(s(zero))) = sK2(sK0) (5:10) [superposition 183,2833]
[SA] passive: 2847. v = a(l8(s(sK2(sK0))),sK0) (4:8) [superposition 349,2833]
[SA] active: 2762. alength != 0 | zero = nl8 (2:6:1) [superposition 108,2477]
[SA] active: 2483. alength != i(l8(X1)) | nl8 = X1 (1:8:1) [backward demodulation 2476,109]
[SA] new: 2848. alength != 0 | zero = nl8 (2:6) [superposition 2483,65]
[SA] new: 2849. alength != 1 | s(zero) = nl8 (2:7) [superposition 2483,119]
[SA] new: 2850. alength != 2 | s(s(zero)) = nl8 (3:8) [superposition 2483,124]
[SA] new: 2851. alength != 3 | s(s(s(zero))) = nl8 (4:9) [superposition 2483,133]
[SA] new: 2852. alength != 4 | s(s(s(s(zero)))) = nl8 (5:10) [superposition 2483,186]
[SA] new: 2853. alength != 5 | s(s(s(s(s(zero))))) = nl8 (6:11) [superposition 2483,312]
[SA] new: 2854. alength != 6 | s(s(s(s(s(s(zero)))))) = nl8 (7:12) [superposition 2483,640]
[SA] new: 2856. alength != X0 | nl8 = sK1(X0,p(nl8)) | a(main_end,X0) = a(l11(p(nl8)),X0) | zero = nl8 (3:21) [superposition 2483,226]
[SA] new: 2857. alength != sK0 | nl8 = sK2(sK0) (2:7) [superposition 2483,2833]
[SA] passive: 2857. alength != sK0 | nl8 = sK2(sK0) (2:7) [superposition 2483,2833]
[SA] active: 2704. ~$less(alength,sK0) | alength = sK0 (5:6:1) [forward demodulation 2567,2476]
[SA] new: 2858. alength = sK0 | $less(0,X0) | $less(X0,alength) | alength = X0 | 0 = sK0 | 0 = X0 (6:18) [resolution 2704,387]
[SA] new: 2859. alength = sK0 | $less(0,X1) | $less(X1,alength) | alength = sK0 | 0 = sK0 | sK0 = X1 (6:18) [resolution 2704,332]
[SA] new: 2860. alength = sK0 | $less(0,X2) | $less(X2,alength) | alength = X2 | 0 = sK0 | alength = sK0 (6:18) [resolution 2704,323]
[SA] new: 2861. alength = sK0 | $less(0,alength) | alength = 0 | alength = sK0 (6:12) [resolution 2704,380]
[SA] new: 2862. alength = sK0 | $less(0,alength) | 0 = sK0 | alength = 0 (6:12) [resolution 2704,229]
[SA] new: 2863. alength = sK0 | $less(0,alength) | alength = sK0 | 0 = sK0 (6:12) [resolution 2704,177]
[SA] new: 2864. alength = sK0 | $less(sK0,alength) | alength = sK0 (6:9) [resolution 2704,15]
[SA] new: 2865. alength = sK0 | $less(sK0,alength) | alength = sK0 (6:9) [resolution 2704,15]
[SA] new: 2866. alength = sK0 | $less(X3,alength) | $less(sK0,X3) | alength = sK0 | sK0 = X3 (6:15) [resolution 2704,174]
[SA] new: 2867. alength = sK0 | $less(sK0,X4) | $less(X4,alength) | alength = X4 | alength = sK0 (6:15) [resolution 2704,174]
[SA] new: 2868. alength = sK0 | $less(X5,alength) | $less(sK0,X5) | sK0 = X5 | alength = X5 (6:15) [resolution 2704,174]
[SA] new: 2869. alength = sK0 | $less(sK0,alength) | $less(sK0,alength) | alength = sK0 | sK0 = sK0 (6:15) [resolution 2704,198]
[SA] new: 2870. alength = sK0 | $less(sK0,alength) | $less(sK0,alength) | alength = sK0 | alength = sK0 (6:15) [resolution 2704,204]
[SA] new: 2871. alength = sK0 | $less(sK0,alength) | $less(sK0,alength) | alength = sK0 | alength = sK0 (6:15) [resolution 2704,246]
[SA] new: 2872. alength = sK0 | $less(sK0,alength) | $less(sK0,alength) | alength = sK0 | alength = alength (6:15) [resolution 2704,253]
[SA] new: 2873. alength = sK0 | $less(0,alength) | $less(sK0,sK0) | alength = sK0 | 0 = sK0 | sK0 = sK0 (6:18) [resolution 2704,323]
[SA] new: 2874. alength = sK0 | $less(0,alength) | $less(sK0,sK0) | sK0 = sK0 | 0 = sK0 | alength = sK0 (6:18) [resolution 2704,332]
[SA] new: 2875. alength = sK0 | $less(0,alength) | $less(sK0,sK0) | alength = sK0 | 0 = sK0 | alength = 0 (6:18) [resolution 2704,387]
[SA] new: 2876. alength = sK0 | $less(sK0,alength) | $less(sK0,alength) | alength = sK0 | alength = sK0 (6:15) [resolution 2704,388]
[SA] new: 2877. alength = sK0 | $less(sK0,alength) | $less(sK0,alength) | alength = sK0 | sK0 = sK0 (6:15) [resolution 2704,408]
[SA] new: 2878. alength = sK0 | $less(sK0,alength) | sK0 = sK0 (6:9) [duplicate literal removal 2877]
[SA] new: 2879. alength = sK0 | $less(sK0,alength) (6:6) [duplicate literal removal 2876]
[SA] new: 2880. alength = sK0 | $less(0,alength) | $less(sK0,sK0) | 0 = sK0 | alength = 0 (6:15) [duplicate literal removal 2875]
[SA] new: 2881. alength = sK0 | $less(0,alength) | $less(sK0,sK0) | sK0 = sK0 | 0 = sK0 (6:15) [duplicate literal removal 2874]
[SA] new: 2882. alength = sK0 | $less(0,alength) | $less(sK0,sK0) | 0 = sK0 | sK0 = sK0 (6:15) [duplicate literal removal 2873]
[SA] new: 2883. alength = sK0 | $less(sK0,alength) | alength = alength (6:9) [duplicate literal removal 2872]
[SA] new: 2884. alength = sK0 | $less(sK0,alength) (6:6) [duplicate literal removal 2871]
[SA] new: 2885. alength = sK0 | $less(sK0,alength) (6:6) [duplicate literal removal 2870]
[SA] new: 2886. alength = sK0 | $less(sK0,alength) | sK0 = sK0 (6:9) [duplicate literal removal 2869]
[SA] new: 2887. alength = sK0 | $less(sK0,X4) | $less(X4,alength) | alength = X4 (6:12) [duplicate literal removal 2867]
[SA] new: 2888. alength = sK0 | $less(X3,alength) | $less(sK0,X3) | sK0 = X3 (6:12) [duplicate literal removal 2866]
[SA] new: 2889. alength = sK0 | $less(sK0,alength) (6:6) [duplicate literal removal 2865]
[SA] new: 2890. alength = sK0 | $less(sK0,alength) (6:6) [duplicate literal removal 2864]
[SA] new: 2891. alength = sK0 | $less(0,alength) | 0 = sK0 (6:9) [duplicate literal removal 2863]
[SA] new: 2892. alength = sK0 | $less(0,alength) | alength = 0 (6:9) [duplicate literal removal 2861]
[SA] new: 2893. alength = sK0 | $less(0,X2) | $less(X2,alength) | alength = X2 | 0 = sK0 (6:15) [duplicate literal removal 2860]
[SA] new: 2894. alength = sK0 | $less(0,X1) | $less(X1,alength) | 0 = sK0 | sK0 = X1 (6:15) [duplicate literal removal 2859]
[SA] active: 2484. i(l8(s(nl8))) = $sum(alength,1) (1:8:1) [backward demodulation 2476,118]
[SA] new: 2895. i(l8(X0)) != $sum(alength,1) | s(nl8) = X0 (2:11) [superposition 48,2484]
[SA] new: 2896. i(l8(X1)) != $sum(alength,1) | s(nl8) = X1 (2:11) [superposition 48,2484]
[SA] new: 2897. v = a(l11(s(nl8)),$sum(alength,1)) (2:9) [superposition 70,2484]
[SA] new: 2898. i(l8(s(s(nl8)))) = $sum($sum(alength,1),1) (2:11) [superposition 74,2484]
[SA] new: 2899. $sum(alength,1) != 0 | zero = s(nl8) (2:9) [superposition 108,2484]
[SA] new: 2900. 1 != $sum(alength,1) | s(zero) = s(nl8) (3:10) [superposition 121,2484]
[SA] new: 2901. $sum(alength,1) != 2 | s(nl8) = s(s(zero)) (4:11) [superposition 130,2484]
[SA] new: 2902. $sum(alength,1) != 3 | s(nl8) = s(s(s(zero))) (5:12) [superposition 183,2484]
[SA] new: 2903. v = a(l8(s(s(nl8))),$sum(alength,1)) (4:10) [superposition 349,2484]
[SA] new: 2904. alength != $sum(alength,1) | s(nl8) = nl8 (2:9) [superposition 2483,2484]
[SA] new: 2905. alength != $sum(3,-1) | s(nl8) = s(s(s(zero))) (5:12) [evaluation 2902]
[SA] new: 2906. alength != 2 | s(nl8) = s(s(s(zero))) (5:10) [evaluation 2905]
[SA] new: 2907. alength != 2 | s(s(zero)) = nl8 (5:8) [term algebras injectivity 2906]
[SA] new: 2908. alength != $sum(2,-1) | s(nl8) = s(s(zero)) (4:11) [evaluation 2901]
[SA] new: 2909. alength != 1 | s(nl8) = s(s(zero)) (4:9) [evaluation 2908]
[SA] new: 2910. alength != 1 | s(zero) = nl8 (4:7) [term algebras injectivity 2909]
[SA] new: 2911. alength != $sum(1,-1) | s(zero) = s(nl8) (3:10) [evaluation 2900]
[SA] new: 2912. alength != 0 | s(zero) = s(nl8) (3:8) [evaluation 2911]
[SA] new: 2913. alength != 0 | zero = nl8 (3:6) [term algebras injectivity 2912]
[SA] new: 2914. alength != $uminus(1) | zero = s(nl8) (2:8) [evaluation 2899]
[SA] new: 2915. alength != -1 | zero = s(nl8) (2:7) [evaluation 2914]
[SA] new: 2916. alength != -1 (2:3) [term algebras distinctness 2915]
[SA] active: 2841. v = a(l11(sK2(sK0)),sK0) (2:7:1) [superposition 70,2833]
[SA] new: 2917. v = a(l8(sK2(sK0)),sK0) | i(l8(sK2(sK0))) = sK0 (3:13) [superposition 2841,72]
[SA] new: 2918. v = a(l8(sK2(sK0)),sK0) | i(l8(sK2(sK0))) = sK0 (3:13) [superposition 72,2841]
[SA] active: 2771. $less(i(l8(sK2(X2))),alength) | ~$less(X2,alength) | $less(X2,0) (1:12:1) [resolution 2479,56]
[SA] new: 2919. ~$less(X0,alength) | $less(X0,0) | $less(i(l8(sK2(X0))),0) | i(l8(sK2(X0))) = i(l8(sK2(i(l8(sK2(X0)))))) (2:24) [resolution 2771,2480]
[SA] new: 2920. ~$less(X1,alength) | $less(X1,0) | ~$less(X2,i(l8(sK2(X1)))) | $less(X2,alength) (2:15) [resolution 2771,14]
[SA] new: 2921. ~$less(X3,alength) | $less(X3,0) | alength = X4 | $less(X4,alength) | $less(i(l8(sK2(X3))),X4) (2:18) [resolution 2771,99]
[SA] new: 2922. ~$less(X5,alength) | $less(X5,0) | alength = sK0 | $less(sK0,alength) | $less(i(l8(sK2(X5))),alength) (4:18) [resolution 2771,205]
[SA] new: 2923. ~$less(X6,alength) | $less(X6,0) | alength = alength | $less(sK0,alength) | $less(i(l8(sK2(X6))),alength) (4:18) [resolution 2771,254]
[SA] new: 2924. ~$less(X7,alength) | $less(X7,0) | alength = sK0 | 0 = sK0 | $less(0,alength) | $less(i(l8(sK2(X7))),sK0) (4:21) [resolution 2771,333]
[SA] new: 2925. $less(sK0,alength) | ~$less(sK0,alength) | $less(sK0,0) (2:9) [superposition 2771,2833]
[SA] passive: 2919. ~$less(X0,alength) | $less(X0,0) | $less(i(l8(sK2(X0))),0) | i(l8(sK2(X0))) = i(l8(sK2(i(l8(sK2(X0)))))) (2:24) [resolution 2771,2480]
[SA] passive: 2920. ~$less(X1,alength) | $less(X1,0) | ~$less(X2,i(l8(sK2(X1)))) | $less(X2,alength) (2:15) [resolution 2771,14]
[SA] passive: 2921. ~$less(X3,alength) | $less(X3,0) | alength = X4 | $less(X4,alength) | $less(i(l8(sK2(X3))),X4) (2:18) [resolution 2771,99]
[SA] active: 2843. 0 != sK0 | zero = sK2(sK0) (2:7:1) [superposition 108,2833]
[SA] active: 2655. 'Sub'(sK2(X0),nl8) | $less(alength,X0) | alength = X0 | $less(X0,0) (1:13:1) [forward demodulation 2485,2476]
[SA] new: 2926. $less(alength,X0) | alength = X0 | $less(X0,0) | a(main_end,X1) = a(l11(sK2(X0)),X1) | i(l8(sK1(X1,sK2(X0)))) = X1 (2:26) [resolution 2655,150]
[SA] new: 2927. $less(alength,X2) | alength = X2 | $less(X2,0) | $less(i(l8(sK2(X2))),alength) (2:15) [resolution 2655,56]
[SA] active: 2857. alength != sK0 | nl8 = sK2(sK0) (2:7:1) [superposition 2483,2833]
[SA] active: 2656. $less(alength,X0) | $less(X0,0) | i(l8(sK2(X0))) = X0 | alength = X0 (1:15:2) [forward demodulation 2486,2476]
[SA] new: 2928. $less(sK0,0) | i(l8(sK2(sK0))) = sK0 | alength = sK0 | alength = sK0 (6:15) [resolution 2656,2704]
[SA] new: 2929. $less(alength,0) | alength = i(l8(sK2(alength))) | alength = alength (2:12) [resolution 2656,13]
[SA] new: 2930. $less(X0,0) | i(l8(sK2(X0))) = X0 | alength = X0 | ~$less(X1,alength) | $less(X1,X0) (2:18) [resolution 2656,14]
[SA] new: 2931. $less(X2,0) | i(l8(sK2(X2))) = X2 | alength = X2 | X2 = X3 | $less(X3,X2) | $less(alength,X3) (2:21) [resolution 2656,99]
[SA] new: 2932. $less(X4,0) | i(l8(sK2(X4))) = X4 | alength = X4 | sK0 = X4 | $less(sK0,X4) | $less(alength,alength) (4:21) [resolution 2656,205]
[SA] new: 2933. $less(X5,0) | i(l8(sK2(X5))) = X5 | alength = X5 | alength = X5 | $less(sK0,X5) | $less(alength,alength) (4:21) [resolution 2656,254]
[SA] new: 2934. $less(X6,0) | i(l8(sK2(X6))) = X6 | alength = X6 | sK0 = X6 | 0 = sK0 | $less(0,X6) | $less(alength,sK0) (4:24) [resolution 2656,333]
[SA] new: 2935. $less(alength,0) | alength = i(l8(sK2(alength))) | alength = alength | $less(alength,0) | alength = i(l8(sK2(alength))) (2:21) [resolution 2656,2480]
[SA] new: 2936. $less(i(l8(p(nl8))),0) | i(l8(p(nl8))) = i(l8(sK2(i(l8(p(nl8)))))) | alength = i(l8(p(nl8))) | zero = nl8 | $less(alength,alength) (4:30) [resolution 2656,194]
[SA] new: 2937. $less(0,0) | i(l8(sK2(0))) = 0 | alength = 0 | sK0 = X7 | 0 = sK0 | $less(X7,sK0) | $less(alength,X7) (4:24) [resolution 2656,324]
[SA] new: 2938. $less(0,0) | i(l8(sK2(0))) = 0 | alength = 0 | alength = 0 | $less(alength,alength) (5:18) [resolution 2656,262]
[SA] new: 2939. $less(0,0) | i(l8(sK2(0))) = 0 | alength = 0 | 0 = sK0 | $less(alength,alength) (4:18) [resolution 2656,116]
[SA] new: 2940. $less(0,0) | i(l8(sK2(0))) = 0 | alength = 0 | 0 = sK0 | $less(alength,sK0) (3:18) [resolution 2656,114]
[SA] new: 2941. $less(sK0,0) | i(l8(sK2(sK0))) = sK0 | alength = sK0 | alength = X8 | $less(X8,alength) | $less(alength,X8) (4:21) [resolution 2656,247]
[SA] new: 2942. $less(sK0,0) | i(l8(sK2(sK0))) = sK0 | alength = sK0 | sK0 = X9 | $less(X9,alength) | $less(alength,X9) (4:21) [resolution 2656,199]
[SA] new: 2943. $less(sK0,0) | i(l8(sK2(sK0))) = sK0 | alength = sK0 | $less(alength,alength) (2:15) [resolution 2656,96]
[SA] new: 2944. $less(alength,sK0) | i(l8(sK2(sK0))) = sK0 | alength = sK0 (2:12) [resolution 2656,46]
[SA] new: 2945. $less(alength,X0) | i(l8(sK2(X0))) = X0 | alength = X0 | sK0 = X1 | 0 = sK0 | $less(X1,sK0) | $less(X0,X1) (4:24) [resolution 2656,324]
[SA] new: 2946. $less(alength,X2) | i(l8(sK2(X2))) = X2 | alength = X2 | alength = 0 | $less(X2,alength) (5:18) [resolution 2656,262]
[SA] new: 2947. $less(alength,X3) | i(l8(sK2(X3))) = X3 | alength = X3 | 0 = sK0 | $less(X3,alength) (4:18) [resolution 2656,116]
[SA] new: 2948. $less(alength,X4) | i(l8(sK2(X4))) = X4 | alength = X4 | 0 = sK0 | $less(X4,sK0) (3:18) [resolution 2656,114]
[SA] new: 2949. $less(alength,0) | i(l8(sK2(0))) = 0 | alength = 0 (2:12) [resolution 2656,13]
[SA] new: 2950. $less(alength,X5) | i(l8(sK2(X5))) = X5 | alength = X5 | ~$less(X6,X5) | $less(X6,0) (2:18) [resolution 2656,14]
[SA] new: 2951. $less(alength,X7) | i(l8(sK2(X7))) = X7 | alength = X7 | 0 = X8 | $less(X8,0) | $less(X7,X8) (2:21) [resolution 2656,99]
[SA] new: 2952. $less(alength,X9) | i(l8(sK2(X9))) = X9 | alength = X9 | 0 = sK0 | $less(sK0,0) | $less(X9,alength) (4:21) [resolution 2656,205]
[SA] new: 2953. $less(alength,X10) | i(l8(sK2(X10))) = X10 | alength = X10 | alength = 0 | $less(sK0,0) | $less(X10,alength) (4:21) [resolution 2656,254]
[SA] new: 2954. $less(alength,X11) | i(l8(sK2(X11))) = X11 | alength = X11 | 0 = sK0 | 0 = sK0 | $less(0,0) | $less(X11,sK0) (4:24) [resolution 2656,333]
[SA] new: 2955. $less(alength,X11) | i(l8(sK2(X11))) = X11 | alength = X11 | 0 = sK0 | $less(0,0) | $less(X11,sK0) (4:21) [duplicate literal removal 2954]
[SA] new: 2956. $less(alength,X11) | i(l8(sK2(X11))) = X11 | alength = X11 | 0 = sK0 | $less(X11,sK0) (4:18) [evaluation 2955]
[SA] new: 2957. i(l8(sK2(0))) = 0 | alength = 0 | 0 = sK0 | $less(alength,sK0) (3:15) [evaluation 2940]
[SA] new: 2958. i(l8(sK2(0))) = 0 | alength = 0 | 0 = sK0 | $less(alength,alength) (4:15) [evaluation 2939]
[SA] new: 2959. $less(0,0) | i(l8(sK2(0))) = 0 | alength = 0 | $less(alength,alength) (5:15) [duplicate literal removal 2938]
[SA] new: 2960. i(l8(sK2(0))) = 0 | alength = 0 | $less(alength,alength) (5:12) [evaluation 2959]
[SA] new: 2961. i(l8(sK2(0))) = 0 | alength = 0 | sK0 = X7 | 0 = sK0 | $less(X7,sK0) | $less(alength,X7) (4:21) [evaluation 2937]
[SA] new: 2962. $less(alength,0) | alength = i(l8(sK2(alength))) | alength = alength (2:12) [duplicate literal removal 2935]
[SA] new: 2963. $less(X5,0) | i(l8(sK2(X5))) = X5 | alength = X5 | $less(sK0,X5) | $less(alength,alength) (4:18) [duplicate literal removal 2933]
[SA] new: 2964. $less(sK0,0) | i(l8(sK2(sK0))) = sK0 | alength = sK0 (6:12) [duplicate literal removal 2928]
[SA] passive: 2945. $less(alength,X0) | i(l8(sK2(X0))) = X0 | alength = X0 | sK0 = X1 | 0 = sK0 | $less(X1,sK0) | $less(X0,X1) (4:24) [resolution 2656,324]
[SA] passive: 2948. $less(alength,X4) | i(l8(sK2(X4))) = X4 | alength = X4 | 0 = sK0 | $less(X4,sK0) (3:18) [resolution 2656,114]
[SA] passive: 2950. $less(alength,X5) | i(l8(sK2(X5))) = X5 | alength = X5 | ~$less(X6,X5) | $less(X6,0) (2:18) [resolution 2656,14]
[SA] passive: 2951. $less(alength,X7) | i(l8(sK2(X7))) = X7 | alength = X7 | 0 = X8 | $less(X8,0) | $less(X7,X8) (2:21) [resolution 2656,99]
[SA] active: 2496. 'Sub'(sK2(sK0),nl8) | alength = sK0 (3:7:1) [backward demodulation 2476,211]
[SA] new: 2965. alength = sK0 | a(main_end,X0) = a(l11(sK2(sK0)),X0) | i(l8(sK1(X0,sK2(sK0)))) = X0 (4:20) [resolution 2496,150]
[SA] new: 2966. alength = sK0 | $less(i(l8(sK2(sK0))),alength) (4:9) [resolution 2496,56]
[SA] new: 2967. $less(sK0,alength) | alength = sK0 (4:6) [forward demodulation 2966,2833]
[SA] active: 2488. alength != $sum(alength,1) | s(nl8) = nl8 (2:9:1) [backward demodulation 2476,155]
[SA] active: 2763. alength != 1 | s(zero) = nl8 (3:7:1) [superposition 121,2477]
[SA] active: 2491. v = a(l11(s(nl8)),$sum(alength,1)) (2:9:1) [backward demodulation 2476,160]
[SA] new: 2968. v = a(l8(s(nl8)),$sum(alength,1)) | i(l8(s(nl8))) = $sum(alength,1) (3:17) [superposition 2491,72]
[SA] new: 2969. v = a(l8(s(nl8)),$sum(alength,1)) | i(l8(s(nl8))) = $sum(alength,1) (3:17) [superposition 72,2491]
[SA] active: 2509. v = a(l8(s(nl8)),alength) (4:7:1) [backward demodulation 2476,354]
[SA] new: 2970. v = a(l8(nl8),alength) | alength = i(l8(nl8)) (5:11) [superposition 2509,440]
[SA] new: 2971. v = a(l11(nl8),alength) (5:6) [superposition 2509,57]
[SA] new: 2972. v = a(l8(nl8),alength) | alength = i(l8(nl8)) (5:11) [superposition 440,2509]
[SA] new: 2973. v = a(l11(nl8),alength) (5:6) [superposition 57,2509]
[SA] active: 2839. i(l8(X0)) != sK0 | sK2(sK0) = X0 (2:9:1) [superposition 48,2833]
[SA] new: 2974. 0 != sK0 | zero = sK2(sK0) (3:7) [superposition 2839,65]
[SA] new: 2975. 1 != sK0 | s(zero) = sK2(sK0) (3:8) [superposition 2839,119]
[SA] new: 2976. sK0 != 2 | s(s(zero)) = sK2(sK0) (3:9) [superposition 2839,124]
[SA] new: 2977. sK0 != 3 | s(s(s(zero))) = sK2(sK0) (4:10) [superposition 2839,133]
[SA] new: 2978. sK0 != 4 | s(s(s(s(zero)))) = sK2(sK0) (5:11) [superposition 2839,186]
[SA] new: 2979. sK0 != 5 | s(s(s(s(s(zero))))) = sK2(sK0) (6:12) [superposition 2839,312]
[SA] new: 2980. sK0 != 6 | s(s(s(s(s(s(zero)))))) = sK2(sK0) (7:13) [superposition 2839,640]
[SA] new: 2981. $sum(alength,1) != sK0 | s(nl8) = sK2(sK0) (3:10) [superposition 2839,2484]
[SA] new: 2982. alength != sK0 | nl8 = sK2(sK0) (3:7) [superposition 2839,2477]
[SA] new: 2983. sK0 != X0 | sK1(X0,p(nl8)) = sK2(sK0) | a(main_end,X0) = a(l11(p(nl8)),X0) | zero = nl8 (3:22) [superposition 2839,226]
[SA] passive: 2978. sK0 != 4 | s(s(s(s(zero)))) = sK2(sK0) (5:11) [superposition 2839,186]
[SA] passive: 2979. sK0 != 5 | s(s(s(s(s(zero))))) = sK2(sK0) (6:12) [superposition 2839,312]
[SA] passive: 2980. sK0 != 6 | s(s(s(s(s(s(zero)))))) = sK2(sK0) (7:13) [superposition 2839,640]
[SA] passive: 2981. $sum(alength,1) != sK0 | s(nl8) = sK2(sK0) (3:10) [superposition 2839,2484]
[SA] passive: 2983. sK0 != X0 | sK1(X0,p(nl8)) = sK2(sK0) | a(main_end,X0) = a(l11(p(nl8)),X0) | zero = nl8 (3:22) [superposition 2839,226]
[SA] active: 2487. alength != 2 | s(s(zero)) = nl8 (3:8:1) [backward demodulation 2476,154]
[SA] active: 2842. i(l8(s(sK2(sK0)))) = $sum(sK0,1) (2:9:1) [superposition 74,2833]
[SA] new: 2985. i(l8(X0)) != $sum(sK0,1) | s(sK2(sK0)) = X0 (3:12) [superposition 48,2842]
[SA] new: 2986. i(l8(X1)) != $sum(sK0,1) | s(sK2(sK0)) = X1 (3:12) [superposition 48,2842]
[SA] new: 2987. v = a(l11(s(sK2(sK0))),$sum(sK0,1)) (3:10) [superposition 70,2842]
[SA] new: 2988. i(l8(s(s(sK2(sK0))))) = $sum($sum(sK0,1),1) (3:12) [superposition 74,2842]
[SA] new: 2989. $sum(sK0,1) != 0 | zero = s(sK2(sK0)) (3:10) [superposition 108,2842]
[SA] new: 2990. 1 != $sum(sK0,1) | s(zero) = s(sK2(sK0)) (3:11) [superposition 121,2842]
[SA] new: 2991. $sum(sK0,1) != 2 | s(s(zero)) = s(sK2(sK0)) (4:12) [superposition 130,2842]
[SA] new: 2992. $sum(sK0,1) != 3 | s(sK2(sK0)) = s(s(s(zero))) (5:13) [superposition 183,2842]
[SA] new: 2993. v = a(l8(s(s(sK2(sK0)))),$sum(sK0,1)) (4:11) [superposition 349,2842]
[SA] new: 2994. alength != $sum(sK0,1) | s(sK2(sK0)) = nl8 (3:10) [superposition 2483,2842]
[SA] new: 2995. $sum(sK0,1) != sK0 | s(sK2(sK0)) = sK2(sK0) (3:11) [superposition 2839,2842]
[SA] new: 2996. $sum(3,-1) != sK0 | s(sK2(sK0)) = s(s(s(zero))) (5:13) [evaluation 2992]
[SA] new: 2997. sK0 != 2 | s(sK2(sK0)) = s(s(s(zero))) (5:11) [evaluation 2996]
[SA] new: 2998. sK0 != 2 | s(s(zero)) = sK2(sK0) (5:9) [term algebras injectivity 2997]
[SA] new: 2999. $sum(2,-1) != sK0 | s(s(zero)) = s(sK2(sK0)) (4:12) [evaluation 2991]
[SA] new: 3000. 1 != sK0 | s(s(zero)) = s(sK2(sK0)) (4:10) [evaluation 2999]
[SA] new: 3001. 1 != sK0 | s(zero) = sK2(sK0) (4:8) [term algebras injectivity 3000]
[SA] new: 3002. $sum(1,-1) != sK0 | s(zero) = s(sK2(sK0)) (3:11) [evaluation 2990]
[SA] new: 3003. 0 != sK0 | s(zero) = s(sK2(sK0)) (3:9) [evaluation 3002]
[SA] new: 3004. 0 != sK0 | zero = sK2(sK0) (3:7) [term algebras injectivity 3003]
[SA] new: 3005. $uminus(1) != sK0 | zero = s(sK2(sK0)) (3:9) [evaluation 2989]
[SA] new: 3006. sK0 != -1 | zero = s(sK2(sK0)) (3:8) [evaluation 3005]
[SA] new: 3007. sK0 != -1 (3:3) [term algebras distinctness 3006]
[SA] passive: 2985. i(l8(X0)) != $sum(sK0,1) | s(sK2(sK0)) = X0 (3:12) [superposition 48,2842]
[SA] passive: 2987. v = a(l11(s(sK2(sK0))),$sum(sK0,1)) (3:10) [superposition 70,2842]
[SA] passive: 2988. i(l8(s(s(sK2(sK0))))) = $sum($sum(sK0,1),1) (3:12) [superposition 74,2842]
[SA] passive: 3007. sK0 != -1 (3:3) [term algebras distinctness 3006]
[SA] passive: 2993. v = a(l8(s(s(sK2(sK0)))),$sum(sK0,1)) (4:11) [superposition 349,2842]
[SA] passive: 2994. alength != $sum(sK0,1) | s(sK2(sK0)) = nl8 (3:10) [superposition 2483,2842]
[SA] passive: 2995. $sum(sK0,1) != sK0 | s(sK2(sK0)) = sK2(sK0) (3:11) [superposition 2839,2842]
[SA] active: 3007. sK0 != -1 (3:3:1) [term algebras distinctness 3006]
[SA] active: 2489. i(l8(s(s(nl8)))) = $sum($sum(alength,1),1) (2:11:1) [backward demodulation 2476,157]
[SA] new: 3008. i(l8(X0)) != $sum($sum(alength,1),1) | s(s(nl8)) = X0 (3:14) [superposition 48,2489]
[SA] new: 3009. i(l8(X1)) != $sum($sum(alength,1),1) | s(s(nl8)) = X1 (3:14) [superposition 48,2489]
[SA] new: 3010. v = a(l11(s(s(nl8))),$sum($sum(alength,1),1)) (3:12) [superposition 70,2489]
[SA] new: 3011. i(l8(s(s(s(nl8))))) = $sum($sum($sum(alength,1),1),1) (3:14) [superposition 74,2489]
[SA] new: 3012. $sum($sum(alength,1),1) != 0 | zero = s(s(nl8)) (3:12) [superposition 108,2489]
[SA] new: 3013. 1 != $sum($sum(alength,1),1) | s(zero) = s(s(nl8)) (3:13) [superposition 121,2489]
[SA] new: 3014. $sum($sum(alength,1),1) != 2 | s(s(zero)) = s(s(nl8)) (4:14) [superposition 130,2489]
[SA] new: 3015. $sum($sum(alength,1),1) != 3 | s(s(nl8)) = s(s(s(zero))) (5:15) [superposition 183,2489]
[SA] new: 3016. v = a(l8(s(s(s(nl8)))),$sum($sum(alength,1),1)) (4:13) [superposition 349,2489]
[SA] new: 3017. alength != $sum($sum(alength,1),1) | s(s(nl8)) = nl8 (3:12) [superposition 2483,2489]
[SA] new: 3018. $sum($sum(alength,1),1) != sK0 | s(s(nl8)) = sK2(sK0) (3:13) [superposition 2839,2489]
[SA] new: 3019. alength != $sum($sum(3,-1),-1) | s(s(nl8)) = s(s(s(zero))) (5:15) [evaluation 3015]
[SA] new: 3020. alength != $sum(2,-1) | s(s(nl8)) = s(s(s(zero))) (5:13) [evaluation 3019]
[SA] new: 3021. alength != 1 | s(s(nl8)) = s(s(s(zero))) (5:11) [evaluation 3020]
[SA] new: 3022. alength != 1 | s(nl8) = s(s(zero)) (5:9) [term algebras injectivity 3021]
[SA] new: 3023. alength != 1 | s(zero) = nl8 (5:7) [term algebras injectivity 3022]
[SA] new: 3024. alength != $sum($sum(2,-1),-1) | s(s(zero)) = s(s(nl8)) (4:14) [evaluation 3014]
[SA] new: 3025. alength != $sum(1,-1) | s(s(zero)) = s(s(nl8)) (4:12) [evaluation 3024]
[SA] new: 3026. alength != 0 | s(s(zero)) = s(s(nl8)) (4:10) [evaluation 3025]
[SA] new: 3027. alength != 0 | s(zero) = s(nl8) (4:8) [term algebras injectivity 3026]
[SA] new: 3028. alength != 0 | zero = nl8 (4:6) [term algebras injectivity 3027]
[SA] new: 3029. alength != $sum($sum(1,-1),-1) | s(zero) = s(s(nl8)) (3:13) [evaluation 3013]
[SA] new: 3030. alength != $sum(0,-1) | s(zero) = s(s(nl8)) (3:11) [evaluation 3029]
[SA] new: 3031. alength != -1 | s(zero) = s(s(nl8)) (3:9) [evaluation 3030]
[SA] new: 3032. alength != -1 | zero = s(nl8) (3:7) [term algebras injectivity 3031]
[SA] new: 3033. alength != -1 (3:3) [term algebras distinctness 3032]
[SA] new: 3034. alength != $sum($uminus(1),-1) | zero = s(s(nl8)) (3:11) [evaluation 3012]
[SA] new: 3035. alength != $sum(-1,-1) | zero = s(s(nl8)) (3:10) [evaluation 3034]
[SA] new: 3036. alength != -2 | zero = s(s(nl8)) (3:8) [evaluation 3035]
[SA] new: 3037. alength != -2 (3:3) [term algebras distinctness 3036]
[SA] passive: 3008. i(l8(X0)) != $sum($sum(alength,1),1) | s(s(nl8)) = X0 (3:14) [superposition 48,2489]
[SA] passive: 3010. v = a(l11(s(s(nl8))),$sum($sum(alength,1),1)) (3:12) [superposition 70,2489]
[SA] passive: 3011. i(l8(s(s(s(nl8))))) = $sum($sum($sum(alength,1),1),1) (3:14) [superposition 74,2489]
[SA] passive: 3037. alength != -2 (3:3) [term algebras distinctness 3036]
[SA] passive: 3016. v = a(l8(s(s(s(nl8)))),$sum($sum(alength,1),1)) (4:13) [superposition 349,2489]
[SA] passive: 3017. alength != $sum($sum(alength,1),1) | s(s(nl8)) = nl8 (3:12) [superposition 2483,2489]
[SA] passive: 3018. $sum($sum(alength,1),1) != sK0 | s(s(nl8)) = sK2(sK0) (3:13) [superposition 2839,2489]
[SA] active: 3037. alength != -2 (3:3:1) [term algebras distinctness 3036]
[SA] active: 2490. i(l8(X0)) != $sum(alength,1) | s(nl8) = X0 (2:11:1) [backward demodulation 2476,158]
[SA] new: 3038. $sum(alength,1) != 0 | zero = s(nl8) (3:9) [superposition 2490,65]
[SA] new: 3039. 1 != $sum(alength,1) | s(zero) = s(nl8) (3:10) [superposition 2490,119]
[SA] new: 3040. $sum(alength,1) != 2 | s(nl8) = s(s(zero)) (3:11) [superposition 2490,124]
[SA] new: 3041. $sum(alength,1) != 3 | s(nl8) = s(s(s(zero))) (4:12) [superposition 2490,133]
[SA] new: 3042. $sum(alength,1) != 4 | s(nl8) = s(s(s(s(zero)))) (5:13) [superposition 2490,186]
[SA] new: 3043. $sum(alength,1) != 5 | s(nl8) = s(s(s(s(s(zero))))) (6:14) [superposition 2490,312]
[SA] new: 3044. $sum(alength,1) != 6 | s(nl8) = s(s(s(s(s(s(zero)))))) (7:15) [superposition 2490,640]
[SA] new: 3045. $sum(alength,1) != $sum($sum(alength,1),1) | s(nl8) = s(s(nl8)) (3:15) [superposition 2490,2489]
[SA] new: 3047. $sum(alength,1) != $sum(sK0,1) | s(nl8) = s(sK2(sK0)) (3:13) [superposition 2490,2842]
[SA] new: 3048. alength != $sum(alength,1) | s(nl8) = nl8 (3:9) [superposition 2490,2477]
[SA] new: 3049. $sum(alength,1) != X0 | s(nl8) = sK1(X0,p(nl8)) | a(main_end,X0) = a(l11(p(nl8)),X0) | zero = nl8 (3:24) [superposition 2490,226]
[SA] new: 3050. $sum(alength,1) != sK0 | s(nl8) = sK2(sK0) (3:10) [superposition 2490,2833]
[SA] new: 3051. $sum(alength,1) != $sum(sK0,1) | nl8 = sK2(sK0) (3:11) [term algebras injectivity 3047]
[SA] new: 3052. $sum(alength,1) != $sum($sum(alength,1),1) | s(nl8) = nl8 (3:13) [term algebras injectivity 3045]
[SA] new: 3053. alength != $sum(6,-1) | s(nl8) = s(s(s(s(s(s(zero)))))) (7:15) [evaluation 3044]
[SA] new: 3054. alength != 5 | s(nl8) = s(s(s(s(s(s(zero)))))) (7:13) [evaluation 3053]
[SA] new: 3055. alength != 5 | s(s(s(s(s(zero))))) = nl8 (7:11) [term algebras injectivity 3054]
[SA] new: 3056. alength != $sum(5,-1) | s(nl8) = s(s(s(s(s(zero))))) (6:14) [evaluation 3043]
[SA] new: 3057. alength != 4 | s(nl8) = s(s(s(s(s(zero))))) (6:12) [evaluation 3056]
[SA] new: 3058. alength != 4 | s(s(s(s(zero)))) = nl8 (6:10) [term algebras injectivity 3057]
[SA] new: 3059. alength != $sum(4,-1) | s(nl8) = s(s(s(s(zero)))) (5:13) [evaluation 3042]
[SA] new: 3060. alength != 3 | s(nl8) = s(s(s(s(zero)))) (5:11) [evaluation 3059]
[SA] new: 3061. alength != 3 | s(s(s(zero))) = nl8 (5:9) [term algebras injectivity 3060]
[SA] new: 3062. alength != $sum(3,-1) | s(nl8) = s(s(s(zero))) (4:12) [evaluation 3041]
[SA] new: 3063. alength != 2 | s(nl8) = s(s(s(zero))) (4:10) [evaluation 3062]
[SA] new: 3064. alength != 2 | s(s(zero)) = nl8 (4:8) [term algebras injectivity 3063]
[SA] new: 3065. alength != $sum(2,-1) | s(nl8) = s(s(zero)) (3:11) [evaluation 3040]
[SA] new: 3066. alength != 1 | s(nl8) = s(s(zero)) (3:9) [evaluation 3065]
[SA] new: 3067. alength != 1 | s(zero) = nl8 (3:7) [term algebras injectivity 3066]
[SA] new: 3068. alength != $sum(1,-1) | s(zero) = s(nl8) (3:10) [evaluation 3039]
[SA] new: 3069. alength != 0 | s(zero) = s(nl8) (3:8) [evaluation 3068]
[SA] new: 3070. alength != 0 | zero = nl8 (3:6) [term algebras injectivity 3069]
[SA] new: 3071. alength != $uminus(1) | zero = s(nl8) (3:8) [evaluation 3038]
[SA] new: 3072. alength != -1 | zero = s(nl8) (3:7) [evaluation 3071]
[SA] new: 3073. alength != -1 (3:3) [term algebras distinctness 3072]
[SA] passive: 3052. $sum(alength,1) != $sum($sum(alength,1),1) | s(nl8) = nl8 (3:13) [term algebras injectivity 3045]
[SA] passive: 3051. $sum(alength,1) != $sum(sK0,1) | nl8 = sK2(sK0) (3:11) [term algebras injectivity 3047]
[SA] active: 2844. 1 != sK0 | s(zero) = sK2(sK0) (3:8:1) [superposition 121,2833]
[SA] active: 2657. $less(i(l8(sK2(X0))),alength) | $less(alength,X0) | alength = X0 | $less(X0,0) (2:15:1) [forward demodulation 2493,2476]
[SA] new: 3074. $less(alength,X0) | alength = X0 | $less(X0,0) | $less(i(l8(sK2(X0))),0) | i(l8(sK2(X0))) = i(l8(sK2(i(l8(sK2(X0)))))) (3:27) [resolution 2657,2480]
[SA] new: 3075. $less(alength,X1) | alength = X1 | $less(X1,0) | ~$less(X2,i(l8(sK2(X1)))) | $less(X2,alength) (3:18) [resolution 2657,14]
[SA] new: 3076. $less(alength,X3) | alength = X3 | $less(X3,0) | alength = X4 | $less(X4,alength) | $less(i(l8(sK2(X3))),X4) (3:21) [resolution 2657,99]
[SA] new: 3077. $less(alength,X5) | alength = X5 | $less(X5,0) | alength = sK0 | $less(sK0,alength) | $less(i(l8(sK2(X5))),alength) (4:21) [resolution 2657,205]
[SA] new: 3078. $less(alength,X6) | alength = X6 | $less(X6,0) | alength = alength | $less(sK0,alength) | $less(i(l8(sK2(X6))),alength) (4:21) [resolution 2657,254]
[SA] new: 3079. $less(alength,X7) | alength = X7 | $less(X7,0) | alength = sK0 | 0 = sK0 | $less(0,alength) | $less(i(l8(sK2(X7))),sK0) (4:24) [resolution 2657,333]
[SA] new: 3080. $less(sK0,alength) | $less(alength,sK0) | alength = sK0 | $less(sK0,0) (3:12) [superposition 2657,2833]
[SA] passive: 3074. $less(alength,X0) | alength = X0 | $less(X0,0) | $less(i(l8(sK2(X0))),0) | i(l8(sK2(X0))) = i(l8(sK2(i(l8(sK2(X0)))))) (3:27) [resolution 2657,2480]
[SA] active: 2847. v = a(l8(s(sK2(sK0))),sK0) (4:8:1) [superposition 349,2833]
[SA] new: 3081. v = a(l8(sK2(sK0)),sK0) | i(l8(sK2(sK0))) = sK0 (5:13) [superposition 2847,440]
[SA] new: 3082. v = a(l11(sK2(sK0)),sK0) (5:7) [superposition 2847,57]
[SA] new: 3083. v = a(l8(sK2(sK0)),sK0) | i(l8(sK2(sK0))) = sK0 (5:13) [superposition 440,2847]
[SA] new: 3084. v = a(l11(sK2(sK0)),sK0) (5:7) [superposition 57,2847]
[SA] active: 2920. ~$less(X2,i(l8(sK2(X1)))) | $less(X1,0) | ~$less(X1,alength) | $less(X2,alength) (2:15:1) [resolution 2771,14]
[SA] new: 3085. $less(X0,0) | ~$less(X0,alength) | $less(X1,alength) | $less(i(l8(sK2(X0))),X1) | i(l8(sK2(X0))) = X1 (3:21) [resolution 2920,15]
[SA] new: 3086. $less(X2,0) | ~$less(X2,alength) | $less(X3,alength) | $less(i(l8(sK2(X2))),X3) | i(l8(sK2(X2))) = X3 (3:21) [resolution 2920,15]
[SA] new: 3087. $less(X4,0) | ~$less(X4,alength) | $less(X5,alength) | $less(X6,X5) | $less(i(l8(sK2(X4))),X6) | i(l8(sK2(X4))) = X5 | i(l8(sK2(X4))) = X6 (3:30) [resolution 2920,174]
[SA] new: 3088. $less(X7,0) | ~$less(X7,alength) | $less(X8,alength) | $less(i(l8(sK2(X7))),X9) | $less(X9,X8) | X8 = X9 | i(l8(sK2(X7))) = X8 (3:27) [resolution 2920,174]
[SA] new: 3089. $less(X10,0) | ~$less(X10,alength) | $less(X11,alength) | $less(X12,X11) | $less(i(l8(sK2(X10))),X12) | i(l8(sK2(X10))) = X12 | X11 = X12 (3:27) [resolution 2920,174]
[SA] new: 3090. $less(X13,0) | ~$less(X13,alength) | $less(X14,alength) | $less(sK0,X14) | $less(i(l8(sK2(X13))),alength) | i(l8(sK2(X13))) = X14 | i(l8(sK2(X13))) = sK0 (4:30) [resolution 2920,198]
[SA] new: 3091. $less(X15,0) | ~$less(X15,alength) | $less(X16,alength) | $less(sK0,X16) | $less(i(l8(sK2(X15))),alength) | alength = i(l8(sK2(X15))) | sK0 = X16 (4:27) [resolution 2920,204]
[SA] new: 3092. $less(X17,0) | ~$less(X17,alength) | $less(X18,alength) | $less(sK0,X18) | $less(i(l8(sK2(X17))),alength) | i(l8(sK2(X17))) = X18 | alength = i(l8(sK2(X17))) (4:30) [resolution 2920,246]
[SA] new: 3093. $less(X19,0) | ~$less(X19,alength) | $less(X20,alength) | $less(sK0,X20) | $less(i(l8(sK2(X19))),alength) | alength = i(l8(sK2(X19))) | alength = X20 (4:27) [resolution 2920,253]
[SA] new: 3094. $less(X21,0) | ~$less(X21,alength) | $less(X22,alength) | $less(0,X22) | $less(i(l8(sK2(X21))),sK0) | i(l8(sK2(X21))) = X22 | 0 = sK0 | i(l8(sK2(X21))) = sK0 (4:33) [resolution 2920,323]
[SA] new: 3095. $less(X23,0) | ~$less(X23,alength) | $less(X24,alength) | $less(0,X24) | $less(i(l8(sK2(X23))),sK0) | i(l8(sK2(X23))) = sK0 | 0 = sK0 | sK0 = X24 (4:30) [resolution 2920,332]
[SA] new: 3096. $less(X25,0) | ~$less(X25,alength) | $less(X26,alength) | $less(0,X26) | $less(i(l8(sK2(X25))),sK0) | i(l8(sK2(X25))) = X26 | 0 = sK0 | 0 = X26 (4:30) [resolution 2920,387]
[SA] new: 3097. $less(X27,0) | ~$less(X27,alength) | $less(X28,alength) | $less(sK0,X28) | $less(i(l8(sK2(X27))),alength) | i(l8(sK2(X27))) = X28 | sK0 = X28 (4:27) [resolution 2920,388]
[SA] new: 3098. $less(X29,0) | ~$less(X29,alength) | $less(X30,alength) | $less(sK0,X30) | $less(i(l8(sK2(X29))),alength) | sK0 = X30 | i(l8(sK2(X29))) = sK0 (4:27) [resolution 2920,408]
[SA] new: 3099. $less(X31,0) | ~$less(X31,alength) | $less(alength,alength) | $less(i(l8(sK2(X31))),0) | i(l8(sK2(X31))) = i(l8(sK2(i(l8(sK2(X31)))))) | alength = i(l8(sK2(X31))) (3:33) [resolution 2920,2656]
[SA] new: 3100. $less(X32,0) | ~$less(X32,alength) | $less(i(l8(p(nl8))),alength) | alength = i(l8(sK2(X32))) | $less(i(l8(sK2(X32))),alength) | zero = nl8 (4:27) [resolution 2920,193]
[SA] new: 3101. $less(X33,0) | ~$less(X33,alength) | $less(0,alength) | $less(i(l8(sK2(X33))),sK0) | i(l8(sK2(X33))) = sK0 | 0 = sK0 (3:24) [resolution 2920,177]
[SA] new: 3102. $less(X34,0) | ~$less(X34,alength) | $less(0,alength) | $less(i(l8(sK2(X34))),alength) | alength = i(l8(sK2(X34))) | 0 = sK0 (4:24) [resolution 2920,178]
[SA] new: 3103. $less(X35,0) | ~$less(X35,alength) | $less(0,alength) | $less(i(l8(sK2(X35))),sK0) | 0 = sK0 | i(l8(sK2(X35))) = 0 (4:24) [resolution 2920,229]
[SA] new: 3104. $less(X36,0) | ~$less(X36,alength) | $less(0,alength) | $less(i(l8(sK2(X36))),alength) | alength = i(l8(sK2(X36))) | alength = 0 (5:24) [resolution 2920,261]
[SA] new: 3105. $less(X37,0) | ~$less(X37,alength) | $less(0,alength) | $less(i(l8(sK2(X37))),alength) | 0 = sK0 | i(l8(sK2(X37))) = 0 (5:24) [resolution 2920,278]
[SA] new: 3106. $less(X38,0) | ~$less(X38,alength) | $less(0,alength) | $less(i(l8(sK2(X38))),X39) | $less(X39,sK0) | i(l8(sK2(X38))) = X39 | 0 = sK0 | sK0 = X39 (4:30) [resolution 2920,323]
[SA] new: 3107. $less(X40,0) | ~$less(X40,alength) | $less(0,alength) | $less(i(l8(sK2(X40))),alength) | 0 = sK0 | i(l8(sK2(X40))) = sK0 (4:24) [resolution 2920,330]
[SA] new: 3108. $less(X41,0) | ~$less(X41,alength) | $less(0,alength) | $less(X42,sK0) | $less(i(l8(sK2(X41))),X42) | sK0 = X42 | 0 = sK0 | i(l8(sK2(X41))) = sK0 (4:30) [resolution 2920,332]
[SA] new: 3109. $less(X43,0) | ~$less(X43,alength) | $less(0,alength) | $less(i(l8(sK2(X43))),sK0) | i(l8(sK2(X43))) = 0 | i(l8(sK2(X43))) = sK0 (4:27) [resolution 2920,380]
[SA] new: 3110. $less(X44,0) | ~$less(X44,alength) | $less(0,alength) | $less(X45,sK0) | $less(i(l8(sK2(X44))),X45) | i(l8(sK2(X44))) = X45 | 0 = sK0 | i(l8(sK2(X44))) = 0 (4:33) [resolution 2920,387]
[SA] new: 3111. $less(X46,0) | ~$less(X46,alength) | $less(sK0,alength) | $less(i(l8(sK2(X46))),alength) | i(l8(sK2(X46))) = sK0 (3:21) [resolution 2920,106]
[SA] new: 3112. $less(X47,0) | ~$less(X47,alength) | $less(sK0,alength) | $less(i(l8(sK2(X47))),alength) | alength = i(l8(sK2(X47))) (3:21) [resolution 2920,179]
[SA] new: 3113. $less(X48,0) | ~$less(X48,alength) | $less(sK0,alength) | $less(i(l8(sK2(X48))),X49) | $less(X49,alength) | i(l8(sK2(X48))) = X49 | sK0 = X49 (4:27) [resolution 2920,198]
[SA] new: 3114. $less(X50,0) | ~$less(X50,alength) | $less(sK0,alength) | $less(X51,alength) | $less(i(l8(sK2(X50))),X51) | alength = X51 | i(l8(sK2(X50))) = sK0 (4:27) [resolution 2920,204]
[SA] new: 3115. $less(X52,0) | ~$less(X52,alength) | $less(sK0,alength) | $less(i(l8(sK2(X52))),X53) | $less(X53,alength) | i(l8(sK2(X52))) = X53 | alength = X53 (4:27) [resolution 2920,246]
[SA] new: 3116. $less(X54,0) | ~$less(X54,alength) | $less(sK0,alength) | $less(X55,alength) | $less(i(l8(sK2(X54))),X55) | alength = X55 | alength = i(l8(sK2(X54))) (4:27) [resolution 2920,253]
[SA] new: 3117. $less(X56,0) | ~$less(X56,alength) | $less(sK0,alength) | $less(X57,alength) | $less(i(l8(sK2(X56))),X57) | i(l8(sK2(X56))) = X57 | i(l8(sK2(X56))) = sK0 (4:30) [resolution 2920,388]
[SA] new: 3118. $less(X58,0) | ~$less(X58,alength) | $less(sK0,alength) | $less(X59,alength) | $less(i(l8(sK2(X58))),X59) | i(l8(sK2(X58))) = sK0 | sK0 = X59 (4:27) [resolution 2920,408]
[SA] new: 3119. ~$less(X0,sK0) | $less(sK0,0) | ~$less(sK0,alength) | $less(X0,alength) (3:12) [superposition 2920,2833]
[SA] passive: 3085. $less(X0,0) | ~$less(X0,alength) | $less(X1,alength) | $less(i(l8(sK2(X0))),X1) | i(l8(sK2(X0))) = X1 (3:21) [resolution 2920,15]
[SA] passive: 3087. $less(X4,0) | ~$less(X4,alength) | $less(X5,alength) | $less(X6,X5) | $less(i(l8(sK2(X4))),X6) | i(l8(sK2(X4))) = X5 | i(l8(sK2(X4))) = X6 (3:30) [resolution 2920,174]
[SA] passive: 3088. $less(X7,0) | ~$less(X7,alength) | $less(X8,alength) | $less(i(l8(sK2(X7))),X9) | $less(X9,X8) | X8 = X9 | i(l8(sK2(X7))) = X8 (3:27) [resolution 2920,174]
[SA] passive: 3089. $less(X10,0) | ~$less(X10,alength) | $less(X11,alength) | $less(X12,X11) | $less(i(l8(sK2(X10))),X12) | i(l8(sK2(X10))) = X12 | X11 = X12 (3:27) [resolution 2920,174]
[SA] passive: 3094. $less(X21,0) | ~$less(X21,alength) | $less(X22,alength) | $less(0,X22) | $less(i(l8(sK2(X21))),sK0) | i(l8(sK2(X21))) = X22 | 0 = sK0 | i(l8(sK2(X21))) = sK0 (4:33) [resolution 2920,323]
[SA] passive: 3109. $less(X43,0) | ~$less(X43,alength) | $less(0,alength) | $less(i(l8(sK2(X43))),sK0) | i(l8(sK2(X43))) = 0 | i(l8(sK2(X43))) = sK0 (4:27) [resolution 2920,380]
[SA] active: 2495. i(l8(sK2(sK0))) = sK0 | alength = sK0 (3:9:1) [backward demodulation 2476,209]
[SA] new: 3120. ~$less(X0,sK0) | $less(sK0,0) | ~$less(sK0,alength) | $less(X0,alength) | alength = sK0 (4:15) [superposition 2920,2495]
[SA] new: 3121. $less(sK0,alength) | $less(alength,sK0) | alength = sK0 | $less(sK0,0) | alength = sK0 (4:15) [superposition 2657,2495]
[SA] new: 3122. $less(sK0,alength) | ~$less(sK0,alength) | $less(sK0,0) | alength = sK0 (4:12) [superposition 2771,2495]
[SA] new: 3123. i(l8(X1)) != sK0 | sK2(sK0) = X1 | alength = sK0 (4:12) [superposition 48,2495]
[SA] new: 3124. i(l8(X2)) != sK0 | sK2(sK0) = X2 | alength = sK0 (4:12) [superposition 48,2495]
[SA] new: 3125. v = a(l11(sK2(sK0)),sK0) | alength = sK0 (4:10) [superposition 70,2495]
[SA] new: 3126. i(l8(s(sK2(sK0)))) = $sum(sK0,1) | alength = sK0 (4:12) [superposition 74,2495]
[SA] new: 3127. 0 != sK0 | zero = sK2(sK0) | alength = sK0 (4:10) [superposition 108,2495]
[SA] new: 3128. 1 != sK0 | s(zero) = sK2(sK0) | alength = sK0 (4:11) [superposition 121,2495]
[SA] new: 3129. sK0 != 2 | s(s(zero)) = sK2(sK0) | alength = sK0 (4:12) [superposition 130,2495]
[SA] new: 3130. sK0 != 3 | s(s(s(zero))) = sK2(sK0) | alength = sK0 (5:13) [superposition 183,2495]
[SA] new: 3131. v = a(l8(s(sK2(sK0))),sK0) | alength = sK0 (4:11) [superposition 349,2495]
[SA] new: 3132. alength != sK0 | nl8 = sK2(sK0) | alength = sK0 (4:10) [superposition 2483,2495]
[SA] new: 3133. $sum(alength,1) != sK0 | s(nl8) = sK2(sK0) | alength = sK0 (4:13) [superposition 2490,2495]
[SA] new: 3135. $less(sK0,alength) | $less(alength,sK0) | alength = sK0 | $less(sK0,0) (4:12) [duplicate literal removal 3121]
[SA] active: 2660. ~$less(X3,alength) | alength = X2 | $less(X2,0) | i(l8(sK2(X2))) = X2 | $less(X3,X2) (2:18:1) [forward demodulation 2498,2476]
[SA] new: 3136. alength = X0 | $less(X0,0) | i(l8(sK2(X0))) = X0 | $less(X1,X0) | $less(sK0,X1) | sK0 = X1 (3:21) [resolution 2660,106]
[SA] new: 3137. alength = X2 | $less(X2,0) | i(l8(sK2(X2))) = X2 | $less(X3,X2) | $less(0,X3) | alength = X3 | 0 = sK0 (4:24) [resolution 2660,178]
[SA] new: 3138. alength = X4 | $less(X4,0) | i(l8(sK2(X4))) = X4 | $less(X5,X4) | $less(sK0,X5) | alength = X5 (3:21) [resolution 2660,179]
[SA] new: 3139. alength = X6 | $less(X6,0) | i(l8(sK2(X6))) = X6 | $less(X7,X6) | $less(sK0,X8) | $less(X8,X7) | X7 = X8 | sK0 = X7 (4:27) [resolution 2660,198]
[SA] new: 3140. alength = X9 | $less(X9,0) | i(l8(sK2(X9))) = X9 | $less(X10,X9) | $less(sK0,X11) | $less(X11,X10) | alength = X10 | sK0 = X11 (4:27) [resolution 2660,204]
[SA] new: 3141. alength = X12 | $less(X12,0) | i(l8(sK2(X12))) = X12 | $less(X13,X12) | $less(sK0,X14) | $less(X14,X13) | X13 = X14 | alength = X13 (4:27) [resolution 2660,246]
[SA] new: 3142. alength = X15 | $less(X15,0) | i(l8(sK2(X15))) = X15 | $less(X16,X15) | $less(sK0,X17) | $less(X17,X16) | alength = X16 | alength = X17 (4:27) [resolution 2660,253]
[SA] new: 3143. alength = X18 | $less(X18,0) | i(l8(sK2(X18))) = X18 | $less(X19,X18) | $less(0,X19) | alength = X19 | alength = 0 (5:24) [resolution 2660,261]
[SA] new: 3144. alength = X20 | $less(X20,0) | i(l8(sK2(X20))) = X20 | $less(X21,X20) | $less(0,X21) | 0 = sK0 | 0 = X21 (5:24) [resolution 2660,278]
[SA] new: 3145. alength = X22 | $less(X22,0) | i(l8(sK2(X22))) = X22 | $less(X23,X22) | $less(0,X23) | 0 = sK0 | sK0 = X23 (4:24) [resolution 2660,330]
[SA] new: 3146. alength = X24 | $less(X24,0) | i(l8(sK2(X24))) = X24 | $less(X25,X24) | $less(sK0,X26) | $less(X26,X25) | X25 = X26 | sK0 = X26 (4:27) [resolution 2660,388]
[SA] new: 3147. alength = X27 | $less(X27,0) | i(l8(sK2(X27))) = X27 | $less(X28,X27) | $less(sK0,X29) | $less(X29,X28) | sK0 = X29 | sK0 = X28 (4:27) [resolution 2660,408]
[SA] new: 3148. alength = X30 | $less(X30,0) | i(l8(sK2(X30))) = X30 | $less(X31,X30) | $less(alength,X31) | alength = X31 (3:21) [resolution 2660,15]
[SA] new: 3149. alength = X32 | $less(X32,0) | i(l8(sK2(X32))) = X32 | $less(X33,X32) | $less(alength,X33) | alength = X33 (3:21) [resolution 2660,15]
[SA] new: 3150. alength = X34 | $less(X34,0) | i(l8(sK2(X34))) = X34 | $less(X35,X34) | $less(X36,X35) | $less(alength,X36) | alength = X35 | alength = X36 (3:27) [resolution 2660,174]
[SA] new: 3151. alength = X37 | $less(X37,0) | i(l8(sK2(X37))) = X37 | $less(X38,X37) | $less(alength,X39) | $less(X39,X38) | X38 = X39 | alength = X38 (3:27) [resolution 2660,174]
[SA] new: 3152. alength = X40 | $less(X40,0) | i(l8(sK2(X40))) = X40 | $less(X41,X40) | $less(X42,X41) | $less(alength,X42) | alength = X42 | X41 = X42 (3:27) [resolution 2660,174]
[SA] new: 3153. alength = X43 | $less(X43,0) | i(l8(sK2(X43))) = X43 | $less(X44,X43) | $less(sK0,X44) | $less(alength,alength) | alength = X44 | alength = sK0 (4:27) [resolution 2660,198]
[SA] new: 3154. alength = X45 | $less(X45,0) | i(l8(sK2(X45))) = X45 | $less(X46,X45) | $less(sK0,X46) | $less(alength,alength) | alength = alength | sK0 = X46 (4:27) [resolution 2660,204]
[SA] new: 3155. alength = X47 | $less(X47,0) | i(l8(sK2(X47))) = X47 | $less(X48,X47) | $less(sK0,X48) | $less(alength,alength) | alength = X48 | alength = alength (4:27) [resolution 2660,246]
[SA] new: 3156. alength = X49 | $less(X49,0) | i(l8(sK2(X49))) = X49 | $less(X50,X49) | $less(sK0,X50) | $less(alength,alength) | alength = alength | alength = X50 (4:27) [resolution 2660,253]
[SA] new: 3157. alength = X51 | $less(X51,0) | i(l8(sK2(X51))) = X51 | $less(X52,X51) | $less(0,X52) | $less(alength,sK0) | alength = X52 | 0 = sK0 | alength = sK0 (4:30) [resolution 2660,323]
[SA] new: 3158. alength = X53 | $less(X53,0) | i(l8(sK2(X53))) = X53 | $less(X54,X53) | $less(0,X54) | $less(alength,sK0) | alength = sK0 | 0 = sK0 | sK0 = X54 (4:30) [resolution 2660,332]
[SA] new: 3159. alength = X55 | $less(X55,0) | i(l8(sK2(X55))) = X55 | $less(X56,X55) | $less(0,X56) | $less(alength,sK0) | alength = X56 | 0 = sK0 | 0 = X56 (4:30) [resolution 2660,387]
[SA] new: 3160. alength = X57 | $less(X57,0) | i(l8(sK2(X57))) = X57 | $less(X58,X57) | $less(sK0,X58) | $less(alength,alength) | alength = X58 | sK0 = X58 (4:27) [resolution 2660,388]
[SA] new: 3161. alength = X59 | $less(X59,0) | i(l8(sK2(X59))) = X59 | $less(X60,X59) | $less(sK0,X60) | $less(alength,alength) | sK0 = X60 | alength = sK0 (4:27) [resolution 2660,408]
[SA] new: 3162. alength = X61 | $less(X61,0) | i(l8(sK2(X61))) = X61 | $less(alength,X61) | $less(alength,0) | alength = i(l8(sK2(alength))) | alength = alength (3:27) [resolution 2660,2656]
[SA] new: 3163. alength = X62 | $less(X62,0) | i(l8(sK2(X62))) = X62 | $less(i(l8(p(nl8))),X62) | zero = nl8 (3:21) [resolution 2660,95]
[SA] new: 3164. alength = X63 | $less(X63,0) | i(l8(sK2(X63))) = X63 | $less(i(l8(p(p(nl8)))),X63) | zero = p(nl8) | zero = nl8 (4:26) [resolution 2660,314]
[SA] new: 3165. alength = X64 | $less(X64,0) | i(l8(sK2(X64))) = X64 | $less(i(l8(sK2(X65))),X64) | $less(alength,X65) | alength = X65 | $less(X65,0) (3:27) [resolution 2660,2657]
[SA] new: 3166. alength = X66 | $less(X66,0) | i(l8(sK2(X66))) = X66 | $less(i(l8(sK2(X67))),X66) | ~$less(X67,alength) | $less(X67,0) (3:24) [resolution 2660,2771]
[SA] new: 3167. alength = X68 | $less(X68,0) | i(l8(sK2(X68))) = X68 | $less(i(l8(p(nl8))),X68) | alength = alength | $less(alength,alength) | zero = nl8 (4:27) [resolution 2660,193]
[SA] new: 3168. alength = X69 | $less(X69,0) | i(l8(sK2(X69))) = X69 | $less(0,X69) | alength = 0 (4:18) [resolution 2660,244]
[SA] new: 3169. alength = X70 | $less(X70,0) | i(l8(sK2(X70))) = X70 | $less(0,X70) | 0 = sK0 (3:18) [resolution 2660,113]
[SA] new: 3170. alength = X71 | $less(X71,0) | i(l8(sK2(X71))) = X71 | $less(0,X71) | $less(alength,sK0) | alength = sK0 | 0 = sK0 (3:24) [resolution 2660,177]
[SA] new: 3171. alength = X72 | $less(X72,0) | i(l8(sK2(X72))) = X72 | $less(0,X72) | $less(alength,alength) | alength = alength | 0 = sK0 (4:24) [resolution 2660,178]
[SA] new: 3172. alength = X73 | $less(X73,0) | i(l8(sK2(X73))) = X73 | $less(0,X73) | $less(alength,sK0) | 0 = sK0 | alength = 0 (4:24) [resolution 2660,229]
[SA] new: 3173. alength = X74 | $less(X74,0) | i(l8(sK2(X74))) = X74 | $less(0,X74) | $less(alength,alength) | alength = alength | alength = 0 (5:24) [resolution 2660,261]
[SA] new: 3174. alength = X75 | $less(X75,0) | i(l8(sK2(X75))) = X75 | $less(0,X75) | $less(alength,alength) | 0 = sK0 | alength = 0 (5:24) [resolution 2660,278]
[SA] new: 3175. alength = X76 | $less(X76,0) | i(l8(sK2(X76))) = X76 | $less(0,X76) | $less(alength,X77) | $less(X77,sK0) | alength = X77 | 0 = sK0 | sK0 = X77 (4:30) [resolution 2660,323]
[SA] new: 3176. alength = X78 | $less(X78,0) | i(l8(sK2(X78))) = X78 | $less(0,X78) | $less(alength,alength) | 0 = sK0 | alength = sK0 (4:24) [resolution 2660,330]
[SA] new: 3177. alength = X79 | $less(X79,0) | i(l8(sK2(X79))) = X79 | $less(0,X79) | $less(X80,sK0) | $less(alength,X80) | sK0 = X80 | 0 = sK0 | alength = sK0 (4:30) [resolution 2660,332]
[SA] new: 3178. alength = X81 | $less(X81,0) | i(l8(sK2(X81))) = X81 | $less(0,X81) | $less(alength,sK0) | alength = 0 | alength = sK0 (4:24) [resolution 2660,380]
[SA] new: 3179. alength = X82 | $less(X82,0) | i(l8(sK2(X82))) = X82 | $less(0,X82) | $less(X83,sK0) | $less(alength,X83) | alength = X83 | 0 = sK0 | alength = 0 (4:30) [resolution 2660,387]
[SA] new: 3180. alength = X84 | $less(X84,0) | i(l8(sK2(X84))) = X84 | $less(sK0,X84) (3:15) [resolution 2660,45]
[SA] new: 3181. alength = X85 | $less(X85,0) | i(l8(sK2(X85))) = X85 | $less(sK0,X85) | $less(alength,alength) | alength = sK0 (3:21) [resolution 2660,106]
[SA] new: 3182. alength = X86 | $less(X86,0) | i(l8(sK2(X86))) = X86 | $less(sK0,X86) | $less(alength,alength) | alength = alength (3:21) [resolution 2660,179]
[SA] new: 3183. alength = X87 | $less(X87,0) | i(l8(sK2(X87))) = X87 | $less(sK0,X87) | $less(alength,X88) | $less(X88,alength) | alength = X88 | sK0 = X88 (4:27) [resolution 2660,198]
[SA] new: 3184. alength = X89 | $less(X89,0) | i(l8(sK2(X89))) = X89 | $less(sK0,X89) | $less(X90,alength) | $less(alength,X90) | alength = X90 | alength = sK0 (4:27) [resolution 2660,204]
[SA] new: 3185. alength = X91 | $less(X91,0) | i(l8(sK2(X91))) = X91 | $less(sK0,X91) | $less(alength,X92) | $less(X92,alength) | alength = X92 | alength = X92 (4:27) [resolution 2660,246]
[SA] new: 3186. alength = X93 | $less(X93,0) | i(l8(sK2(X93))) = X93 | $less(sK0,X93) | $less(X94,alength) | $less(alength,X94) | alength = X94 | alength = alength (4:27) [resolution 2660,253]
[SA] new: 3187. alength = X95 | $less(X95,0) | i(l8(sK2(X95))) = X95 | $less(sK0,X95) | $less(X96,alength) | $less(alength,X96) | alength = X96 | alength = sK0 (4:27) [resolution 2660,388]
[SA] new: 3188. alength = X97 | $less(X97,0) | i(l8(sK2(X97))) = X97 | $less(sK0,X97) | $less(X98,alength) | $less(alength,X98) | alength = sK0 | sK0 = X98 (4:27) [resolution 2660,408]
[SA] new: 3189. alength = X91 | $less(X91,0) | i(l8(sK2(X91))) = X91 | $less(sK0,X91) | $less(alength,X92) | $less(X92,alength) | alength = X92 (4:24) [duplicate literal removal 3185]
[SA] passive: 3137. alength = X2 | $less(X2,0) | i(l8(sK2(X2))) = X2 | $less(X3,X2) | $less(0,X3) | alength = X3 | 0 = sK0 (4:24) [resolution 2660,178]
[SA] passive: 3139. alength = X6 | $less(X6,0) | i(l8(sK2(X6))) = X6 | $less(X7,X6) | $less(sK0,X8) | $less(X8,X7) | X7 = X8 | sK0 = X7 (4:27) [resolution 2660,198]
[SA] passive: 3140. alength = X9 | $less(X9,0) | i(l8(sK2(X9))) = X9 | $less(X10,X9) | $less(sK0,X11) | $less(X11,X10) | alength = X10 | sK0 = X11 (4:27) [resolution 2660,204]
[SA] passive: 3141. alength = X12 | $less(X12,0) | i(l8(sK2(X12))) = X12 | $less(X13,X12) | $less(sK0,X14) | $less(X14,X13) | X13 = X14 | alength = X13 (4:27) [resolution 2660,246]
[SA] passive: 3142. alength = X15 | $less(X15,0) | i(l8(sK2(X15))) = X15 | $less(X16,X15) | $less(sK0,X17) | $less(X17,X16) | alength = X16 | alength = X17 (4:27) [resolution 2660,253]
[SA] passive: 3143. alength = X18 | $less(X18,0) | i(l8(sK2(X18))) = X18 | $less(X19,X18) | $less(0,X19) | alength = X19 | alength = 0 (5:24) [resolution 2660,261]
[SA] passive: 3144. alength = X20 | $less(X20,0) | i(l8(sK2(X20))) = X20 | $less(X21,X20) | $less(0,X21) | 0 = sK0 | 0 = X21 (5:24) [resolution 2660,278]
[SA] passive: 3145. alength = X22 | $less(X22,0) | i(l8(sK2(X22))) = X22 | $less(X23,X22) | $less(0,X23) | 0 = sK0 | sK0 = X23 (4:24) [resolution 2660,330]
[SA] passive: 3146. alength = X24 | $less(X24,0) | i(l8(sK2(X24))) = X24 | $less(X25,X24) | $less(sK0,X26) | $less(X26,X25) | X25 = X26 | sK0 = X26 (4:27) [resolution 2660,388]
[SA] passive: 3147. alength = X27 | $less(X27,0) | i(l8(sK2(X27))) = X27 | $less(X28,X27) | $less(sK0,X29) | $less(X29,X28) | sK0 = X29 | sK0 = X28 (4:27) [resolution 2660,408]
[SA] passive: 3163. alength = X62 | $less(X62,0) | i(l8(sK2(X62))) = X62 | $less(i(l8(p(nl8))),X62) | zero = nl8 (3:21) [resolution 2660,95]
[SA] passive: 3164. alength = X63 | $less(X63,0) | i(l8(sK2(X63))) = X63 | $less(i(l8(p(p(nl8)))),X63) | zero = p(nl8) | zero = nl8 (4:26) [resolution 2660,314]
[SA] passive: 3165. alength = X64 | $less(X64,0) | i(l8(sK2(X64))) = X64 | $less(i(l8(sK2(X65))),X64) | $less(alength,X65) | alength = X65 | $less(X65,0) (3:27) [resolution 2660,2657]
[SA] passive: 3166. alength = X66 | $less(X66,0) | i(l8(sK2(X66))) = X66 | $less(i(l8(sK2(X67))),X66) | ~$less(X67,alength) | $less(X67,0) (3:24) [resolution 2660,2771]
[SA] active: 2831. i(l8(sK2(0))) = 0 | 0 = sK0 (3:9:1) [evaluation 2802]
[SA] new: 3190. ~$less(X0,0) | $less(0,0) | ~$less(0,alength) | $less(X0,alength) | 0 = sK0 (4:15) [superposition 2920,2831]
[SA] new: 3191. $less(0,alength) | $less(alength,0) | alength = 0 | $less(0,0) | 0 = sK0 (4:15) [superposition 2657,2831]
[SA] new: 3192. $less(0,alength) | ~$less(0,alength) | $less(0,0) | 0 = sK0 (4:12) [superposition 2771,2831]
[SA] new: 3193. i(l8(X1)) != 0 | sK2(0) = X1 | 0 = sK0 (4:12) [superposition 48,2831]
[SA] new: 3194. i(l8(X2)) != 0 | sK2(0) = X2 | 0 = sK0 (4:12) [superposition 48,2831]
[SA] new: 3195. v = a(l11(sK2(0)),0) | 0 = sK0 (4:10) [superposition 70,2831]
[SA] new: 3196. i(l8(s(sK2(0)))) = $sum(0,1) | 0 = sK0 (4:12) [superposition 74,2831]
[SA] new: 3197. 0 != 0 | zero = sK2(0) | 0 = sK0 (4:10) [superposition 108,2831]
[SA] new: 3198. 1 != 0 | s(zero) = sK2(0) | 0 = sK0 (4:11) [superposition 121,2831]
[SA] new: 3199. 0 != 2 | s(s(zero)) = sK2(0) | 0 = sK0 (4:12) [superposition 130,2831]
[SA] new: 3200. 0 != 3 | s(s(s(zero))) = sK2(0) | 0 = sK0 (5:13) [superposition 183,2831]
[SA] new: 3201. v = a(l8(s(sK2(0))),0) | 0 = sK0 (4:11) [superposition 349,2831]
[SA] new: 3202. alength != 0 | nl8 = sK2(0) | 0 = sK0 (4:10) [superposition 2483,2831]
[SA] new: 3203. $sum(alength,1) != 0 | s(nl8) = sK2(0) | 0 = sK0 (4:13) [superposition 2490,2831]
[SA] new: 3204. 0 != sK0 | sK2(0) = sK2(sK0) | 0 = sK0 (4:11) [superposition 2839,2831]
[SA] new: 3205. alength != $uminus(1) | s(nl8) = sK2(0) | 0 = sK0 (4:12) [evaluation 3203]
[SA] new: 3206. alength != -1 | s(nl8) = sK2(0) | 0 = sK0 (4:11) [evaluation 3205]
[SA] new: 3207. zero = sK2(0) | 0 = sK0 (4:7) [trivial inequality removal 3197]
[SA] new: 3208. i(l8(s(sK2(0)))) = 1 | 0 = sK0 (4:10) [evaluation 3196]
[SA] new: 3209. $less(0,alength) | $less(alength,0) | alength = 0 | 0 = sK0 (4:12) [evaluation 3191]
[SA] new: 3210. ~$less(X0,0) | ~$less(0,alength) | $less(X0,alength) | 0 = sK0 (4:12) [evaluation 3190]
[SA] passive: 3193. i(l8(X1)) != 0 | sK2(0) = X1 | 0 = sK0 (4:12) [superposition 48,2831]
[SA] passive: 3195. v = a(l11(sK2(0)),0) | 0 = sK0 (4:10) [superposition 70,2831]
[SA] passive: 3208. i(l8(s(sK2(0)))) = 1 | 0 = sK0 (4:10) [evaluation 3196]
[SA] passive: 3207. zero = sK2(0) | 0 = sK0 (4:7) [trivial inequality removal 3197]
[SA] passive: 3201. v = a(l8(s(sK2(0))),0) | 0 = sK0 (4:11) [superposition 349,2831]
[SA] passive: 3202. alength != 0 | nl8 = sK2(0) | 0 = sK0 (4:10) [superposition 2483,2831]
[SA] active: 2921. $less(i(l8(sK2(X3))),X4) | $less(X3,0) | alength = X4 | $less(X4,alength) | ~$less(X3,alength) (2:18:1) [resolution 2771,99]
[SA] new: 3211. $less(X0,0) | alength = i(l8(sK2(X0))) | $less(i(l8(sK2(X0))),alength) | ~$less(X0,alength) (3:18) [resolution 2921,13]
[SA] new: 3212. $less(X1,0) | alength = X2 | $less(X2,alength) | ~$less(X1,alength) | ~$less(X3,i(l8(sK2(X1)))) | $less(X3,X2) (3:21) [resolution 2921,14]
[SA] new: 3213. $less(X4,0) | alength = X5 | $less(X5,alength) | ~$less(X4,alength) | X5 = X6 | $less(X6,X5) | $less(i(l8(sK2(X4))),X6) (3:24) [resolution 2921,99]
[SA] new: 3214. $less(X7,0) | alength = X8 | $less(X8,alength) | ~$less(X7,alength) | sK0 = X8 | $less(sK0,X8) | $less(i(l8(sK2(X7))),alength) (4:24) [resolution 2921,205]
[SA] new: 3215. $less(X9,0) | alength = X10 | $less(X10,alength) | ~$less(X9,alength) | alength = X10 | $less(sK0,X10) | $less(i(l8(sK2(X9))),alength) (4:24) [resolution 2921,254]
[SA] new: 3216. $less(X11,0) | alength = X12 | $less(X12,alength) | ~$less(X11,alength) | sK0 = X12 | 0 = sK0 | $less(0,X12) | $less(i(l8(sK2(X11))),sK0) (4:27) [resolution 2921,333]
[SA] new: 3217. $less(X13,0) | alength = alength | $less(alength,alength) | ~$less(X13,alength) | alength = X14 | $less(X14,0) | i(l8(sK2(X14))) = X14 | $less(i(l8(sK2(X13))),X14) (3:30) [resolution 2921,2660]
[SA] new: 3218. $less(X15,0) | alength = alength | $less(alength,alength) | ~$less(X15,alength) | $less(i(l8(sK2(X15))),0) | i(l8(sK2(X15))) = i(l8(sK2(i(l8(sK2(X15)))))) (3:30) [resolution 2921,2480]
[SA] new: 3219. $less(X16,0) | alength = i(l8(p(nl8))) | $less(i(l8(p(nl8))),alength) | ~$less(X16,alength) | zero = nl8 | $less(i(l8(sK2(X16))),alength) (4:27) [resolution 2921,194]
[SA] new: 3220. $less(X17,0) | alength = i(l8(sK2(X18))) | $less(i(l8(sK2(X18))),alength) | ~$less(X17,alength) | $less(X18,0) | ~$less(X18,alength) | $less(i(l8(sK2(X17))),alength) (3:30) [resolution 2921,2920]
[SA] new: 3221. $less(X19,0) | alength = 0 | $less(0,alength) | ~$less(X19,alength) | sK0 = X20 | 0 = sK0 | $less(X20,sK0) | $less(i(l8(sK2(X19))),X20) (4:27) [resolution 2921,324]
[SA] new: 3222. $less(X21,0) | alength = 0 | $less(0,alength) | ~$less(X21,alength) | alength = 0 | $less(i(l8(sK2(X21))),alength) (5:21) [resolution 2921,262]
[SA] new: 3223. $less(X22,0) | alength = 0 | $less(0,alength) | ~$less(X22,alength) | 0 = sK0 | $less(i(l8(sK2(X22))),alength) (4:21) [resolution 2921,116]
[SA] new: 3224. $less(X23,0) | alength = 0 | $less(0,alength) | ~$less(X23,alength) | 0 = sK0 | $less(i(l8(sK2(X23))),sK0) (3:21) [resolution 2921,114]
[SA] new: 3225. $less(X24,0) | alength = sK0 | $less(sK0,alength) | ~$less(X24,alength) | alength = X25 | $less(X25,alength) | $less(i(l8(sK2(X24))),X25) (4:24) [resolution 2921,247]
[SA] new: 3226. $less(X26,0) | alength = sK0 | $less(sK0,alength) | ~$less(X26,alength) | sK0 = X27 | $less(X27,alength) | $less(i(l8(sK2(X26))),X27) (4:24) [resolution 2921,199]
[SA] new: 3227. $less(X28,0) | alength = sK0 | $less(sK0,alength) | ~$less(X28,alength) | $less(i(l8(sK2(X28))),alength) (3:18) [resolution 2921,96]
[SA] new: 3228. $less(sK0,X0) | $less(sK0,0) | alength = X0 | $less(X0,alength) | ~$less(sK0,alength) | alength = sK0 (4:18) [superposition 2921,2495]
[SA] new: 3229. $less(sK0,X1) | $less(sK0,0) | alength = X1 | $less(X1,alength) | ~$less(sK0,alength) (3:15) [superposition 2921,2833]
[SA] new: 3230. $less(0,X2) | $less(0,0) | alength = X2 | $less(X2,alength) | ~$less(0,alength) | 0 = sK0 (4:18) [superposition 2921,2831]
[SA] new: 3231. $less(0,X2) | alength = X2 | $less(X2,alength) | ~$less(0,alength) | 0 = sK0 (4:15) [evaluation 3230]
[SA] new: 3232. $less(X21,0) | alength = 0 | $less(0,alength) | ~$less(X21,alength) | $less(i(l8(sK2(X21))),alength) (5:18) [duplicate literal removal 3222]
[SA] new: 3233. $less(X9,0) | alength = X10 | $less(X10,alength) | ~$less(X9,alength) | $less(sK0,X10) | $less(i(l8(sK2(X9))),alength) (4:21) [duplicate literal removal 3215]
[SA] passive: 3212. $less(X1,0) | alength = X2 | $less(X2,alength) | ~$less(X1,alength) | ~$less(X3,i(l8(sK2(X1)))) | $less(X3,X2) (3:21) [resolution 2921,14]
[SA] passive: 3213. $less(X4,0) | alength = X5 | $less(X5,alength) | ~$less(X4,alength) | X5 = X6 | $less(X6,X5) | $less(i(l8(sK2(X4))),X6) (3:24) [resolution 2921,99]
[SA] active: 3207. zero = sK2(0) | 0 = sK0 (4:7:1) [trivial inequality removal 3197]
[SA] new: 3234. i(l8(zero)) = 0 | 0 = sK0 | 0 = sK0 (5:11) [superposition 2831,3207]
[SA] new: 3235. $less(i(l8(zero)),X0) | $less(0,0) | alength = X0 | $less(X0,alength) | ~$less(0,alength) | 0 = sK0 (5:20) [superposition 2921,3207]
[SA] new: 3236. ~$less(X1,i(l8(zero))) | $less(0,0) | ~$less(0,alength) | $less(X1,alength) | 0 = sK0 (5:17) [superposition 2920,3207]
[SA] new: 3237. $less(i(l8(zero)),alength) | $less(alength,0) | alength = 0 | $less(0,0) | 0 = sK0 (5:17) [superposition 2657,3207]
[SA] new: 3238. 'Sub'(zero,nl8) | $less(alength,0) | alength = 0 | $less(0,0) | 0 = sK0 (5:15) [superposition 2655,3207]
[SA] new: 3239. $less(i(l8(zero)),alength) | ~$less(0,alength) | $less(0,0) | 0 = sK0 (5:14) [superposition 2771,3207]
[SA] new: 3240. 'Sub'(zero,nl8) | $less(0,0) | ~$less(0,alength) | 0 = sK0 (5:12) [superposition 2479,3207]
[SA] new: 3241. 'Sub'(zero,nl8) | ~$less(0,alength) | 0 = sK0 (5:9) [evaluation 3240]
[SA] new: 3242. $less(i(l8(zero)),alength) | ~$less(0,alength) | 0 = sK0 (5:11) [evaluation 3239]
[SA] new: 3243. 'Sub'(zero,nl8) | $less(alength,0) | alength = 0 | 0 = sK0 (5:12) [evaluation 3238]
[SA] new: 3244. $less(i(l8(zero)),alength) | $less(alength,0) | alength = 0 | 0 = sK0 (5:14) [evaluation 3237]
[SA] new: 3245. ~$less(X1,i(l8(zero))) | ~$less(0,alength) | $less(X1,alength) | 0 = sK0 (5:14) [evaluation 3236]
[SA] new: 3246. $less(i(l8(zero)),X0) | alength = X0 | $less(X0,alength) | ~$less(0,alength) | 0 = sK0 (5:17) [evaluation 3235]
[SA] new: 3247. i(l8(zero)) = 0 | 0 = sK0 (5:8) [duplicate literal removal 3234]
[SA] new: 3248. $less(i(l8(zero)),X0) | alength = X0 | $less(X0,alength) | 0 = sK0 (5:14) [subsumption resolution 3246,113]
[SA] new: 3249. $less(0,X0) | alength = X0 | $less(X0,alength) | 0 = sK0 (5:12) [forward demodulation 3248,65]
[SA] new: 3250. ~$less(X1,i(l8(zero))) | $less(X1,alength) | 0 = sK0 (5:11) [subsumption resolution 3245,113]
[SA] new: 3251. ~$less(X1,0) | $less(X1,alength) | 0 = sK0 (5:9) [forward demodulation 3250,65]
[SA] new: 3252. $less(0,alength) | $less(alength,0) | alength = 0 | 0 = sK0 (5:12) [forward demodulation 3244,65]
[SA] passive: 3243. 'Sub'(zero,nl8) | $less(alength,0) | alength = 0 | 0 = sK0 (5:12) [evaluation 3238]
[SA] new: 3253. $less(i(l8(zero)),alength) | 0 = sK0 (5:8) [subsumption resolution 3242,113]
[SA] new: 3254. $less(0,alength) | 0 = sK0 (5:6) [forward demodulation 3253,65]
[SA] new: 3255. 'Sub'(zero,nl8) | 0 = sK0 (5:6) [subsumption resolution 3241,113]
[SA] passive: 3255. 'Sub'(zero,nl8) | 0 = sK0 (5:6) [subsumption resolution 3241,113]
[SA] active: 2950. ~$less(X6,X5) | i(l8(sK2(X5))) = X5 | alength = X5 | $less(alength,X5) | $less(X6,0) (2:18:1) [resolution 2656,14]
[SA] new: 3256. i(l8(sK2(X0))) = X0 | alength = X0 | $less(alength,X0) | $less(X1,0) | $less(X0,X1) | X0 = X1 (3:21) [resolution 2950,15]
[SA] new: 3257. i(l8(sK2(X2))) = X2 | alength = X2 | $less(alength,X2) | $less(X3,0) | $less(X2,X3) | X2 = X3 (3:21) [resolution 2950,15]
[SA] new: 3258. i(l8(sK2(X4))) = X4 | alength = X4 | $less(alength,X4) | $less(X5,0) | $less(X6,X5) | $less(X4,X6) | X4 = X5 | X4 = X6 (3:27) [resolution 2950,174]
[SA] new: 3259. i(l8(sK2(X7))) = X7 | alength = X7 | $less(alength,X7) | $less(X8,0) | $less(X7,X9) | $less(X9,X8) | X8 = X9 | X7 = X8 (3:27) [resolution 2950,174]
[SA] new: 3260. i(l8(sK2(X10))) = X10 | alength = X10 | $less(alength,X10) | $less(X11,0) | $less(X12,X11) | $less(X10,X12) | X10 = X12 | X11 = X12 (3:27) [resolution 2950,174]
[SA] new: 3261. i(l8(sK2(X13))) = X13 | alength = X13 | $less(alength,X13) | $less(X14,0) | $less(sK0,X14) | $less(X13,alength) | X13 = X14 | sK0 = X13 (4:27) [resolution 2950,198]
[SA] new: 3262. i(l8(sK2(X15))) = X15 | alength = X15 | $less(alength,X15) | $less(X16,0) | $less(sK0,X16) | $less(X15,alength) | alength = X15 | sK0 = X16 (4:27) [resolution 2950,204]
[SA] new: 3263. i(l8(sK2(X17))) = X17 | alength = X17 | $less(alength,X17) | $less(X18,0) | $less(sK0,X18) | $less(X17,alength) | X17 = X18 | alength = X17 (4:27) [resolution 2950,246]
[SA] new: 3264. i(l8(sK2(X19))) = X19 | alength = X19 | $less(alength,X19) | $less(X20,0) | $less(sK0,X20) | $less(X19,alength) | alength = X19 | alength = X20 (4:27) [resolution 2950,253]
[SA] new: 3265. i(l8(sK2(X21))) = X21 | alength = X21 | $less(alength,X21) | $less(X22,0) | $less(0,X22) | $less(X21,sK0) | X21 = X22 | 0 = sK0 | sK0 = X21 (4:30) [resolution 2950,323]
[SA] new: 3266. i(l8(sK2(X23))) = X23 | alength = X23 | $less(alength,X23) | $less(X24,0) | $less(0,X24) | $less(X23,sK0) | sK0 = X23 | 0 = sK0 | sK0 = X24 (4:30) [resolution 2950,332]
[SA] new: 3267. i(l8(sK2(X25))) = X25 | alength = X25 | $less(alength,X25) | $less(X26,0) | $less(0,X26) | $less(X25,sK0) | X25 = X26 | 0 = sK0 | 0 = X26 (4:30) [resolution 2950,387]
[SA] new: 3268. i(l8(sK2(X27))) = X27 | alength = X27 | $less(alength,X27) | $less(X28,0) | $less(sK0,X28) | $less(X27,alength) | X27 = X28 | sK0 = X28 (4:27) [resolution 2950,388]
[SA] new: 3269. i(l8(sK2(X29))) = X29 | alength = X29 | $less(alength,X29) | $less(X30,0) | $less(sK0,X30) | $less(X29,alength) | sK0 = X30 | sK0 = X29 (4:27) [resolution 2950,408]
[SA] new: 3270. alength = i(l8(sK2(alength))) | alength = alength | $less(alength,alength) | $less(X31,0) | $less(sK0,X31) | sK0 = X31 (3:21) [resolution 2950,106]
[SA] new: 3271. alength = i(l8(sK2(alength))) | alength = alength | $less(alength,alength) | $less(X32,0) | $less(0,X32) | alength = X32 | 0 = sK0 (4:24) [resolution 2950,178]
[SA] new: 3272. alength = i(l8(sK2(alength))) | alength = alength | $less(alength,alength) | $less(X33,0) | $less(sK0,X33) | alength = X33 (3:21) [resolution 2950,179]
[SA] new: 3273. alength = i(l8(sK2(alength))) | alength = alength | $less(alength,alength) | $less(X34,0) | $less(sK0,X35) | $less(X35,X34) | X34 = X35 | sK0 = X34 (4:27) [resolution 2950,198]
[SA] new: 3274. alength = i(l8(sK2(alength))) | alength = alength | $less(alength,alength) | $less(X36,0) | $less(sK0,X37) | $less(X37,X36) | alength = X36 | sK0 = X37 (4:27) [resolution 2950,204]
[SA] new: 3275. alength = i(l8(sK2(alength))) | alength = alength | $less(alength,alength) | $less(X38,0) | $less(sK0,X39) | $less(X39,X38) | X38 = X39 | alength = X38 (4:27) [resolution 2950,246]
[SA] new: 3276. alength = i(l8(sK2(alength))) | alength = alength | $less(alength,alength) | $less(X40,0) | $less(sK0,X41) | $less(X41,X40) | alength = X40 | alength = X41 (4:27) [resolution 2950,253]
[SA] new: 3277. alength = i(l8(sK2(alength))) | alength = alength | $less(alength,alength) | $less(X42,0) | $less(0,X42) | alength = X42 | alength = 0 (5:24) [resolution 2950,261]
[SA] new: 3278. alength = i(l8(sK2(alength))) | alength = alength | $less(alength,alength) | $less(X43,0) | $less(0,X43) | 0 = sK0 | 0 = X43 (5:24) [resolution 2950,278]
[SA] new: 3279. alength = i(l8(sK2(alength))) | alength = alength | $less(alength,alength) | $less(X44,0) | $less(0,X44) | 0 = sK0 | sK0 = X44 (4:24) [resolution 2950,330]
[SA] new: 3280. alength = i(l8(sK2(alength))) | alength = alength | $less(alength,alength) | $less(X45,0) | $less(sK0,X46) | $less(X46,X45) | X45 = X46 | sK0 = X46 (4:27) [resolution 2950,388]
[SA] new: 3281. alength = i(l8(sK2(alength))) | alength = alength | $less(alength,alength) | $less(X47,0) | $less(sK0,X48) | $less(X48,X47) | sK0 = X48 | sK0 = X47 (4:27) [resolution 2950,408]
[SA] new: 3282. i(l8(sK2(0))) = 0 | alength = 0 | $less(alength,0) | $less(X49,0) | $less(alength,X49) | i(l8(sK2(X49))) = X49 | alength = X49 (3:27) [resolution 2950,2656]
[SA] new: 3283. i(l8(sK2(sK0))) = sK0 | alength = sK0 | $less(alength,sK0) | $less(X50,0) | $less(0,X51) | $less(X51,X50) | X50 = X51 | 0 = sK0 | 0 = X51 (4:30) [resolution 2950,387]
[SA] new: 3284. i(l8(sK2(sK0))) = sK0 | alength = sK0 | $less(alength,sK0) | $less(X52,0) | $less(0,X53) | $less(X53,X52) | sK0 = X52 | 0 = sK0 | sK0 = X53 (4:30) [resolution 2950,332]
[SA] new: 3285. i(l8(sK2(sK0))) = sK0 | alength = sK0 | $less(alength,sK0) | $less(X54,0) | $less(0,X55) | $less(X55,X54) | X54 = X55 | 0 = sK0 | sK0 = X54 (4:30) [resolution 2950,323]
[SA] new: 3286. i(l8(sK2(sK0))) = sK0 | alength = sK0 | $less(alength,sK0) | $less(X56,0) | $less(0,X56) | 0 = X56 | sK0 = X56 (4:24) [resolution 2950,380]
[SA] new: 3287. i(l8(sK2(sK0))) = sK0 | alength = sK0 | $less(alength,sK0) | $less(X57,0) | $less(0,X57) | 0 = sK0 | 0 = X57 (4:24) [resolution 2950,229]
[SA] new: 3288. i(l8(sK2(sK0))) = sK0 | alength = sK0 | $less(alength,sK0) | $less(X58,0) | $less(0,X58) | sK0 = X58 | 0 = sK0 (3:24) [resolution 2950,177]
[SA] new: 3289. i(l8(sK2(X59))) = X59 | alength = X59 | $less(alength,X59) | $less(alength,0) | $less(X59,0) | i(l8(sK2(X59))) = X59 | alength = X59 (3:27) [resolution 2950,2656]
[SA] new: 3290. alength = i(l8(sK2(alength))) | alength = alength | $less(alength,alength) | $less(i(l8(p(nl8))),0) | zero = nl8 (3:21) [resolution 2950,95]
[SA] new: 3291. alength = i(l8(sK2(alength))) | alength = alength | $less(alength,alength) | $less(i(l8(p(p(nl8)))),0) | zero = p(nl8) | zero = nl8 (4:26) [resolution 2950,314]
[SA] new: 3292. alength = i(l8(sK2(alength))) | alength = alength | $less(alength,alength) | $less(i(l8(sK2(X60))),0) | $less(alength,X60) | alength = X60 | $less(X60,0) (3:27) [resolution 2950,2657]
[SA] new: 3293. alength = i(l8(sK2(alength))) | alength = alength | $less(alength,alength) | $less(i(l8(sK2(X61))),0) | ~$less(X61,alength) | $less(X61,0) (3:24) [resolution 2950,2771]
[SA] new: 3294. i(l8(sK2(X62))) = X62 | alength = X62 | $less(alength,X62) | $less(i(l8(p(nl8))),0) | alength = X62 | $less(X62,alength) | zero = nl8 (4:27) [resolution 2950,193]
[SA] new: 3295. i(l8(sK2(X63))) = X63 | alength = X63 | $less(alength,X63) | $less(i(l8(sK2(X64))),0) | $less(X64,0) | alength = X63 | $less(X63,alength) | ~$less(X64,alength) (3:30) [resolution 2950,2921]
[SA] new: 3296. i(l8(sK2(X65))) = X65 | alength = X65 | $less(alength,X65) | $less(0,0) | $less(X65,sK0) | sK0 = X65 | 0 = sK0 (3:24) [resolution 2950,177]
[SA] new: 3297. i(l8(sK2(X66))) = X66 | alength = X66 | $less(alength,X66) | $less(0,0) | $less(X66,alength) | alength = X66 | 0 = sK0 (4:24) [resolution 2950,178]
[SA] new: 3298. i(l8(sK2(X67))) = X67 | alength = X67 | $less(alength,X67) | $less(0,0) | $less(X67,sK0) | 0 = sK0 | 0 = X67 (4:24) [resolution 2950,229]
[SA] new: 3299. i(l8(sK2(X68))) = X68 | alength = X68 | $less(alength,X68) | $less(0,0) | $less(X68,alength) | alength = X68 | alength = 0 (5:24) [resolution 2950,261]
[SA] new: 3300. i(l8(sK2(X69))) = X69 | alength = X69 | $less(alength,X69) | $less(0,0) | $less(X69,alength) | 0 = sK0 | 0 = X69 (5:24) [resolution 2950,278]
[SA] new: 3301. i(l8(sK2(X70))) = X70 | alength = X70 | $less(alength,X70) | $less(0,0) | $less(X70,X71) | $less(X71,sK0) | X70 = X71 | 0 = sK0 | sK0 = X71 (4:30) [resolution 2950,323]
[SA] new: 3302. i(l8(sK2(X72))) = X72 | alength = X72 | $less(alength,X72) | $less(0,0) | $less(X72,alength) | 0 = sK0 | sK0 = X72 (4:24) [resolution 2950,330]
[SA] new: 3303. i(l8(sK2(X73))) = X73 | alength = X73 | $less(alength,X73) | $less(0,0) | $less(X74,sK0) | $less(X73,X74) | sK0 = X74 | 0 = sK0 | sK0 = X73 (4:30) [resolution 2950,332]
[SA] new: 3304. i(l8(sK2(X75))) = X75 | alength = X75 | $less(alength,X75) | $less(0,0) | $less(X75,sK0) | 0 = X75 | sK0 = X75 (4:24) [resolution 2950,380]
[SA] new: 3305. i(l8(sK2(X76))) = X76 | alength = X76 | $less(alength,X76) | $less(0,0) | $less(X77,sK0) | $less(X76,X77) | X76 = X77 | 0 = sK0 | 0 = X76 (4:30) [resolution 2950,387]
[SA] new: 3306. alength = i(l8(sK2(alength))) | alength = alength | $less(alength,alength) | $less(0,0) | alength = 0 (4:18) [resolution 2950,244]
[SA] new: 3307. alength = i(l8(sK2(alength))) | alength = alength | $less(alength,alength) | $less(0,0) | 0 = sK0 (3:18) [resolution 2950,113]
[SA] new: 3308. i(l8(sK2(sK0))) = sK0 | alength = sK0 | $less(alength,sK0) | $less(0,0) | 0 = sK0 (3:18) [resolution 2950,97]
[SA] new: 3309. alength = i(l8(sK2(alength))) | alength = alength | $less(alength,alength) | $less(sK0,0) (3:15) [resolution 2950,45]
[SA] new: 3310. i(l8(sK2(X78))) = X78 | alength = X78 | $less(alength,X78) | $less(sK0,0) | $less(X78,alength) | sK0 = X78 (3:21) [resolution 2950,106]
[SA] new: 3311. i(l8(sK2(X79))) = X79 | alength = X79 | $less(alength,X79) | $less(sK0,0) | $less(X79,alength) | alength = X79 (3:21) [resolution 2950,179]
[SA] new: 3312. i(l8(sK2(X80))) = X80 | alength = X80 | $less(alength,X80) | $less(sK0,0) | $less(X80,X81) | $less(X81,alength) | X80 = X81 | sK0 = X81 (4:27) [resolution 2950,198]
[SA] new: 3313. i(l8(sK2(X82))) = X82 | alength = X82 | $less(alength,X82) | $less(sK0,0) | $less(X83,alength) | $less(X82,X83) | alength = X83 | sK0 = X82 (4:27) [resolution 2950,204]
[SA] new: 3314. i(l8(sK2(X84))) = X84 | alength = X84 | $less(alength,X84) | $less(sK0,0) | $less(X84,X85) | $less(X85,alength) | X84 = X85 | alength = X85 (4:27) [resolution 2950,246]
[SA] new: 3315. i(l8(sK2(X86))) = X86 | alength = X86 | $less(alength,X86) | $less(sK0,0) | $less(X87,alength) | $less(X86,X87) | alength = X87 | alength = X86 (4:27) [resolution 2950,253]
[SA] new: 3316. i(l8(sK2(X88))) = X88 | alength = X88 | $less(alength,X88) | $less(sK0,0) | $less(X89,alength) | $less(X88,X89) | X88 = X89 | sK0 = X88 (4:27) [resolution 2950,388]
[SA] new: 3317. i(l8(sK2(X90))) = X90 | alength = X90 | $less(alength,X90) | $less(sK0,0) | $less(X91,alength) | $less(X90,X91) | sK0 = X90 | sK0 = X91 (4:27) [resolution 2950,408]
[SA] new: 3318. i(l8(sK2(X86))) = X86 | alength = X86 | $less(alength,X86) | $less(sK0,0) | $less(X87,alength) | $less(X86,X87) | alength = X87 (4:24) [duplicate literal removal 3315]
[SA] new: 3319. i(l8(sK2(X79))) = X79 | alength = X79 | $less(alength,X79) | $less(sK0,0) | $less(X79,alength) (3:18) [duplicate literal removal 3311]
[SA] new: 3320. i(l8(sK2(sK0))) = sK0 | alength = sK0 | $less(alength,sK0) | 0 = sK0 (3:15) [evaluation 3308]
[SA] new: 3321. i(l8(sK2(X76))) = X76 | alength = X76 | $less(alength,X76) | $less(X77,sK0) | $less(X76,X77) | X76 = X77 | 0 = sK0 | 0 = X76 (4:27) [evaluation 3305]
[SA] new: 3322. i(l8(sK2(X75))) = X75 | alength = X75 | $less(alength,X75) | $less(X75,sK0) | 0 = X75 | sK0 = X75 (4:21) [evaluation 3304]
[SA] new: 3323. i(l8(sK2(X73))) = X73 | alength = X73 | $less(alength,X73) | $less(X74,sK0) | $less(X73,X74) | sK0 = X74 | 0 = sK0 | sK0 = X73 (4:27) [evaluation 3303]
[SA] new: 3324. i(l8(sK2(X72))) = X72 | alength = X72 | $less(alength,X72) | $less(X72,alength) | 0 = sK0 | sK0 = X72 (4:21) [evaluation 3302]
[SA] new: 3325. i(l8(sK2(X70))) = X70 | alength = X70 | $less(alength,X70) | $less(X70,X71) | $less(X71,sK0) | X70 = X71 | 0 = sK0 | sK0 = X71 (4:27) [evaluation 3301]
[SA] new: 3326. i(l8(sK2(X69))) = X69 | alength = X69 | $less(alength,X69) | $less(X69,alength) | 0 = sK0 | 0 = X69 (5:21) [evaluation 3300]
[SA] new: 3327. i(l8(sK2(X68))) = X68 | alength = X68 | $less(alength,X68) | $less(0,0) | $less(X68,alength) | alength = 0 (5:21) [duplicate literal removal 3299]
[SA] new: 3328. i(l8(sK2(X68))) = X68 | alength = X68 | $less(alength,X68) | $less(X68,alength) | alength = 0 (5:18) [evaluation 3327]
[SA] new: 3329. i(l8(sK2(X67))) = X67 | alength = X67 | $less(alength,X67) | $less(X67,sK0) | 0 = sK0 | 0 = X67 (4:21) [evaluation 3298]
[SA] new: 3330. i(l8(sK2(X66))) = X66 | alength = X66 | $less(alength,X66) | $less(0,0) | $less(X66,alength) | 0 = sK0 (4:21) [duplicate literal removal 3297]
[SA] new: 3331. i(l8(sK2(X66))) = X66 | alength = X66 | $less(alength,X66) | $less(X66,alength) | 0 = sK0 (4:18) [evaluation 3330]
[SA] new: 3332. i(l8(sK2(X65))) = X65 | alength = X65 | $less(alength,X65) | $less(X65,sK0) | sK0 = X65 | 0 = sK0 (3:21) [evaluation 3296]
[SA] new: 3333. i(l8(sK2(X63))) = X63 | alength = X63 | $less(alength,X63) | $less(i(l8(sK2(X64))),0) | $less(X64,0) | $less(X63,alength) | ~$less(X64,alength) (3:27) [duplicate literal removal 3295]
[SA] new: 3334. i(l8(sK2(X62))) = X62 | alength = X62 | $less(alength,X62) | $less(i(l8(p(nl8))),0) | $less(X62,alength) | zero = nl8 (4:24) [duplicate literal removal 3294]
[SA] new: 3335. i(l8(sK2(X59))) = X59 | alength = X59 | $less(alength,X59) | $less(alength,0) | $less(X59,0) (3:18) [duplicate literal removal 3289]
[SA] new: 3336. i(l8(sK2(X19))) = X19 | alength = X19 | $less(alength,X19) | $less(X20,0) | $less(sK0,X20) | $less(X19,alength) | alength = X20 (4:24) [duplicate literal removal 3264]
[SA] new: 3337. i(l8(sK2(X17))) = X17 | alength = X17 | $less(alength,X17) | $less(X18,0) | $less(sK0,X18) | $less(X17,alength) | X17 = X18 (4:24) [duplicate literal removal 3263]
[SA] new: 3338. i(l8(sK2(X15))) = X15 | alength = X15 | $less(alength,X15) | $less(X16,0) | $less(sK0,X16) | $less(X15,alength) | sK0 = X16 (4:24) [duplicate literal removal 3262]
[SA] passive: 3256. i(l8(sK2(X0))) = X0 | alength = X0 | $less(alength,X0) | $less(X1,0) | $less(X0,X1) | X0 = X1 (3:21) [resolution 2950,15]
[SA] passive: 3258. i(l8(sK2(X4))) = X4 | alength = X4 | $less(alength,X4) | $less(X5,0) | $less(X6,X5) | $less(X4,X6) | X4 = X5 | X4 = X6 (3:27) [resolution 2950,174]
[SA] passive: 3259. i(l8(sK2(X7))) = X7 | alength = X7 | $less(alength,X7) | $less(X8,0) | $less(X7,X9) | $less(X9,X8) | X8 = X9 | X7 = X8 (3:27) [resolution 2950,174]
[SA] passive: 3260. i(l8(sK2(X10))) = X10 | alength = X10 | $less(alength,X10) | $less(X11,0) | $less(X12,X11) | $less(X10,X12) | X10 = X12 | X11 = X12 (3:27) [resolution 2950,174]
[SA] passive: 3322. i(l8(sK2(X75))) = X75 | alength = X75 | $less(alength,X75) | $less(X75,sK0) | 0 = X75 | sK0 = X75 (4:21) [evaluation 3304]
[SA] passive: 3321. i(l8(sK2(X76))) = X76 | alength = X76 | $less(alength,X76) | $less(X77,sK0) | $less(X76,X77) | X76 = X77 | 0 = sK0 | 0 = X76 (4:27) [evaluation 3305]
[SA] new: 3339. i(l8(sK2(X90))) = X90 | alength = X90 | $less(alength,X90) | $less(X91,alength) | $less(X90,X91) | sK0 = X90 | sK0 = X91 (4:24) [subsumption resolution 3317,46]
[SA] passive: 3339. i(l8(sK2(X90))) = X90 | alength = X90 | $less(alength,X90) | $less(X91,alength) | $less(X90,X91) | sK0 = X90 | sK0 = X91 (4:24) [subsumption resolution 3317,46]
[SA] active: 3255. 'Sub'(zero,nl8) | 0 = sK0 (5:6:1) [subsumption resolution 3241,113]
[SA] new: 3340. 0 = sK0 | a(main_end,X0) = a(l11(zero),X0) | i(l8(sK1(X0,zero))) = X0 (6:18) [resolution 3255,150]
[SA] new: 3341. 0 = sK0 | $less(i(l8(zero)),alength) (6:8) [resolution 3255,56]
[SA] passive: 3340. 0 = sK0 | a(main_end,X0) = a(l11(zero),X0) | i(l8(sK1(X0,zero))) = X0 (6:18) [resolution 3255,150]
[SA] new: 3342. $less(0,alength) | 0 = sK0 (6:6) [forward demodulation 3341,65]
[SA] active: 2659. $less(alength,X1) | $less(X1,X0) | $less(X0,0) | alength = X0 | X0 = X1 | i(l8(sK2(X0))) = X0 (2:21:3) [forward demodulation 2497,2476]
[SA] new: 3343. $less(sK0,X0) | $less(X0,0) | alength = X0 | sK0 = X0 | i(l8(sK2(X0))) = X0 | alength = sK0 (6:21) [resolution 2659,2704]
[SA] new: 3344. $less(alength,X1) | $less(X1,0) | alength = X1 | alength = X1 | i(l8(sK2(X1))) = X1 (3:18) [resolution 2659,13]
[SA] new: 3345. $less(X2,X3) | $less(X3,0) | alength = X3 | X2 = X3 | i(l8(sK2(X3))) = X3 | ~$less(X4,alength) | $less(X4,X2) (3:24) [resolution 2659,14]
[SA] new: 3346. $less(X5,X6) | $less(X6,0) | alength = X6 | X5 = X6 | i(l8(sK2(X6))) = X6 | X5 = X7 | $less(X7,X5) | $less(alength,X7) (3:27) [resolution 2659,99]
[SA] new: 3347. $less(X8,X9) | $less(X9,0) | alength = X9 | X8 = X9 | i(l8(sK2(X9))) = X9 | sK0 = X8 | $less(sK0,X8) | $less(alength,alength) (4:27) [resolution 2659,205]
[SA] new: 3348. $less(X10,X11) | $less(X11,0) | alength = X11 | X10 = X11 | i(l8(sK2(X11))) = X11 | alength = X10 | $less(sK0,X10) | $less(alength,alength) (4:27) [resolution 2659,254]
[SA] new: 3349. $less(X12,X13) | $less(X13,0) | alength = X13 | X12 = X13 | i(l8(sK2(X13))) = X13 | sK0 = X12 | 0 = sK0 | $less(0,X12) | $less(alength,sK0) (4:30) [resolution 2659,333]
[SA] new: 3350. $less(X14,X15) | $less(X15,0) | alength = X15 | X14 = X15 | i(l8(sK2(X15))) = X15 | i(l8(sK2(X14))) = X14 | alength = X14 | $less(alength,X14) | $less(alength,0) (3:33) [resolution 2659,2950]
[SA] new: 3351. $less(alength,X16) | $less(X16,0) | alength = X16 | alength = X16 | i(l8(sK2(X16))) = X16 | alength = X17 | $less(X17,0) | i(l8(sK2(X17))) = X17 | $less(alength,X17) (3:33) [resolution 2659,2660]
[SA] new: 3352. $less(alength,X18) | $less(X18,0) | alength = X18 | alength = X18 | i(l8(sK2(X18))) = X18 | $less(alength,0) | alength = i(l8(sK2(alength))) (3:27) [resolution 2659,2480]
[SA] new: 3353. $less(i(l8(p(nl8))),X19) | $less(X19,0) | alength = X19 | i(l8(p(nl8))) = X19 | i(l8(sK2(X19))) = X19 | zero = nl8 | $less(alength,alength) (4:30) [resolution 2659,194]
[SA] new: 3354. $less(i(l8(sK2(X20))),X21) | $less(X21,0) | alength = X21 | i(l8(sK2(X20))) = X21 | i(l8(sK2(X21))) = X21 | $less(X20,0) | ~$less(X20,alength) | $less(alength,alength) (3:33) [resolution 2659,2920]
[SA] new: 3355. $less(0,X22) | $less(X22,0) | alength = X22 | 0 = X22 | i(l8(sK2(X22))) = X22 | sK0 = X23 | 0 = sK0 | $less(X23,sK0) | $less(alength,X23) (4:30) [resolution 2659,324]
[SA] new: 3356. $less(0,X24) | $less(X24,0) | alength = X24 | 0 = X24 | i(l8(sK2(X24))) = X24 | alength = 0 | $less(alength,alength) (5:24) [resolution 2659,262]
[SA] new: 3357. $less(0,X25) | $less(X25,0) | alength = X25 | 0 = X25 | i(l8(sK2(X25))) = X25 | 0 = sK0 | $less(alength,alength) (4:24) [resolution 2659,116]
[SA] new: 3358. $less(0,X26) | $less(X26,0) | alength = X26 | 0 = X26 | i(l8(sK2(X26))) = X26 | 0 = sK0 | $less(alength,sK0) (3:24) [resolution 2659,114]
[SA] new: 3359. $less(sK0,X27) | $less(X27,0) | alength = X27 | sK0 = X27 | i(l8(sK2(X27))) = X27 | alength = X28 | $less(X28,alength) | $less(alength,X28) (4:27) [resolution 2659,247]
[SA] new: 3360. $less(sK0,X29) | $less(X29,0) | alength = X29 | sK0 = X29 | i(l8(sK2(X29))) = X29 | sK0 = X30 | $less(X30,alength) | $less(alength,X30) (4:27) [resolution 2659,199]
[SA] new: 3361. $less(sK0,X31) | $less(X31,0) | alength = X31 | sK0 = X31 | i(l8(sK2(X31))) = X31 | $less(alength,alength) (3:21) [resolution 2659,96]
[SA] new: 3362. $less(alength,sK0) | $less(0,0) | alength = 0 | 0 = sK0 | i(l8(sK2(0))) = 0 (3:18) [resolution 2659,46]
[SA] new: 3363. $less(alength,X0) | $less(X0,0) | alength = X0 | X0 = X0 | i(l8(sK2(X0))) = X0 (3:18) [resolution 2659,13]
[SA] new: 3364. $less(alength,X1) | $less(X2,0) | alength = X2 | X1 = X2 | i(l8(sK2(X2))) = X2 | ~$less(X3,X1) | $less(X3,X2) (3:24) [resolution 2659,14]
[SA] new: 3365. $less(alength,X4) | $less(X5,0) | alength = X5 | X4 = X5 | i(l8(sK2(X5))) = X5 | X5 = X6 | $less(X6,X5) | $less(X4,X6) (3:27) [resolution 2659,99]
[SA] new: 3366. $less(alength,X7) | $less(X8,0) | alength = X8 | X7 = X8 | i(l8(sK2(X8))) = X8 | sK0 = X8 | $less(sK0,X8) | $less(X7,alength) (4:27) [resolution 2659,205]
[SA] new: 3367. $less(alength,X9) | $less(X10,0) | alength = X10 | X9 = X10 | i(l8(sK2(X10))) = X10 | alength = X10 | $less(sK0,X10) | $less(X9,alength) (4:27) [resolution 2659,254]
[SA] new: 3368. $less(alength,X11) | $less(X12,0) | alength = X12 | X11 = X12 | i(l8(sK2(X12))) = X12 | sK0 = X12 | 0 = sK0 | $less(0,X12) | $less(X11,sK0) (4:30) [resolution 2659,333]
[SA] new: 3369. $less(alength,X13) | $less(X14,0) | alength = X14 | X13 = X14 | i(l8(sK2(X14))) = X14 | i(l8(sK2(X14))) = X14 | alength = X14 | $less(alength,X14) | $less(X13,0) (3:33) [resolution 2659,2950]
[SA] new: 3370. $less(alength,X15) | $less(alength,0) | alength = alength | alength = X15 | alength = i(l8(sK2(alength))) | alength = X16 | $less(X16,0) | i(l8(sK2(X16))) = X16 | $less(X15,X16) (3:33) [resolution 2659,2660]
[SA] new: 3371. $less(alength,X17) | $less(alength,0) | alength = alength | alength = X17 | alength = i(l8(sK2(alength))) | $less(X17,0) | i(l8(sK2(X17))) = X17 (3:27) [resolution 2659,2480]
[SA] new: 3372. $less(alength,X18) | $less(i(l8(p(nl8))),0) | alength = i(l8(p(nl8))) | i(l8(p(nl8))) = X18 | i(l8(p(nl8))) = i(l8(sK2(i(l8(p(nl8)))))) | zero = nl8 | $less(X18,alength) (4:39) [resolution 2659,194]
[SA] new: 3373. $less(alength,X19) | $less(i(l8(sK2(X20))),0) | alength = i(l8(sK2(X20))) | i(l8(sK2(X20))) = X19 | i(l8(sK2(X20))) = i(l8(sK2(i(l8(sK2(X20)))))) | $less(X20,0) | ~$less(X20,alength) | $less(X19,alength) (3:42) [resolution 2659,2920]
[SA] new: 3374. $less(alength,X21) | $less(0,0) | alength = 0 | 0 = X21 | i(l8(sK2(0))) = 0 | sK0 = X22 | 0 = sK0 | $less(X22,sK0) | $less(X21,X22) (4:30) [resolution 2659,324]
[SA] new: 3375. $less(alength,X23) | $less(0,0) | alength = 0 | 0 = X23 | i(l8(sK2(0))) = 0 | alength = 0 | $less(X23,alength) (5:24) [resolution 2659,262]
[SA] new: 3376. $less(alength,X24) | $less(0,0) | alength = 0 | 0 = X24 | i(l8(sK2(0))) = 0 | 0 = sK0 | $less(X24,alength) (4:24) [resolution 2659,116]
[SA] new: 3377. $less(alength,X25) | $less(0,0) | alength = 0 | 0 = X25 | i(l8(sK2(0))) = 0 | 0 = sK0 | $less(X25,sK0) (3:24) [resolution 2659,114]
[SA] new: 3378. $less(alength,X26) | $less(sK0,0) | alength = sK0 | sK0 = X26 | i(l8(sK2(sK0))) = sK0 | alength = X27 | $less(X27,alength) | $less(X26,X27) (4:27) [resolution 2659,247]
[SA] new: 3379. $less(alength,X28) | $less(sK0,0) | alength = sK0 | sK0 = X28 | i(l8(sK2(sK0))) = sK0 | sK0 = X29 | $less(X29,alength) | $less(X28,X29) (4:27) [resolution 2659,199]
[SA] new: 3380. $less(alength,X30) | $less(sK0,0) | alength = sK0 | sK0 = X30 | i(l8(sK2(sK0))) = sK0 | $less(X30,alength) (3:21) [resolution 2659,96]
[SA] new: 3381. $less(alength,alength) | $less(sK0,0) | alength = sK0 | alength = sK0 | i(l8(sK2(sK0))) = sK0 | alength = sK0 (6:21) [resolution 2659,2704]
[SA] new: 3382. $less(alength,X0) | $less(X0,sK0) | alength = sK0 | sK0 = X0 | i(l8(sK2(sK0))) = sK0 (3:18) [resolution 2659,46]
[SA] new: 3383. $less(alength,X1) | $less(X1,X2) | alength = X2 | X1 = X2 | i(l8(sK2(X2))) = X2 | sK0 = X3 | 0 = sK0 | $less(X3,sK0) | $less(X2,X3) (4:30) [resolution 2659,324]
[SA] new: 3384. $less(alength,X4) | $less(X4,X5) | alength = X5 | X4 = X5 | i(l8(sK2(X5))) = X5 | alength = 0 | $less(X5,alength) (5:24) [resolution 2659,262]
[SA] new: 3385. $less(alength,X6) | $less(X6,X7) | alength = X7 | X6 = X7 | i(l8(sK2(X7))) = X7 | 0 = sK0 | $less(X7,alength) (4:24) [resolution 2659,116]
[SA] new: 3386. $less(alength,X8) | $less(X8,X9) | alength = X9 | X8 = X9 | i(l8(sK2(X9))) = X9 | 0 = sK0 | $less(X9,sK0) (3:24) [resolution 2659,114]
[SA] new: 3387. $less(alength,X10) | $less(X10,0) | alength = 0 | 0 = X10 | i(l8(sK2(0))) = 0 (3:18) [resolution 2659,13]
[SA] new: 3388. $less(alength,X11) | $less(X11,X12) | alength = X12 | X11 = X12 | i(l8(sK2(X12))) = X12 | ~$less(X13,X12) | $less(X13,0) (3:24) [resolution 2659,14]
[SA] new: 3389. $less(alength,X14) | $less(X14,X15) | alength = X15 | X14 = X15 | i(l8(sK2(X15))) = X15 | 0 = X16 | $less(X16,0) | $less(X15,X16) (3:27) [resolution 2659,99]
[SA] new: 3390. $less(alength,X17) | $less(X17,X18) | alength = X18 | X17 = X18 | i(l8(sK2(X18))) = X18 | 0 = sK0 | $less(sK0,0) | $less(X18,alength) (4:27) [resolution 2659,205]
[SA] new: 3391. $less(alength,X19) | $less(X19,X20) | alength = X20 | X19 = X20 | i(l8(sK2(X20))) = X20 | alength = 0 | $less(sK0,0) | $less(X20,alength) (4:27) [resolution 2659,254]
[SA] new: 3392. $less(alength,X21) | $less(X21,X22) | alength = X22 | X21 = X22 | i(l8(sK2(X22))) = X22 | 0 = sK0 | 0 = sK0 | $less(0,0) | $less(X22,sK0) (4:30) [resolution 2659,333]
[SA] new: 3393. $less(alength,X23) | $less(X23,X24) | alength = X24 | X23 = X24 | i(l8(sK2(X24))) = X24 | i(l8(sK2(0))) = 0 | alength = 0 | $less(alength,0) | $less(X24,0) (3:33) [resolution 2659,2950]
[SA] new: 3394. $less(alength,alength) | $less(alength,0) | alength = alength | alength = alength | alength = i(l8(sK2(alength))) (3:18) [factoring 2659]
[SA] new: 3395. $less(alength,0) | $less(0,alength) | alength = alength | alength = 0 | alength = i(l8(sK2(alength))) (3:18) [factoring 2659]
[SA] new: 3396. $less(alength,0) | $less(0,0) | alength = 0 | 0 = 0 | i(l8(sK2(0))) = 0 (3:18) [factoring 2659]
[SA] new: 3397. $less(alength,alength) | $less(alength,0) | alength = alength | alength = i(l8(sK2(alength))) (3:15) [duplicate literal removal 3394]
[SA] new: 3398. $less(alength,X21) | $less(X21,X22) | alength = X22 | X21 = X22 | i(l8(sK2(X22))) = X22 | 0 = sK0 | $less(0,0) | $less(X22,sK0) (4:27) [duplicate literal removal 3392]
[SA] new: 3399. $less(alength,X21) | $less(X21,X22) | alength = X22 | X21 = X22 | i(l8(sK2(X22))) = X22 | 0 = sK0 | $less(X22,sK0) (4:24) [evaluation 3398]
[SA] new: 3400. $less(alength,alength) | $less(sK0,0) | alength = sK0 | i(l8(sK2(sK0))) = sK0 (6:15) [duplicate literal removal 3381]
[SA] new: 3401. $less(alength,X25) | alength = 0 | 0 = X25 | i(l8(sK2(0))) = 0 | 0 = sK0 | $less(X25,sK0) (3:21) [evaluation 3377]
[SA] new: 3402. $less(alength,X24) | alength = 0 | 0 = X24 | i(l8(sK2(0))) = 0 | 0 = sK0 | $less(X24,alength) (4:21) [evaluation 3376]
[SA] new: 3403. $less(alength,X23) | $less(0,0) | alength = 0 | 0 = X23 | i(l8(sK2(0))) = 0 | $less(X23,alength) (5:21) [duplicate literal removal 3375]
[SA] new: 3404. $less(alength,X23) | alength = 0 | 0 = X23 | i(l8(sK2(0))) = 0 | $less(X23,alength) (5:18) [evaluation 3403]
[SA] new: 3405. $less(alength,X21) | alength = 0 | 0 = X21 | i(l8(sK2(0))) = 0 | sK0 = X22 | 0 = sK0 | $less(X22,sK0) | $less(X21,X22) (4:27) [evaluation 3374]
[SA] new: 3406. $less(alength,X13) | $less(X14,0) | alength = X14 | X13 = X14 | i(l8(sK2(X14))) = X14 | $less(alength,X14) | $less(X13,0) (3:24) [duplicate literal removal 3369]
[SA] new: 3407. $less(alength,X9) | $less(X10,0) | alength = X10 | X9 = X10 | i(l8(sK2(X10))) = X10 | $less(sK0,X10) | $less(X9,alength) (4:24) [duplicate literal removal 3367]
[SA] new: 3408. $less(alength,sK0) | alength = 0 | 0 = sK0 | i(l8(sK2(0))) = 0 (3:15) [evaluation 3362]
[SA] new: 3409. $less(alength,X18) | $less(X18,0) | alength = X18 | i(l8(sK2(X18))) = X18 | $less(alength,0) | alength = i(l8(sK2(alength))) (3:24) [duplicate literal removal 3352]
[SA] new: 3410. $less(alength,X16) | $less(X16,0) | alength = X16 | i(l8(sK2(X16))) = X16 | alength = X17 | $less(X17,0) | i(l8(sK2(X17))) = X17 | $less(alength,X17) (3:30) [duplicate literal removal 3351]
[SA] new: 3411. $less(alength,X1) | $less(X1,0) | alength = X1 | i(l8(sK2(X1))) = X1 (3:15) [duplicate literal removal 3344]
[SA] active: 2492. alength != 3 | s(s(s(zero))) = nl8 (4:9:1) [backward demodulation 2476,180]
[SA] active: 2951. $less(alength,X7) | $less(X8,0) | $less(X7,X8) | 0 = X8 | alength = X7 | i(l8(sK2(X7))) = X7 (2:21:3) [resolution 2656,99]
[SA] new: 3412. $less(X0,0) | $less(sK0,X0) | 0 = X0 | alength = sK0 | i(l8(sK2(sK0))) = sK0 | alength = sK0 (6:21) [resolution 2951,2704]
[SA] new: 3413. $less(X1,0) | $less(alength,X1) | 0 = X1 | alength = alength | alength = i(l8(sK2(alength))) (3:18) [resolution 2951,13]
[SA] new: 3414. $less(X2,0) | $less(X3,X2) | 0 = X2 | alength = X3 | i(l8(sK2(X3))) = X3 | ~$less(X4,alength) | $less(X4,X3) (3:24) [resolution 2951,14]
[SA] new: 3415. $less(X5,0) | $less(X6,X5) | 0 = X5 | alength = X6 | i(l8(sK2(X6))) = X6 | X6 = X7 | $less(X7,X6) | $less(alength,X7) (3:27) [resolution 2951,99]
[SA] new: 3416. $less(X8,0) | $less(X9,X8) | 0 = X8 | alength = X9 | i(l8(sK2(X9))) = X9 | sK0 = X9 | $less(sK0,X9) | $less(alength,alength) (4:27) [resolution 2951,205]
[SA] new: 3417. $less(X10,0) | $less(X11,X10) | 0 = X10 | alength = X11 | i(l8(sK2(X11))) = X11 | alength = X11 | $less(sK0,X11) | $less(alength,alength) (4:27) [resolution 2951,254]
[SA] new: 3418. $less(X12,0) | $less(X13,X12) | 0 = X12 | alength = X13 | i(l8(sK2(X13))) = X13 | sK0 = X13 | 0 = sK0 | $less(0,X13) | $less(alength,sK0) (4:30) [resolution 2951,333]
[SA] new: 3419. $less(X14,0) | $less(X15,X14) | 0 = X14 | alength = X15 | i(l8(sK2(X15))) = X15 | i(l8(sK2(X15))) = X15 | alength = X15 | $less(alength,X15) | $less(alength,0) (3:33) [resolution 2951,2950]
[SA] new: 3420. $less(X16,0) | $less(alength,X16) | 0 = X16 | alength = alength | alength = i(l8(sK2(alength))) | alength = X17 | $less(X17,0) | i(l8(sK2(X17))) = X17 | $less(alength,X17) (3:33) [resolution 2951,2660]
[SA] new: 3421. $less(X18,0) | $less(alength,X18) | 0 = X18 | alength = alength | alength = i(l8(sK2(alength))) | $less(alength,0) | alength = i(l8(sK2(alength))) (3:27) [resolution 2951,2480]
[SA] new: 3422. $less(X19,0) | $less(i(l8(p(nl8))),X19) | 0 = X19 | alength = i(l8(p(nl8))) | i(l8(p(nl8))) = i(l8(sK2(i(l8(p(nl8)))))) | zero = nl8 | $less(alength,alength) (4:36) [resolution 2951,194]
[SA] new: 3423. $less(X20,0) | $less(i(l8(sK2(X21))),X20) | 0 = X20 | alength = i(l8(sK2(X21))) | i(l8(sK2(X21))) = i(l8(sK2(i(l8(sK2(X21)))))) | $less(X21,0) | ~$less(X21,alength) | $less(alength,alength) (3:39) [resolution 2951,2920]
[SA] new: 3424. $less(X22,0) | $less(0,X22) | 0 = X22 | alength = 0 | i(l8(sK2(0))) = 0 | sK0 = X23 | 0 = sK0 | $less(X23,sK0) | $less(alength,X23) (4:30) [resolution 2951,324]
[SA] new: 3425. $less(X24,0) | $less(0,X24) | 0 = X24 | alength = 0 | i(l8(sK2(0))) = 0 | alength = 0 | $less(alength,alength) (5:24) [resolution 2951,262]
[SA] new: 3426. $less(X25,0) | $less(0,X25) | 0 = X25 | alength = 0 | i(l8(sK2(0))) = 0 | 0 = sK0 | $less(alength,alength) (4:24) [resolution 2951,116]
[SA] new: 3427. $less(X26,0) | $less(0,X26) | 0 = X26 | alength = 0 | i(l8(sK2(0))) = 0 | 0 = sK0 | $less(alength,sK0) (3:24) [resolution 2951,114]
[SA] new: 3428. $less(X27,0) | $less(sK0,X27) | 0 = X27 | alength = sK0 | i(l8(sK2(sK0))) = sK0 | alength = X28 | $less(X28,alength) | $less(alength,X28) (4:27) [resolution 2951,247]
[SA] new: 3429. $less(X29,0) | $less(sK0,X29) | 0 = X29 | alength = sK0 | i(l8(sK2(sK0))) = sK0 | sK0 = X30 | $less(X30,alength) | $less(alength,X30) (4:27) [resolution 2951,199]
[SA] new: 3430. $less(X31,0) | $less(sK0,X31) | 0 = X31 | alength = sK0 | i(l8(sK2(sK0))) = sK0 | $less(alength,alength) (3:21) [resolution 2951,96]
[SA] new: 3431. $less(alength,X0) | $less(X0,sK0) | 0 = sK0 | alength = X0 | i(l8(sK2(X0))) = X0 (3:18) [resolution 2951,46]
[SA] new: 3432. $less(alength,X1) | $less(X1,X2) | 0 = X2 | alength = X1 | i(l8(sK2(X1))) = X1 | sK0 = X3 | 0 = sK0 | $less(X3,sK0) | $less(X2,X3) (4:30) [resolution 2951,324]
[SA] new: 3433. $less(alength,X4) | $less(X4,X5) | 0 = X5 | alength = X4 | i(l8(sK2(X4))) = X4 | alength = 0 | $less(X5,alength) (5:24) [resolution 2951,262]
[SA] new: 3434. $less(alength,X6) | $less(X6,X7) | 0 = X7 | alength = X6 | i(l8(sK2(X6))) = X6 | 0 = sK0 | $less(X7,alength) (4:24) [resolution 2951,116]
[SA] new: 3435. $less(alength,X8) | $less(X8,X9) | 0 = X9 | alength = X8 | i(l8(sK2(X8))) = X8 | 0 = sK0 | $less(X9,sK0) (3:24) [resolution 2951,114]
[SA] new: 3436. $less(alength,X10) | $less(X10,0) | 0 = 0 | alength = X10 | i(l8(sK2(X10))) = X10 (3:18) [resolution 2951,13]
[SA] new: 3437. $less(alength,X11) | $less(X11,X12) | 0 = X12 | alength = X11 | i(l8(sK2(X11))) = X11 | ~$less(X13,X12) | $less(X13,0) (3:24) [resolution 2951,14]
[SA] new: 3438. $less(alength,X14) | $less(X14,X15) | 0 = X15 | alength = X14 | i(l8(sK2(X14))) = X14 | 0 = X16 | $less(X16,0) | $less(X15,X16) (3:27) [resolution 2951,99]
[SA] new: 3439. $less(alength,X17) | $less(X17,X18) | 0 = X18 | alength = X17 | i(l8(sK2(X17))) = X17 | 0 = sK0 | $less(sK0,0) | $less(X18,alength) (4:27) [resolution 2951,205]
[SA] new: 3440. $less(alength,X19) | $less(X19,X20) | 0 = X20 | alength = X19 | i(l8(sK2(X19))) = X19 | alength = 0 | $less(sK0,0) | $less(X20,alength) (4:27) [resolution 2951,254]
[SA] new: 3441. $less(alength,X21) | $less(X21,X22) | 0 = X22 | alength = X21 | i(l8(sK2(X21))) = X21 | 0 = sK0 | 0 = sK0 | $less(0,0) | $less(X22,sK0) (4:30) [resolution 2951,333]
[SA] new: 3442. $less(alength,X23) | $less(X23,X24) | 0 = X24 | alength = X23 | i(l8(sK2(X23))) = X23 | i(l8(sK2(0))) = 0 | alength = 0 | $less(alength,0) | $less(X24,0) (3:33) [resolution 2951,2950]
[SA] new: 3443. $less(alength,sK0) | $less(0,0) | 0 = 0 | alength = sK0 | i(l8(sK2(sK0))) = sK0 (3:18) [resolution 2951,46]
[SA] new: 3444. $less(alength,X0) | $less(X0,0) | 0 = X0 | alength = X0 | i(l8(sK2(X0))) = X0 (3:18) [resolution 2951,13]
[SA] new: 3445. $less(alength,X1) | $less(X2,0) | 0 = X2 | alength = X1 | i(l8(sK2(X1))) = X1 | ~$less(X3,X1) | $less(X3,X2) (3:24) [resolution 2951,14]
[SA] new: 3446. $less(alength,X4) | $less(X5,0) | 0 = X5 | alength = X4 | i(l8(sK2(X4))) = X4 | X5 = X6 | $less(X6,X5) | $less(X4,X6) (3:27) [resolution 2951,99]
[SA] new: 3447. $less(alength,X7) | $less(X8,0) | 0 = X8 | alength = X7 | i(l8(sK2(X7))) = X7 | sK0 = X8 | $less(sK0,X8) | $less(X7,alength) (4:27) [resolution 2951,205]
[SA] new: 3448. $less(alength,X9) | $less(X10,0) | 0 = X10 | alength = X9 | i(l8(sK2(X9))) = X9 | alength = X10 | $less(sK0,X10) | $less(X9,alength) (4:27) [resolution 2951,254]
[SA] new: 3449. $less(alength,X11) | $less(X12,0) | 0 = X12 | alength = X11 | i(l8(sK2(X11))) = X11 | sK0 = X12 | 0 = sK0 | $less(0,X12) | $less(X11,sK0) (4:30) [resolution 2951,333]
[SA] new: 3450. $less(alength,X13) | $less(X14,0) | 0 = X14 | alength = X13 | i(l8(sK2(X13))) = X13 | i(l8(sK2(X14))) = X14 | alength = X14 | $less(alength,X14) | $less(X13,0) (3:33) [resolution 2951,2950]
[SA] new: 3451. $less(alength,X15) | $less(alength,0) | alength = 0 | alength = X15 | i(l8(sK2(X15))) = X15 | alength = X16 | $less(X16,0) | i(l8(sK2(X16))) = X16 | $less(X15,X16) (3:33) [resolution 2951,2660]
[SA] new: 3452. $less(alength,X17) | $less(alength,0) | alength = 0 | alength = X17 | i(l8(sK2(X17))) = X17 | $less(X17,0) | i(l8(sK2(X17))) = X17 (3:27) [resolution 2951,2480]
[SA] new: 3453. $less(alength,X18) | $less(i(l8(p(nl8))),0) | i(l8(p(nl8))) = 0 | alength = X18 | i(l8(sK2(X18))) = X18 | zero = nl8 | $less(X18,alength) (4:30) [resolution 2951,194]
[SA] new: 3454. $less(alength,X19) | $less(i(l8(sK2(X20))),0) | i(l8(sK2(X20))) = 0 | alength = X19 | i(l8(sK2(X19))) = X19 | $less(X20,0) | ~$less(X20,alength) | $less(X19,alength) (3:33) [resolution 2951,2920]
[SA] new: 3455. $less(alength,X21) | $less(0,0) | 0 = 0 | alength = X21 | i(l8(sK2(X21))) = X21 | sK0 = X22 | 0 = sK0 | $less(X22,sK0) | $less(X21,X22) (4:30) [resolution 2951,324]
[SA] new: 3456. $less(alength,X23) | $less(0,0) | 0 = 0 | alength = X23 | i(l8(sK2(X23))) = X23 | alength = 0 | $less(X23,alength) (5:24) [resolution 2951,262]
[SA] new: 3457. $less(alength,X24) | $less(0,0) | 0 = 0 | alength = X24 | i(l8(sK2(X24))) = X24 | 0 = sK0 | $less(X24,alength) (4:24) [resolution 2951,116]
[SA] new: 3458. $less(alength,X25) | $less(0,0) | 0 = 0 | alength = X25 | i(l8(sK2(X25))) = X25 | 0 = sK0 | $less(X25,sK0) (3:24) [resolution 2951,114]
[SA] new: 3459. $less(alength,X26) | $less(sK0,0) | 0 = sK0 | alength = X26 | i(l8(sK2(X26))) = X26 | alength = X27 | $less(X27,alength) | $less(X26,X27) (4:27) [resolution 2951,247]
[SA] new: 3460. $less(alength,X28) | $less(sK0,0) | 0 = sK0 | alength = X28 | i(l8(sK2(X28))) = X28 | sK0 = X29 | $less(X29,alength) | $less(X28,X29) (4:27) [resolution 2951,199]
[SA] new: 3461. $less(alength,X30) | $less(sK0,0) | 0 = sK0 | alength = X30 | i(l8(sK2(X30))) = X30 | $less(X30,alength) (3:21) [resolution 2951,96]
[SA] new: 3462. $less(alength,alength) | $less(sK0,0) | 0 = sK0 | alength = alength | alength = i(l8(sK2(alength))) | alength = sK0 (6:21) [resolution 2951,2704]
[SA] new: 3463. $less(alength,0) | $less(0,alength) | alength = 0 | alength = 0 | i(l8(sK2(0))) = 0 (3:18) [factoring 2951]
[SA] new: 3464. $less(alength,alength) | $less(alength,0) | alength = 0 | alength = alength | alength = i(l8(sK2(alength))) (3:18) [factoring 2951]
[SA] new: 3465. $less(alength,0) | $less(0,0) | 0 = 0 | alength = 0 | i(l8(sK2(0))) = 0 (3:18) [factoring 2951]
[SA] new: 3466. $less(alength,0) | $less(0,alength) | alength = 0 | i(l8(sK2(0))) = 0 (3:15) [duplicate literal removal 3463]
[SA] new: 3467. $less(alength,X17) | $less(alength,0) | alength = 0 | alength = X17 | i(l8(sK2(X17))) = X17 | $less(X17,0) (3:21) [duplicate literal removal 3452]
[SA] new: 3468. $less(alength,X21) | $less(X21,X22) | 0 = X22 | alength = X21 | i(l8(sK2(X21))) = X21 | 0 = sK0 | $less(0,0) | $less(X22,sK0) (4:27) [duplicate literal removal 3441]
[SA] new: 3469. $less(alength,X21) | $less(X21,X22) | 0 = X22 | alength = X21 | i(l8(sK2(X21))) = X21 | 0 = sK0 | $less(X22,sK0) (4:24) [evaluation 3468]
[SA] new: 3470. $less(X24,0) | $less(0,X24) | 0 = X24 | alength = 0 | i(l8(sK2(0))) = 0 | $less(alength,alength) (5:21) [duplicate literal removal 3425]
[SA] new: 3471. $less(X18,0) | $less(alength,X18) | 0 = X18 | alength = alength | alength = i(l8(sK2(alength))) | $less(alength,0) (3:21) [duplicate literal removal 3421]
[SA] new: 3472. $less(X14,0) | $less(X15,X14) | 0 = X14 | alength = X15 | i(l8(sK2(X15))) = X15 | $less(alength,X15) | $less(alength,0) (3:24) [duplicate literal removal 3419]
[SA] new: 3473. $less(X10,0) | $less(X11,X10) | 0 = X10 | alength = X11 | i(l8(sK2(X11))) = X11 | $less(sK0,X11) | $less(alength,alength) (4:24) [duplicate literal removal 3417]
[SA] new: 3474. $less(X0,0) | $less(sK0,X0) | 0 = X0 | alength = sK0 | i(l8(sK2(sK0))) = sK0 (6:18) [duplicate literal removal 3412]
[SA] passive: 3414. $less(X2,0) | $less(X3,X2) | 0 = X2 | alength = X3 | i(l8(sK2(X3))) = X3 | ~$less(X4,alength) | $less(X4,X3) (3:24) [resolution 2951,14]
[SA] new: 3475. $less(X8,0) | $less(X9,X8) | 0 = X8 | alength = X9 | i(l8(sK2(X9))) = X9 | sK0 = X9 | $less(sK0,X9) (4:24) [subsumption resolution 3416,13]
[SA] passive: 3475. $less(X8,0) | $less(X9,X8) | 0 = X8 | alength = X9 | i(l8(sK2(X9))) = X9 | sK0 = X9 | $less(sK0,X9) (4:24) [subsumption resolution 3416,13]
[SA] new: 3476. $less(X10,0) | $less(X11,X10) | 0 = X10 | alength = X11 | i(l8(sK2(X11))) = X11 | $less(sK0,X11) (4:21) [subsumption resolution 3473,13]
[SA] passive: 3476. $less(X10,0) | $less(X11,X10) | 0 = X10 | alength = X11 | i(l8(sK2(X11))) = X11 | $less(sK0,X11) (4:21) [subsumption resolution 3473,13]
[SA] passive: 3418. $less(X12,0) | $less(X13,X12) | 0 = X12 | alength = X13 | i(l8(sK2(X13))) = X13 | sK0 = X13 | 0 = sK0 | $less(0,X13) | $less(alength,sK0) (4:30) [resolution 2951,333]
[SA] new: 3477. $less(X19,0) | $less(i(l8(p(nl8))),X19) | 0 = X19 | alength = i(l8(p(nl8))) | i(l8(p(nl8))) = i(l8(sK2(i(l8(p(nl8)))))) | zero = nl8 (4:33) [subsumption resolution 3422,13]
[SA] passive: 3477. $less(X19,0) | $less(i(l8(p(nl8))),X19) | 0 = X19 | alength = i(l8(p(nl8))) | i(l8(p(nl8))) = i(l8(sK2(i(l8(p(nl8)))))) | zero = nl8 (4:33) [subsumption resolution 3422,13]
[SA] new: 3478. $less(X20,0) | $less(i(l8(sK2(X21))),X20) | 0 = X20 | alength = i(l8(sK2(X21))) | i(l8(sK2(X21))) = i(l8(sK2(i(l8(sK2(X21)))))) | $less(X21,0) | ~$less(X21,alength) (3:36) [subsumption resolution 3423,13]
[SA] passive: 3478. $less(X20,0) | $less(i(l8(sK2(X21))),X20) | 0 = X20 | alength = i(l8(sK2(X21))) | i(l8(sK2(X21))) = i(l8(sK2(i(l8(sK2(X21)))))) | $less(X21,0) | ~$less(X21,alength) (3:36) [subsumption resolution 3423,13]
[SA] passive: 3432. $less(alength,X1) | $less(X1,X2) | 0 = X2 | alength = X1 | i(l8(sK2(X1))) = X1 | sK0 = X3 | 0 = sK0 | $less(X3,sK0) | $less(X2,X3) (4:30) [resolution 2951,324]
[SA] passive: 3433. $less(alength,X4) | $less(X4,X5) | 0 = X5 | alength = X4 | i(l8(sK2(X4))) = X4 | alength = 0 | $less(X5,alength) (5:24) [resolution 2951,262]
[SA] passive: 3434. $less(alength,X6) | $less(X6,X7) | 0 = X7 | alength = X6 | i(l8(sK2(X6))) = X6 | 0 = sK0 | $less(X7,alength) (4:24) [resolution 2951,116]
[SA] passive: 3435. $less(alength,X8) | $less(X8,X9) | 0 = X9 | alength = X8 | i(l8(sK2(X8))) = X8 | 0 = sK0 | $less(X9,sK0) (3:24) [resolution 2951,114]
[SA] passive: 3437. $less(alength,X11) | $less(X11,X12) | 0 = X12 | alength = X11 | i(l8(sK2(X11))) = X11 | ~$less(X13,X12) | $less(X13,0) (3:24) [resolution 2951,14]
[SA] passive: 3438. $less(alength,X14) | $less(X14,X15) | 0 = X15 | alength = X14 | i(l8(sK2(X14))) = X14 | 0 = X16 | $less(X16,0) | $less(X15,X16) (3:27) [resolution 2951,99]
[SA] passive: 3445. $less(alength,X1) | $less(X2,0) | 0 = X2 | alength = X1 | i(l8(sK2(X1))) = X1 | ~$less(X3,X1) | $less(X3,X2) (3:24) [resolution 2951,14]
[SA] passive: 3446. $less(alength,X4) | $less(X5,0) | 0 = X5 | alength = X4 | i(l8(sK2(X4))) = X4 | X5 = X6 | $less(X6,X5) | $less(X4,X6) (3:27) [resolution 2951,99]
[SA] new: 3479. $less(alength,X28) | 0 = sK0 | alength = X28 | i(l8(sK2(X28))) = X28 | sK0 = X29 | $less(X29,alength) | $less(X28,X29) (4:24) [subsumption resolution 3460,46]
[SA] passive: 3479. $less(alength,X28) | 0 = sK0 | alength = X28 | i(l8(sK2(X28))) = X28 | sK0 = X29 | $less(X29,alength) | $less(X28,X29) (4:24) [subsumption resolution 3460,46]
[SA] active: 2845. sK0 != 2 | s(s(zero)) = sK2(sK0) (4:9:1) [superposition 130,2833]
[SA] active: 2770. ~$less(X0,alength) | $less(X0,0) | a(main_end,X1) = a(l11(sK2(X0)),X1) | i(l8(sK1(X1,sK2(X0)))) = X1 (2:23:1) [resolution 2479,150]
[SA] new: 3480. $less(X0,0) | a(main_end,X1) = a(l11(sK2(X0)),X1) | i(l8(sK1(X1,sK2(X0)))) = X1 | $less(sK0,X0) | sK0 = X0 (3:26) [resolution 2770,106]
[SA] new: 3481. $less(X2,0) | a(main_end,X3) = a(l11(sK2(X2)),X3) | i(l8(sK1(X3,sK2(X2)))) = X3 | $less(0,X2) | alength = X2 | 0 = sK0 (4:29) [resolution 2770,178]
[SA] new: 3482. $less(X4,0) | a(main_end,X5) = a(l11(sK2(X4)),X5) | i(l8(sK1(X5,sK2(X4)))) = X5 | $less(sK0,X4) | alength = X4 (3:26) [resolution 2770,179]
[SA] new: 3483. $less(X6,0) | a(main_end,X7) = a(l11(sK2(X6)),X7) | i(l8(sK1(X7,sK2(X6)))) = X7 | $less(sK0,X8) | $less(X8,X6) | X6 = X8 | sK0 = X6 (4:32) [resolution 2770,198]
[SA] new: 3484. $less(X9,0) | a(main_end,X10) = a(l11(sK2(X9)),X10) | i(l8(sK1(X10,sK2(X9)))) = X10 | $less(sK0,X11) | $less(X11,X9) | alength = X9 | sK0 = X11 (4:32) [resolution 2770,204]
[SA] new: 3485. $less(X12,0) | a(main_end,X13) = a(l11(sK2(X12)),X13) | i(l8(sK1(X13,sK2(X12)))) = X13 | $less(sK0,X14) | $less(X14,X12) | X12 = X14 | alength = X12 (4:32) [resolution 2770,246]
[SA] new: 3486. $less(X15,0) | a(main_end,X16) = a(l11(sK2(X15)),X16) | i(l8(sK1(X16,sK2(X15)))) = X16 | $less(sK0,X17) | $less(X17,X15) | alength = X15 | alength = X17 (4:32) [resolution 2770,253]
[SA] new: 3487. $less(X18,0) | a(main_end,X19) = a(l11(sK2(X18)),X19) | i(l8(sK1(X19,sK2(X18)))) = X19 | $less(0,X18) | alength = X18 | alength = 0 (5:29) [resolution 2770,261]
[SA] new: 3488. $less(X20,0) | a(main_end,X21) = a(l11(sK2(X20)),X21) | i(l8(sK1(X21,sK2(X20)))) = X21 | $less(0,X20) | 0 = sK0 | 0 = X20 (5:29) [resolution 2770,278]
[SA] new: 3489. $less(X22,0) | a(main_end,X23) = a(l11(sK2(X22)),X23) | i(l8(sK1(X23,sK2(X22)))) = X23 | $less(0,X22) | 0 = sK0 | sK0 = X22 (4:29) [resolution 2770,330]
[SA] new: 3490. $less(X24,0) | a(main_end,X25) = a(l11(sK2(X24)),X25) | i(l8(sK1(X25,sK2(X24)))) = X25 | $less(sK0,X26) | $less(X26,X24) | X24 = X26 | sK0 = X26 (4:32) [resolution 2770,388]
[SA] new: 3491. $less(X27,0) | a(main_end,X28) = a(l11(sK2(X27)),X28) | i(l8(sK1(X28,sK2(X27)))) = X28 | $less(sK0,X29) | $less(X29,X27) | sK0 = X29 | sK0 = X27 (4:32) [resolution 2770,408]
[SA] new: 3492. $less(X30,0) | a(main_end,X31) = a(l11(sK2(X30)),X31) | i(l8(sK1(X31,sK2(X30)))) = X31 | $less(alength,X30) | alength = X30 (3:26) [resolution 2770,15]
[SA] new: 3493. $less(X32,0) | a(main_end,X33) = a(l11(sK2(X32)),X33) | i(l8(sK1(X33,sK2(X32)))) = X33 | $less(alength,X32) | alength = X32 (3:26) [resolution 2770,15]
[SA] new: 3494. $less(X34,0) | a(main_end,X35) = a(l11(sK2(X34)),X35) | i(l8(sK1(X35,sK2(X34)))) = X35 | $less(X36,X34) | $less(alength,X36) | alength = X34 | alength = X36 (3:32) [resolution 2770,174]
[SA] new: 3495. $less(X37,0) | a(main_end,X38) = a(l11(sK2(X37)),X38) | i(l8(sK1(X38,sK2(X37)))) = X38 | $less(alength,X39) | $less(X39,X37) | X37 = X39 | alength = X37 (3:32) [resolution 2770,174]
[SA] new: 3496. $less(X40,0) | a(main_end,X41) = a(l11(sK2(X40)),X41) | i(l8(sK1(X41,sK2(X40)))) = X41 | $less(X42,X40) | $less(alength,X42) | alength = X42 | X40 = X42 (3:32) [resolution 2770,174]
[SA] new: 3497. $less(X43,0) | a(main_end,X44) = a(l11(sK2(X43)),X44) | i(l8(sK1(X44,sK2(X43)))) = X44 | $less(sK0,X43) | $less(alength,alength) | alength = X43 | alength = sK0 (4:32) [resolution 2770,198]
[SA] new: 3498. $less(X45,0) | a(main_end,X46) = a(l11(sK2(X45)),X46) | i(l8(sK1(X46,sK2(X45)))) = X46 | $less(sK0,X45) | $less(alength,alength) | alength = alength | sK0 = X45 (4:32) [resolution 2770,204]
[SA] new: 3499. $less(X47,0) | a(main_end,X48) = a(l11(sK2(X47)),X48) | i(l8(sK1(X48,sK2(X47)))) = X48 | $less(sK0,X47) | $less(alength,alength) | alength = X47 | alength = alength (4:32) [resolution 2770,246]
[SA] new: 3500. $less(X49,0) | a(main_end,X50) = a(l11(sK2(X49)),X50) | i(l8(sK1(X50,sK2(X49)))) = X50 | $less(sK0,X49) | $less(alength,alength) | alength = alength | alength = X49 (4:32) [resolution 2770,253]
[SA] new: 3501. $less(X51,0) | a(main_end,X52) = a(l11(sK2(X51)),X52) | i(l8(sK1(X52,sK2(X51)))) = X52 | $less(0,X51) | $less(alength,sK0) | alength = X51 | 0 = sK0 | alength = sK0 (4:35) [resolution 2770,323]
[SA] new: 3502. $less(X53,0) | a(main_end,X54) = a(l11(sK2(X53)),X54) | i(l8(sK1(X54,sK2(X53)))) = X54 | $less(0,X53) | $less(alength,sK0) | alength = sK0 | 0 = sK0 | sK0 = X53 (4:35) [resolution 2770,332]
[SA] new: 3503. $less(X55,0) | a(main_end,X56) = a(l11(sK2(X55)),X56) | i(l8(sK1(X56,sK2(X55)))) = X56 | $less(0,X55) | $less(alength,sK0) | alength = X55 | 0 = sK0 | 0 = X55 (4:35) [resolution 2770,387]
[SA] new: 3504. $less(X57,0) | a(main_end,X58) = a(l11(sK2(X57)),X58) | i(l8(sK1(X58,sK2(X57)))) = X58 | $less(sK0,X57) | $less(alength,alength) | alength = X57 | sK0 = X57 (4:32) [resolution 2770,388]
[SA] new: 3505. $less(X59,0) | a(main_end,X60) = a(l11(sK2(X59)),X60) | i(l8(sK1(X60,sK2(X59)))) = X60 | $less(sK0,X59) | $less(alength,alength) | sK0 = X59 | alength = sK0 (4:32) [resolution 2770,408]
[SA] new: 3506. $less(X61,0) | a(main_end,X62) = a(l11(sK2(X61)),X62) | i(l8(sK1(X62,sK2(X61)))) = X62 | $less(alength,X61) | $less(alength,0) | alength = alength | alength = X61 | alength = i(l8(sK2(alength))) (3:38) [resolution 2770,2659]
[SA] new: 3507. $less(X63,0) | a(main_end,X64) = a(l11(sK2(X63)),X64) | i(l8(sK1(X64,sK2(X63)))) = X64 | $less(alength,X63) | $less(alength,0) | alength = 0 | alength = X63 | i(l8(sK2(X63))) = X63 (3:38) [resolution 2770,2951]
[SA] new: 3508. $less(alength,0) | a(main_end,X65) = a(l11(sK2(alength)),X65) | i(l8(sK1(X65,sK2(alength)))) = X65 | $less(X66,0) | $less(alength,X66) | 0 = X66 | alength = alength | alength = i(l8(sK2(alength))) (3:38) [resolution 2770,2951]
[SA] new: 3509. $less(alength,0) | a(main_end,X67) = a(l11(sK2(alength)),X67) | i(l8(sK1(X67,sK2(alength)))) = X67 | $less(alength,X68) | $less(X68,0) | alength = X68 | alength = X68 | i(l8(sK2(X68))) = X68 (3:38) [resolution 2770,2659]
[SA] new: 3510. $less(alength,0) | a(main_end,X69) = a(l11(sK2(alength)),X69) | i(l8(sK1(X69,sK2(alength)))) = X69 | $less(alength,0) | alength = i(l8(sK2(alength))) | alength = alength (3:32) [resolution 2770,2656]
[SA] new: 3511. $less(i(l8(p(nl8))),0) | a(main_end,X70) = a(l11(sK2(i(l8(p(nl8))))),X70) | i(l8(sK1(X70,sK2(i(l8(p(nl8))))))) = X70 | zero = nl8 (3:32) [resolution 2770,95]
[SA] new: 3512. $less(i(l8(p(p(nl8)))),0) | a(main_end,X71) = a(l11(sK2(i(l8(p(p(nl8)))))),X71) | i(l8(sK1(X71,sK2(i(l8(p(p(nl8)))))))) = X71 | zero = p(nl8) | zero = nl8 (4:39) [resolution 2770,314]
[SA] new: 3513. $less(i(l8(sK2(X72))),0) | a(main_end,X73) = a(l11(sK2(i(l8(sK2(X72))))),X73) | i(l8(sK1(X73,sK2(i(l8(sK2(X72))))))) = X73 | $less(alength,X72) | alength = X72 | $less(X72,0) (3:38) [resolution 2770,2657]
[SA] new: 3514. $less(i(l8(sK2(X74))),0) | a(main_end,X75) = a(l11(sK2(i(l8(sK2(X74))))),X75) | i(l8(sK1(X75,sK2(i(l8(sK2(X74))))))) = X75 | ~$less(X74,alength) | $less(X74,0) (3:35) [resolution 2770,2771]
[SA] new: 3515. $less(i(l8(p(nl8))),0) | a(main_end,X76) = a(l11(sK2(i(l8(p(nl8))))),X76) | i(l8(sK1(X76,sK2(i(l8(p(nl8))))))) = X76 | alength = alength | $less(alength,alength) | zero = nl8 (4:38) [resolution 2770,193]
[SA] new: 3516. $less(i(l8(sK2(X77))),0) | a(main_end,X78) = a(l11(sK2(i(l8(sK2(X77))))),X78) | i(l8(sK1(X78,sK2(i(l8(sK2(X77))))))) = X78 | $less(X77,0) | alength = alength | $less(alength,alength) | ~$less(X77,alength) (3:41) [resolution 2770,2921]
[SA] new: 3517. $less(0,0) | a(main_end,X79) = a(l11(sK2(0)),X79) | i(l8(sK1(X79,sK2(0)))) = X79 | alength = 0 (4:23) [resolution 2770,244]
[SA] new: 3518. $less(0,0) | a(main_end,X80) = a(l11(sK2(0)),X80) | i(l8(sK1(X80,sK2(0)))) = X80 | 0 = sK0 (3:23) [resolution 2770,113]
[SA] new: 3519. $less(0,0) | a(main_end,X81) = a(l11(sK2(0)),X81) | i(l8(sK1(X81,sK2(0)))) = X81 | $less(alength,sK0) | alength = sK0 | 0 = sK0 (3:29) [resolution 2770,177]
[SA] new: 3520. $less(0,0) | a(main_end,X82) = a(l11(sK2(0)),X82) | i(l8(sK1(X82,sK2(0)))) = X82 | $less(alength,alength) | alength = alength | 0 = sK0 (4:29) [resolution 2770,178]
[SA] new: 3521. $less(0,0) | a(main_end,X83) = a(l11(sK2(0)),X83) | i(l8(sK1(X83,sK2(0)))) = X83 | $less(alength,sK0) | 0 = sK0 | alength = 0 (4:29) [resolution 2770,229]
[SA] new: 3522. $less(0,0) | a(main_end,X84) = a(l11(sK2(0)),X84) | i(l8(sK1(X84,sK2(0)))) = X84 | $less(alength,alength) | alength = alength | alength = 0 (5:29) [resolution 2770,261]
[SA] new: 3523. $less(0,0) | a(main_end,X85) = a(l11(sK2(0)),X85) | i(l8(sK1(X85,sK2(0)))) = X85 | $less(alength,alength) | 0 = sK0 | alength = 0 (5:29) [resolution 2770,278]
[SA] new: 3524. $less(0,0) | a(main_end,X86) = a(l11(sK2(0)),X86) | i(l8(sK1(X86,sK2(0)))) = X86 | $less(alength,X87) | $less(X87,sK0) | alength = X87 | 0 = sK0 | sK0 = X87 (4:35) [resolution 2770,323]
[SA] new: 3525. $less(0,0) | a(main_end,X88) = a(l11(sK2(0)),X88) | i(l8(sK1(X88,sK2(0)))) = X88 | $less(alength,alength) | 0 = sK0 | alength = sK0 (4:29) [resolution 2770,330]
[SA] new: 3526. $less(0,0) | a(main_end,X89) = a(l11(sK2(0)),X89) | i(l8(sK1(X89,sK2(0)))) = X89 | $less(X90,sK0) | $less(alength,X90) | sK0 = X90 | 0 = sK0 | alength = sK0 (4:35) [resolution 2770,332]
[SA] new: 3527. $less(0,0) | a(main_end,X91) = a(l11(sK2(0)),X91) | i(l8(sK1(X91,sK2(0)))) = X91 | $less(alength,sK0) | alength = 0 | alength = sK0 (4:29) [resolution 2770,380]
[SA] new: 3528. $less(0,0) | a(main_end,X92) = a(l11(sK2(0)),X92) | i(l8(sK1(X92,sK2(0)))) = X92 | $less(X93,sK0) | $less(alength,X93) | alength = X93 | 0 = sK0 | alength = 0 (4:35) [resolution 2770,387]
[SA] new: 3529. $less(sK0,0) | a(main_end,X94) = a(l11(sK2(sK0)),X94) | i(l8(sK1(X94,sK2(sK0)))) = X94 (3:20) [resolution 2770,45]
[SA] new: 3530. $less(sK0,0) | a(main_end,X95) = a(l11(sK2(sK0)),X95) | i(l8(sK1(X95,sK2(sK0)))) = X95 | $less(alength,alength) | alength = sK0 (3:26) [resolution 2770,106]
[SA] new: 3531. $less(sK0,0) | a(main_end,X96) = a(l11(sK2(sK0)),X96) | i(l8(sK1(X96,sK2(sK0)))) = X96 | $less(alength,alength) | alength = alength (3:26) [resolution 2770,179]
[SA] new: 3532. $less(sK0,0) | a(main_end,X97) = a(l11(sK2(sK0)),X97) | i(l8(sK1(X97,sK2(sK0)))) = X97 | $less(alength,X98) | $less(X98,alength) | alength = X98 | sK0 = X98 (4:32) [resolution 2770,198]
[SA] new: 3533. $less(sK0,0) | a(main_end,X99) = a(l11(sK2(sK0)),X99) | i(l8(sK1(X99,sK2(sK0)))) = X99 | $less(X100,alength) | $less(alength,X100) | alength = X100 | alength = sK0 (4:32) [resolution 2770,204]
[SA] new: 3534. $less(sK0,0) | a(main_end,X101) = a(l11(sK2(sK0)),X101) | i(l8(sK1(X101,sK2(sK0)))) = X101 | $less(alength,X102) | $less(X102,alength) | alength = X102 | alength = X102 (4:32) [resolution 2770,246]
[SA] new: 3535. $less(sK0,0) | a(main_end,X103) = a(l11(sK2(sK0)),X103) | i(l8(sK1(X103,sK2(sK0)))) = X103 | $less(X104,alength) | $less(alength,X104) | alength = X104 | alength = alength (4:32) [resolution 2770,253]
[SA] new: 3536. $less(sK0,0) | a(main_end,X105) = a(l11(sK2(sK0)),X105) | i(l8(sK1(X105,sK2(sK0)))) = X105 | $less(X106,alength) | $less(alength,X106) | alength = X106 | alength = sK0 (4:32) [resolution 2770,388]
[SA] new: 3537. $less(sK0,0) | a(main_end,X107) = a(l11(sK2(sK0)),X107) | i(l8(sK1(X107,sK2(sK0)))) = X107 | $less(X108,alength) | $less(alength,X108) | alength = sK0 | sK0 = X108 (4:32) [resolution 2770,408]
[SA] new: 3538. $less(sK0,0) | a(main_end,X101) = a(l11(sK2(sK0)),X101) | i(l8(sK1(X101,sK2(sK0)))) = X101 | $less(alength,X102) | $less(X102,alength) | alength = X102 (4:29) [duplicate literal removal 3534]
[SA] new: 3539. a(main_end,X92) = a(l11(sK2(0)),X92) | i(l8(sK1(X92,sK2(0)))) = X92 | $less(X93,sK0) | $less(alength,X93) | alength = X93 | 0 = sK0 | alength = 0 (4:32) [evaluation 3528]
[SA] new: 3540. a(main_end,X91) = a(l11(sK2(0)),X91) | i(l8(sK1(X91,sK2(0)))) = X91 | $less(alength,sK0) | alength = 0 | alength = sK0 (4:26) [evaluation 3527]
[SA] new: 3541. a(main_end,X89) = a(l11(sK2(0)),X89) | i(l8(sK1(X89,sK2(0)))) = X89 | $less(X90,sK0) | $less(alength,X90) | sK0 = X90 | 0 = sK0 | alength = sK0 (4:32) [evaluation 3526]
[SA] new: 3542. a(main_end,X88) = a(l11(sK2(0)),X88) | i(l8(sK1(X88,sK2(0)))) = X88 | $less(alength,alength) | 0 = sK0 | alength = sK0 (4:26) [evaluation 3525]
[SA] new: 3543. a(main_end,X86) = a(l11(sK2(0)),X86) | i(l8(sK1(X86,sK2(0)))) = X86 | $less(alength,X87) | $less(X87,sK0) | alength = X87 | 0 = sK0 | sK0 = X87 (4:32) [evaluation 3524]
[SA] new: 3544. a(main_end,X85) = a(l11(sK2(0)),X85) | i(l8(sK1(X85,sK2(0)))) = X85 | $less(alength,alength) | 0 = sK0 | alength = 0 (5:26) [evaluation 3523]
[SA] new: 3545. a(main_end,X83) = a(l11(sK2(0)),X83) | i(l8(sK1(X83,sK2(0)))) = X83 | $less(alength,sK0) | 0 = sK0 | alength = 0 (4:26) [evaluation 3521]
[SA] new: 3546. a(main_end,X81) = a(l11(sK2(0)),X81) | i(l8(sK1(X81,sK2(0)))) = X81 | $less(alength,sK0) | alength = sK0 | 0 = sK0 (3:26) [evaluation 3519]
[SA] new: 3547. a(main_end,X80) = a(l11(sK2(0)),X80) | i(l8(sK1(X80,sK2(0)))) = X80 | 0 = sK0 (3:20) [evaluation 3518]
[SA] new: 3548. a(main_end,X79) = a(l11(sK2(0)),X79) | i(l8(sK1(X79,sK2(0)))) = X79 | alength = 0 (4:20) [evaluation 3517]
[SA] new: 3549. $less(alength,0) | a(main_end,X69) = a(l11(sK2(alength)),X69) | i(l8(sK1(X69,sK2(alength)))) = X69 | alength = i(l8(sK2(alength))) | alength = alength (3:29) [duplicate literal removal 3510]
[SA] new: 3550. $less(alength,0) | a(main_end,X67) = a(l11(sK2(alength)),X67) | i(l8(sK1(X67,sK2(alength)))) = X67 | $less(alength,X68) | $less(X68,0) | alength = X68 | i(l8(sK2(X68))) = X68 (3:35) [duplicate literal removal 3509]
[SA] passive: 3480. $less(X0,0) | a(main_end,X1) = a(l11(sK2(X0)),X1) | i(l8(sK1(X1,sK2(X0)))) = X1 | $less(sK0,X0) | sK0 = X0 (3:26) [resolution 2770,106]
[SA] passive: 3481. $less(X2,0) | a(main_end,X3) = a(l11(sK2(X2)),X3) | i(l8(sK1(X3,sK2(X2)))) = X3 | $less(0,X2) | alength = X2 | 0 = sK0 (4:29) [resolution 2770,178]
[SA] passive: 3483. $less(X6,0) | a(main_end,X7) = a(l11(sK2(X6)),X7) | i(l8(sK1(X7,sK2(X6)))) = X7 | $less(sK0,X8) | $less(X8,X6) | X6 = X8 | sK0 = X6 (4:32) [resolution 2770,198]
[SA] passive: 3484. $less(X9,0) | a(main_end,X10) = a(l11(sK2(X9)),X10) | i(l8(sK1(X10,sK2(X9)))) = X10 | $less(sK0,X11) | $less(X11,X9) | alength = X9 | sK0 = X11 (4:32) [resolution 2770,204]
[SA] passive: 3485. $less(X12,0) | a(main_end,X13) = a(l11(sK2(X12)),X13) | i(l8(sK1(X13,sK2(X12)))) = X13 | $less(sK0,X14) | $less(X14,X12) | X12 = X14 | alength = X12 (4:32) [resolution 2770,246]
[SA] passive: 3486. $less(X15,0) | a(main_end,X16) = a(l11(sK2(X15)),X16) | i(l8(sK1(X16,sK2(X15)))) = X16 | $less(sK0,X17) | $less(X17,X15) | alength = X15 | alength = X17 (4:32) [resolution 2770,253]
[SA] passive: 3487. $less(X18,0) | a(main_end,X19) = a(l11(sK2(X18)),X19) | i(l8(sK1(X19,sK2(X18)))) = X19 | $less(0,X18) | alength = X18 | alength = 0 (5:29) [resolution 2770,261]
[SA] passive: 3489. $less(X22,0) | a(main_end,X23) = a(l11(sK2(X22)),X23) | i(l8(sK1(X23,sK2(X22)))) = X23 | $less(0,X22) | 0 = sK0 | sK0 = X22 (4:29) [resolution 2770,330]
[SA] passive: 3490. $less(X24,0) | a(main_end,X25) = a(l11(sK2(X24)),X25) | i(l8(sK1(X25,sK2(X24)))) = X25 | $less(sK0,X26) | $less(X26,X24) | X24 = X26 | sK0 = X26 (4:32) [resolution 2770,388]
[SA] passive: 3491. $less(X27,0) | a(main_end,X28) = a(l11(sK2(X27)),X28) | i(l8(sK1(X28,sK2(X27)))) = X28 | $less(sK0,X29) | $less(X29,X27) | sK0 = X29 | sK0 = X27 (4:32) [resolution 2770,408]
[SA] passive: 3494. $less(X34,0) | a(main_end,X35) = a(l11(sK2(X34)),X35) | i(l8(sK1(X35,sK2(X34)))) = X35 | $less(X36,X34) | $less(alength,X36) | alength = X34 | alength = X36 (3:32) [resolution 2770,174]
[SA] passive: 3496. $less(X40,0) | a(main_end,X41) = a(l11(sK2(X40)),X41) | i(l8(sK1(X41,sK2(X40)))) = X41 | $less(X42,X40) | $less(alength,X42) | alength = X42 | X40 = X42 (3:32) [resolution 2770,174]
[SA] passive: 3511. $less(i(l8(p(nl8))),0) | a(main_end,X70) = a(l11(sK2(i(l8(p(nl8))))),X70) | i(l8(sK1(X70,sK2(i(l8(p(nl8))))))) = X70 | zero = nl8 (3:32) [resolution 2770,95]
[SA] passive: 3512. $less(i(l8(p(p(nl8)))),0) | a(main_end,X71) = a(l11(sK2(i(l8(p(p(nl8)))))),X71) | i(l8(sK1(X71,sK2(i(l8(p(p(nl8)))))))) = X71 | zero = p(nl8) | zero = nl8 (4:39) [resolution 2770,314]
[SA] passive: 3513. $less(i(l8(sK2(X72))),0) | a(main_end,X73) = a(l11(sK2(i(l8(sK2(X72))))),X73) | i(l8(sK1(X73,sK2(i(l8(sK2(X72))))))) = X73 | $less(alength,X72) | alength = X72 | $less(X72,0) (3:38) [resolution 2770,2657]
[SA] passive: 3514. $less(i(l8(sK2(X74))),0) | a(main_end,X75) = a(l11(sK2(i(l8(sK2(X74))))),X75) | i(l8(sK1(X75,sK2(i(l8(sK2(X74))))))) = X75 | ~$less(X74,alength) | $less(X74,0) (3:35) [resolution 2770,2771]
[SA] passive: 3547. a(main_end,X80) = a(l11(sK2(0)),X80) | i(l8(sK1(X80,sK2(0)))) = X80 | 0 = sK0 (3:20) [evaluation 3518]
[SA] new: 3551. a(main_end,X94) = a(l11(sK2(sK0)),X94) | i(l8(sK1(X94,sK2(sK0)))) = X94 (3:17) [subsumption resolution 3529,46]
[SA] passive: 3551. a(main_end,X94) = a(l11(sK2(sK0)),X94) | i(l8(sK1(X94,sK2(sK0)))) = X94 (3:17) [subsumption resolution 3529,46]
[SA] active: 2654. i(l8(sK2(0))) = 0 | alength = 0 (5:9:1) [duplicate literal removal 2517]
[SA] new: 3552. i(l8(zero)) = 0 | alength = 0 | 0 = sK0 (6:11) [superposition 2654,3207]
[SA] new: 3553. $less(0,X0) | $less(0,0) | alength = X0 | $less(X0,alength) | ~$less(0,alength) | alength = 0 (6:18) [superposition 2921,2654]
[SA] new: 3554. ~$less(X1,0) | $less(0,0) | ~$less(0,alength) | $less(X1,alength) | alength = 0 (6:15) [superposition 2920,2654]
[SA] new: 3555. $less(0,alength) | $less(alength,0) | alength = 0 | $less(0,0) | alength = 0 (6:15) [superposition 2657,2654]
[SA] new: 3556. $less(0,alength) | ~$less(0,alength) | $less(0,0) | alength = 0 (6:12) [superposition 2771,2654]
[SA] new: 3557. i(l8(X2)) != 0 | sK2(0) = X2 | alength = 0 (6:12) [superposition 48,2654]
[SA] new: 3558. i(l8(X3)) != 0 | sK2(0) = X3 | alength = 0 (6:12) [superposition 48,2654]
[SA] new: 3559. v = a(l11(sK2(0)),0) | alength = 0 (6:10) [superposition 70,2654]
[SA] new: 3560. i(l8(s(sK2(0)))) = $sum(0,1) | alength = 0 (6:12) [superposition 74,2654]
[SA] new: 3561. 0 != 0 | zero = sK2(0) | alength = 0 (6:10) [superposition 108,2654]
[SA] new: 3562. 1 != 0 | s(zero) = sK2(0) | alength = 0 (6:11) [superposition 121,2654]
[SA] new: 3563. 0 != 2 | s(s(zero)) = sK2(0) | alength = 0 (6:12) [superposition 130,2654]
[SA] new: 3564. 0 != 3 | s(s(s(zero))) = sK2(0) | alength = 0 (6:13) [superposition 183,2654]
[SA] new: 3565. v = a(l8(s(sK2(0))),0) | alength = 0 (6:11) [superposition 349,2654]
[SA] new: 3566. alength != 0 | nl8 = sK2(0) | alength = 0 (6:10) [superposition 2483,2654]
[SA] new: 3567. $sum(alength,1) != 0 | s(nl8) = sK2(0) | alength = 0 (6:13) [superposition 2490,2654]
[SA] new: 3568. 0 != sK0 | sK2(0) = sK2(sK0) | alength = 0 (6:11) [superposition 2839,2654]
[SA] new: 3569. alength != $uminus(1) | s(nl8) = sK2(0) | alength = 0 (6:12) [evaluation 3567]
[SA] new: 3570. alength != -1 | s(nl8) = sK2(0) | alength = 0 (6:11) [evaluation 3569]
[SA] new: 3571. zero = sK2(0) | alength = 0 (6:7) [trivial inequality removal 3561]
[SA] new: 3572. i(l8(s(sK2(0)))) = 1 | alength = 0 (6:10) [evaluation 3560]
[SA] new: 3573. $less(0,alength) | $less(alength,0) | alength = 0 | $less(0,0) (6:12) [duplicate literal removal 3555]
[SA] new: 3574. $less(0,alength) | $less(alength,0) | alength = 0 (6:9) [evaluation 3573]
[SA] new: 3575. ~$less(X1,0) | ~$less(0,alength) | $less(X1,alength) | alength = 0 (6:12) [evaluation 3554]
[SA] new: 3576. $less(0,X0) | alength = X0 | $less(X0,alength) | ~$less(0,alength) | alength = 0 (6:15) [evaluation 3553]
[SA] passive: 3557. i(l8(X2)) != 0 | sK2(0) = X2 | alength = 0 (6:12) [superposition 48,2654]
[SA] passive: 3559. v = a(l11(sK2(0)),0) | alength = 0 (6:10) [superposition 70,2654]
[SA] passive: 3572. i(l8(s(sK2(0)))) = 1 | alength = 0 (6:10) [evaluation 3560]
[SA] passive: 3571. zero = sK2(0) | alength = 0 (6:7) [trivial inequality removal 3561]
[SA] passive: 3565. v = a(l8(s(sK2(0))),0) | alength = 0 (6:11) [superposition 349,2654]
[SA] passive: 3568. 0 != sK0 | sK2(0) = sK2(sK0) | alength = 0 (6:11) [superposition 2839,2654]
[SA] active: 2919. $less(i(l8(sK2(X0))),0) | $less(X0,0) | ~$less(X0,alength) | i(l8(sK2(X0))) = i(l8(sK2(i(l8(sK2(X0)))))) (2:24:1) [resolution 2771,2480]
[SA] new: 3577. $less(X0,0) | ~$less(X0,alength) | i(l8(sK2(X0))) = i(l8(sK2(i(l8(sK2(X0)))))) | sK0 = X1 | 0 = sK0 | $less(X1,sK0) | $less(i(l8(sK2(X0))),X1) (4:33) [resolution 2919,324]
[SA] new: 3578. $less(X2,0) | ~$less(X2,alength) | i(l8(sK2(X2))) = i(l8(sK2(i(l8(sK2(X2)))))) | alength = 0 | $less(i(l8(sK2(X2))),alength) (5:27) [resolution 2919,262]
[SA] new: 3579. $less(X3,0) | ~$less(X3,alength) | i(l8(sK2(X3))) = i(l8(sK2(i(l8(sK2(X3)))))) | 0 = sK0 | $less(i(l8(sK2(X3))),alength) (4:27) [resolution 2919,116]
[SA] new: 3580. $less(X4,0) | ~$less(X4,alength) | i(l8(sK2(X4))) = i(l8(sK2(i(l8(sK2(X4)))))) | 0 = sK0 | $less(i(l8(sK2(X4))),sK0) (3:27) [resolution 2919,114]
[SA] new: 3581. $less(X5,0) | ~$less(X5,alength) | i(l8(sK2(X5))) = i(l8(sK2(i(l8(sK2(X5)))))) | ~$less(X6,i(l8(sK2(X5)))) | $less(X6,0) (3:27) [resolution 2919,14]
[SA] new: 3582. $less(X7,0) | ~$less(X7,alength) | i(l8(sK2(X7))) = i(l8(sK2(i(l8(sK2(X7)))))) | 0 = X8 | $less(X8,0) | $less(i(l8(sK2(X7))),X8) (3:30) [resolution 2919,99]
[SA] new: 3583. $less(X9,0) | ~$less(X9,alength) | i(l8(sK2(X9))) = i(l8(sK2(i(l8(sK2(X9)))))) | 0 = sK0 | $less(sK0,0) | $less(i(l8(sK2(X9))),alength) (4:30) [resolution 2919,205]
[SA] new: 3584. $less(X10,0) | ~$less(X10,alength) | i(l8(sK2(X10))) = i(l8(sK2(i(l8(sK2(X10)))))) | alength = 0 | $less(sK0,0) | $less(i(l8(sK2(X10))),alength) (4:30) [resolution 2919,254]
[SA] new: 3585. $less(X11,0) | ~$less(X11,alength) | i(l8(sK2(X11))) = i(l8(sK2(i(l8(sK2(X11)))))) | 0 = sK0 | 0 = sK0 | $less(0,0) | $less(i(l8(sK2(X11))),sK0) (4:33) [resolution 2919,333]
[SA] new: 3586. $less(X12,0) | ~$less(X12,alength) | i(l8(sK2(X12))) = i(l8(sK2(i(l8(sK2(X12)))))) | i(l8(sK2(0))) = 0 | alength = 0 | $less(alength,0) | $less(i(l8(sK2(X12))),0) (3:36) [resolution 2919,2950]
[SA] new: 3587. $less(i(l8(zero)),0) | $less(0,0) | ~$less(0,alength) | i(l8(sK2(0))) = i(l8(sK2(i(l8(sK2(0)))))) | 0 = sK0 (5:26) [superposition 2919,3207]
[SA] new: 3588. $less(sK0,0) | $less(sK0,0) | ~$less(sK0,alength) | i(l8(sK2(sK0))) = i(l8(sK2(i(l8(sK2(sK0)))))) | alength = sK0 (4:24) [superposition 2919,2495]
[SA] new: 3589. $less(sK0,0) | $less(sK0,0) | ~$less(sK0,alength) | i(l8(sK2(sK0))) = i(l8(sK2(i(l8(sK2(sK0)))))) (3:21) [superposition 2919,2833]
[SA] new: 3590. $less(0,0) | $less(0,0) | ~$less(0,alength) | i(l8(sK2(0))) = i(l8(sK2(i(l8(sK2(0)))))) | alength = 0 (6:24) [superposition 2919,2654]
[SA] new: 3591. $less(0,0) | $less(0,0) | ~$less(0,alength) | i(l8(sK2(0))) = i(l8(sK2(i(l8(sK2(0)))))) | 0 = sK0 (4:24) [superposition 2919,2831]
[SA] new: 3592. $less(0,0) | ~$less(0,alength) | i(l8(sK2(0))) = i(l8(sK2(i(l8(sK2(0)))))) | 0 = sK0 (4:21) [duplicate literal removal 3591]
[SA] new: 3593. ~$less(0,alength) | i(l8(sK2(0))) = i(l8(sK2(i(l8(sK2(0)))))) | 0 = sK0 (4:18) [evaluation 3592]
[SA] new: 3594. $less(0,0) | ~$less(0,alength) | i(l8(sK2(0))) = i(l8(sK2(i(l8(sK2(0)))))) | alength = 0 (6:21) [duplicate literal removal 3590]
[SA] new: 3595. ~$less(0,alength) | i(l8(sK2(0))) = i(l8(sK2(i(l8(sK2(0)))))) | alength = 0 (6:18) [evaluation 3594]
[SA] new: 3596. $less(sK0,0) | ~$less(sK0,alength) | i(l8(sK2(sK0))) = i(l8(sK2(i(l8(sK2(sK0)))))) (3:18) [duplicate literal removal 3589]
[SA] new: 3597. $less(sK0,0) | ~$less(sK0,alength) | i(l8(sK2(sK0))) = i(l8(sK2(i(l8(sK2(sK0)))))) | alength = sK0 (4:21) [duplicate literal removal 3588]
[SA] new: 3598. $less(i(l8(zero)),0) | ~$less(0,alength) | i(l8(sK2(0))) = i(l8(sK2(i(l8(sK2(0)))))) | 0 = sK0 (5:23) [evaluation 3587]
[SA] new: 3599. $less(X11,0) | ~$less(X11,alength) | i(l8(sK2(X11))) = i(l8(sK2(i(l8(sK2(X11)))))) | 0 = sK0 | $less(0,0) | $less(i(l8(sK2(X11))),sK0) (4:30) [duplicate literal removal 3585]
[SA] new: 3600. $less(X11,0) | ~$less(X11,alength) | i(l8(sK2(X11))) = i(l8(sK2(i(l8(sK2(X11)))))) | 0 = sK0 | $less(i(l8(sK2(X11))),sK0) (4:27) [evaluation 3599]
[SA] passive: 3577. $less(X0,0) | ~$less(X0,alength) | i(l8(sK2(X0))) = i(l8(sK2(i(l8(sK2(X0)))))) | sK0 = X1 | 0 = sK0 | $less(X1,sK0) | $less(i(l8(sK2(X0))),X1) (4:33) [resolution 2919,324]
[SA] passive: 3580. $less(X4,0) | ~$less(X4,alength) | i(l8(sK2(X4))) = i(l8(sK2(i(l8(sK2(X4)))))) | 0 = sK0 | $less(i(l8(sK2(X4))),sK0) (3:27) [resolution 2919,114]
[SA] passive: 3581. $less(X5,0) | ~$less(X5,alength) | i(l8(sK2(X5))) = i(l8(sK2(i(l8(sK2(X5)))))) | ~$less(X6,i(l8(sK2(X5)))) | $less(X6,0) (3:27) [resolution 2919,14]
[SA] passive: 3582. $less(X7,0) | ~$less(X7,alength) | i(l8(sK2(X7))) = i(l8(sK2(i(l8(sK2(X7)))))) | 0 = X8 | $less(X8,0) | $less(i(l8(sK2(X7))),X8) (3:30) [resolution 2919,99]
[SA] new: 3601. $less(i(l8(zero)),0) | i(l8(sK2(0))) = i(l8(sK2(i(l8(sK2(0)))))) | 0 = sK0 (5:20) [subsumption resolution 3598,113]
[SA] new: 3602. $less(0,0) | i(l8(sK2(0))) = i(l8(sK2(i(l8(sK2(0)))))) | 0 = sK0 (5:18) [forward demodulation 3601,65]
[SA] new: 3603. i(l8(sK2(0))) = i(l8(sK2(i(l8(sK2(0)))))) | 0 = sK0 (5:15) [evaluation 3602]
[SA] passive: 3603. i(l8(sK2(0))) = i(l8(sK2(i(l8(sK2(0)))))) | 0 = sK0 (5:15) [evaluation 3602]
[SA] new: 3604. ~$less(sK0,alength) | i(l8(sK2(sK0))) = i(l8(sK2(i(l8(sK2(sK0)))))) | alength = sK0 (4:18) [subsumption resolution 3597,46]
[SA] new: 3605. i(l8(sK2(sK0))) = i(l8(sK2(i(l8(sK2(sK0)))))) | alength = sK0 (4:15) [subsumption resolution 3604,45]
[SA] new: 3606. i(l8(sK2(sK0))) = sK0 | alength = sK0 (4:9) [forward demodulation 3605,2833]
[SA] new: 3607. ~$less(sK0,alength) | i(l8(sK2(sK0))) = i(l8(sK2(i(l8(sK2(sK0)))))) (3:15) [subsumption resolution 3596,46]
[SA] new: 3608. i(l8(sK2(sK0))) = i(l8(sK2(i(l8(sK2(sK0)))))) (3:12) [subsumption resolution 3607,45]
[SA] new: 3609. i(l8(sK2(sK0))) = sK0 (3:6) [forward demodulation 3608,2833]
[SA] new: 3610. i(l8(sK2(0))) = i(l8(sK2(i(l8(sK2(0)))))) | alength = 0 (6:15) [subsumption resolution 3595,244]
[SA] passive: 3610. i(l8(sK2(0))) = i(l8(sK2(i(l8(sK2(0)))))) | alength = 0 (6:15) [subsumption resolution 3595,244]
[SA] active: 3571. zero = sK2(0) | alength = 0 (6:7:1) [trivial inequality removal 3561]
[SA] new: 3611. i(l8(zero)) = 0 | alength = 0 | alength = 0 (7:11) [superposition 2654,3571]
[SA] new: 3612. i(l8(zero)) = 0 | 0 = sK0 | alength = 0 (7:11) [superposition 2831,3571]
[SA] new: 3613. 'Sub'(zero,nl8) | $less(0,0) | ~$less(0,alength) | alength = 0 (7:12) [superposition 2479,3571]
[SA] new: 3614. 'Sub'(zero,nl8) | $less(alength,0) | alength = 0 | $less(0,0) | alength = 0 (7:15) [superposition 2655,3571]
[SA] new: 3615. $less(i(l8(zero)),alength) | $less(alength,0) | alength = 0 | $less(0,0) | alength = 0 (7:17) [superposition 2657,3571]
[SA] new: 3616. $less(i(l8(zero)),alength) | ~$less(0,alength) | $less(0,0) | alength = 0 (7:14) [superposition 2771,3571]
[SA] new: 3617. $less(i(l8(zero)),0) | $less(0,0) | ~$less(0,alength) | i(l8(sK2(0))) = i(l8(sK2(i(l8(sK2(0)))))) | alength = 0 (7:26) [superposition 2919,3571]
[SA] new: 3618. ~$less(X0,i(l8(zero))) | $less(0,0) | ~$less(0,alength) | $less(X0,alength) | alength = 0 (7:17) [superposition 2920,3571]
[SA] new: 3619. $less(i(l8(zero)),X1) | $less(0,0) | alength = X1 | $less(X1,alength) | ~$less(0,alength) | alength = 0 (7:20) [superposition 2921,3571]
[SA] new: 3620. $less(i(l8(zero)),X1) | alength = X1 | $less(X1,alength) | ~$less(0,alength) | alength = 0 (7:17) [evaluation 3619]
[SA] new: 3621. ~$less(X0,i(l8(zero))) | ~$less(0,alength) | $less(X0,alength) | alength = 0 (7:14) [evaluation 3618]
[SA] new: 3622. $less(i(l8(zero)),0) | ~$less(0,alength) | i(l8(sK2(0))) = i(l8(sK2(i(l8(sK2(0)))))) | alength = 0 (7:23) [evaluation 3617]
[SA] new: 3623. $less(i(l8(zero)),alength) | ~$less(0,alength) | alength = 0 (7:11) [evaluation 3616]
[SA] new: 3624. $less(i(l8(zero)),alength) | $less(alength,0) | alength = 0 | $less(0,0) (7:14) [duplicate literal removal 3615]
[SA] new: 3625. $less(i(l8(zero)),alength) | $less(alength,0) | alength = 0 (7:11) [evaluation 3624]
[SA] new: 3626. 'Sub'(zero,nl8) | $less(alength,0) | alength = 0 | $less(0,0) (7:12) [duplicate literal removal 3614]
[SA] new: 3627. 'Sub'(zero,nl8) | $less(alength,0) | alength = 0 (7:9) [evaluation 3626]
[SA] new: 3628. 'Sub'(zero,nl8) | ~$less(0,alength) | alength = 0 (7:9) [evaluation 3613]
[SA] new: 3629. i(l8(zero)) = 0 | alength = 0 (7:8) [duplicate literal removal 3611]
[SA] new: 3630. 'Sub'(zero,nl8) | alength = 0 (7:6) [subsumption resolution 3628,244]
[SA] passive: 3630. 'Sub'(zero,nl8) | alength = 0 (7:6) [subsumption resolution 3628,244]
[SA] new: 3631. $less(0,alength) | $less(alength,0) | alength = 0 (7:9) [forward demodulation 3625,65]
[SA] new: 3632. $less(i(l8(zero)),alength) | alength = 0 (7:8) [subsumption resolution 3623,244]
[SA] new: 3633. $less(0,alength) | alength = 0 (7:6) [forward demodulation 3632,65]
[SA] new: 3634. ~$less(X0,i(l8(zero))) | $less(X0,alength) | alength = 0 (7:11) [subsumption resolution 3621,244]
[SA] new: 3635. ~$less(X0,0) | $less(X0,alength) | alength = 0 (7:9) [forward demodulation 3634,65]
[SA] new: 3636. $less(i(l8(zero)),X1) | alength = X1 | $less(X1,alength) | alength = 0 (7:14) [subsumption resolution 3620,244]
[SA] new: 3637. $less(0,X1) | alength = X1 | $less(X1,alength) | alength = 0 (7:12) [forward demodulation 3636,65]
[SA] active: 2664. $less(alength,X1) | $less(X1,0) | i(l8(sK1(X0,sK2(X1)))) = X0 | alength = X1 | a(main_end,X0) = a(l11(sK2(X1)),X0) (2:26:2) [forward demodulation 2501,2476]
[SA] new: 3638. $less(sK0,0) | i(l8(sK1(X0,sK2(sK0)))) = X0 | alength = sK0 | a(main_end,X0) = a(l11(sK2(sK0)),X0) | alength = sK0 (6:26) [resolution 2664,2704]
[SA] new: 3639. $less(alength,0) | i(l8(sK1(X1,sK2(alength)))) = X1 | alength = alength | a(main_end,X1) = a(l11(sK2(alength)),X1) (3:23) [resolution 2664,13]
[SA] new: 3640. $less(X2,0) | i(l8(sK1(X3,sK2(X2)))) = X3 | alength = X2 | a(main_end,X3) = a(l11(sK2(X2)),X3) | ~$less(X4,alength) | $less(X4,X2) (3:29) [resolution 2664,14]
[SA] new: 3641. $less(X5,0) | i(l8(sK1(X6,sK2(X5)))) = X6 | alength = X5 | a(main_end,X6) = a(l11(sK2(X5)),X6) | X5 = X7 | $less(X7,X5) | $less(alength,X7) (3:32) [resolution 2664,99]
[SA] new: 3642. $less(X8,0) | i(l8(sK1(X9,sK2(X8)))) = X9 | alength = X8 | a(main_end,X9) = a(l11(sK2(X8)),X9) | sK0 = X8 | $less(sK0,X8) | $less(alength,alength) (4:32) [resolution 2664,205]
[SA] new: 3643. $less(X10,0) | i(l8(sK1(X11,sK2(X10)))) = X11 | alength = X10 | a(main_end,X11) = a(l11(sK2(X10)),X11) | alength = X10 | $less(sK0,X10) | $less(alength,alength) (4:32) [resolution 2664,254]
[SA] new: 3644. $less(X12,0) | i(l8(sK1(X13,sK2(X12)))) = X13 | alength = X12 | a(main_end,X13) = a(l11(sK2(X12)),X13) | sK0 = X12 | 0 = sK0 | $less(0,X12) | $less(alength,sK0) (4:35) [resolution 2664,333]
[SA] new: 3645. $less(X14,0) | i(l8(sK1(X15,sK2(X14)))) = X15 | alength = X14 | a(main_end,X15) = a(l11(sK2(X14)),X15) | i(l8(sK2(X14))) = X14 | alength = X14 | $less(alength,X14) | $less(alength,0) (3:38) [resolution 2664,2950]
[SA] new: 3646. $less(alength,0) | i(l8(sK1(X16,sK2(alength)))) = X16 | alength = alength | a(main_end,X16) = a(l11(sK2(alength)),X16) | $less(alength,0) | a(main_end,X17) = a(l11(sK2(alength)),X17) | i(l8(sK1(X17,sK2(alength)))) = X17 (3:43) [resolution 2664,2770]
[SA] new: 3647. $less(alength,0) | i(l8(sK1(X18,sK2(alength)))) = X18 | alength = alength | a(main_end,X18) = a(l11(sK2(alength)),X18) | alength = X19 | $less(X19,0) | i(l8(sK2(X19))) = X19 | $less(alength,X19) (3:38) [resolution 2664,2660]
[SA] new: 3648. $less(alength,0) | i(l8(sK1(X20,sK2(alength)))) = X20 | alength = alength | a(main_end,X20) = a(l11(sK2(alength)),X20) | $less(alength,0) | alength = i(l8(sK2(alength))) (3:32) [resolution 2664,2480]
[SA] new: 3649. $less(i(l8(p(nl8))),0) | i(l8(sK1(X21,sK2(i(l8(p(nl8))))))) = X21 | alength = i(l8(p(nl8))) | a(main_end,X21) = a(l11(sK2(i(l8(p(nl8))))),X21) | zero = nl8 | $less(alength,alength) (4:41) [resolution 2664,194]
[SA] new: 3650. $less(i(l8(sK2(X22))),0) | i(l8(sK1(X23,sK2(i(l8(sK2(X22))))))) = X23 | alength = i(l8(sK2(X22))) | a(main_end,X23) = a(l11(sK2(i(l8(sK2(X22))))),X23) | $less(X22,0) | ~$less(X22,alength) | $less(alength,alength) (3:44) [resolution 2664,2920]
[SA] new: 3651. $less(0,0) | i(l8(sK1(X24,sK2(0)))) = X24 | alength = 0 | a(main_end,X24) = a(l11(sK2(0)),X24) | sK0 = X25 | 0 = sK0 | $less(X25,sK0) | $less(alength,X25) (4:35) [resolution 2664,324]
[SA] new: 3652. $less(0,0) | i(l8(sK1(X26,sK2(0)))) = X26 | alength = 0 | a(main_end,X26) = a(l11(sK2(0)),X26) | alength = 0 | $less(alength,alength) (5:29) [resolution 2664,262]
[SA] new: 3653. $less(0,0) | i(l8(sK1(X27,sK2(0)))) = X27 | alength = 0 | a(main_end,X27) = a(l11(sK2(0)),X27) | 0 = sK0 | $less(alength,alength) (4:29) [resolution 2664,116]
[SA] new: 3654. $less(0,0) | i(l8(sK1(X28,sK2(0)))) = X28 | alength = 0 | a(main_end,X28) = a(l11(sK2(0)),X28) | 0 = sK0 | $less(alength,sK0) (3:29) [resolution 2664,114]
[SA] new: 3655. $less(sK0,0) | i(l8(sK1(X29,sK2(sK0)))) = X29 | alength = sK0 | a(main_end,X29) = a(l11(sK2(sK0)),X29) | alength = X30 | $less(X30,alength) | $less(alength,X30) (4:32) [resolution 2664,247]
[SA] new: 3656. $less(sK0,0) | i(l8(sK1(X31,sK2(sK0)))) = X31 | alength = sK0 | a(main_end,X31) = a(l11(sK2(sK0)),X31) | sK0 = X32 | $less(X32,alength) | $less(alength,X32) (4:32) [resolution 2664,199]
[SA] new: 3657. $less(sK0,0) | i(l8(sK1(X33,sK2(sK0)))) = X33 | alength = sK0 | a(main_end,X33) = a(l11(sK2(sK0)),X33) | $less(alength,alength) (3:26) [resolution 2664,96]
[SA] new: 3658. $less(alength,sK0) | i(l8(sK1(X0,sK2(sK0)))) = X0 | alength = sK0 | a(main_end,X0) = a(l11(sK2(sK0)),X0) (3:23) [resolution 2664,46]
[SA] new: 3659. $less(alength,X1) | i(l8(sK1(X2,sK2(X1)))) = X2 | alength = X1 | a(main_end,X2) = a(l11(sK2(X1)),X2) | sK0 = X3 | 0 = sK0 | $less(X3,sK0) | $less(X1,X3) (4:35) [resolution 2664,324]
[SA] new: 3660. $less(alength,X4) | i(l8(sK1(X5,sK2(X4)))) = X5 | alength = X4 | a(main_end,X5) = a(l11(sK2(X4)),X5) | alength = 0 | $less(X4,alength) (5:29) [resolution 2664,262]
[SA] new: 3661. $less(alength,X6) | i(l8(sK1(X7,sK2(X6)))) = X7 | alength = X6 | a(main_end,X7) = a(l11(sK2(X6)),X7) | 0 = sK0 | $less(X6,alength) (4:29) [resolution 2664,116]
[SA] new: 3662. $less(alength,X8) | i(l8(sK1(X9,sK2(X8)))) = X9 | alength = X8 | a(main_end,X9) = a(l11(sK2(X8)),X9) | 0 = sK0 | $less(X8,sK0) (3:29) [resolution 2664,114]
[SA] new: 3663. $less(alength,0) | i(l8(sK1(X10,sK2(0)))) = X10 | alength = 0 | a(main_end,X10) = a(l11(sK2(0)),X10) (3:23) [resolution 2664,13]
[SA] new: 3664. $less(alength,X11) | i(l8(sK1(X12,sK2(X11)))) = X12 | alength = X11 | a(main_end,X12) = a(l11(sK2(X11)),X12) | ~$less(X13,X11) | $less(X13,0) (3:29) [resolution 2664,14]
[SA] new: 3665. $less(alength,X14) | i(l8(sK1(X15,sK2(X14)))) = X15 | alength = X14 | a(main_end,X15) = a(l11(sK2(X14)),X15) | 0 = X16 | $less(X16,0) | $less(X14,X16) (3:32) [resolution 2664,99]
[SA] new: 3666. $less(alength,X17) | i(l8(sK1(X18,sK2(X17)))) = X18 | alength = X17 | a(main_end,X18) = a(l11(sK2(X17)),X18) | 0 = sK0 | $less(sK0,0) | $less(X17,alength) (4:32) [resolution 2664,205]
[SA] new: 3667. $less(alength,X19) | i(l8(sK1(X20,sK2(X19)))) = X20 | alength = X19 | a(main_end,X20) = a(l11(sK2(X19)),X20) | alength = 0 | $less(sK0,0) | $less(X19,alength) (4:32) [resolution 2664,254]
[SA] new: 3668. $less(alength,X21) | i(l8(sK1(X22,sK2(X21)))) = X22 | alength = X21 | a(main_end,X22) = a(l11(sK2(X21)),X22) | 0 = sK0 | 0 = sK0 | $less(0,0) | $less(X21,sK0) (4:35) [resolution 2664,333]
[SA] new: 3669. $less(alength,X23) | i(l8(sK1(X24,sK2(X23)))) = X24 | alength = X23 | a(main_end,X24) = a(l11(sK2(X23)),X24) | i(l8(sK2(0))) = 0 | alength = 0 | $less(alength,0) | $less(X23,0) (3:38) [resolution 2664,2950]
[SA] new: 3670. $less(alength,X21) | i(l8(sK1(X22,sK2(X21)))) = X22 | alength = X21 | a(main_end,X22) = a(l11(sK2(X21)),X22) | 0 = sK0 | $less(0,0) | $less(X21,sK0) (4:32) [duplicate literal removal 3668]
[SA] new: 3671. $less(alength,X21) | i(l8(sK1(X22,sK2(X21)))) = X22 | alength = X21 | a(main_end,X22) = a(l11(sK2(X21)),X22) | 0 = sK0 | $less(X21,sK0) (4:29) [evaluation 3670]
[SA] new: 3672. i(l8(sK1(X28,sK2(0)))) = X28 | alength = 0 | a(main_end,X28) = a(l11(sK2(0)),X28) | 0 = sK0 | $less(alength,sK0) (3:26) [evaluation 3654]
[SA] new: 3673. i(l8(sK1(X27,sK2(0)))) = X27 | alength = 0 | a(main_end,X27) = a(l11(sK2(0)),X27) | 0 = sK0 | $less(alength,alength) (4:26) [evaluation 3653]
[SA] new: 3674. $less(0,0) | i(l8(sK1(X26,sK2(0)))) = X26 | alength = 0 | a(main_end,X26) = a(l11(sK2(0)),X26) | $less(alength,alength) (5:26) [duplicate literal removal 3652]
[SA] new: 3675. i(l8(sK1(X26,sK2(0)))) = X26 | alength = 0 | a(main_end,X26) = a(l11(sK2(0)),X26) | $less(alength,alength) (5:23) [evaluation 3674]
[SA] new: 3676. i(l8(sK1(X24,sK2(0)))) = X24 | alength = 0 | a(main_end,X24) = a(l11(sK2(0)),X24) | sK0 = X25 | 0 = sK0 | $less(X25,sK0) | $less(alength,X25) (4:32) [evaluation 3651]
[SA] new: 3677. $less(alength,0) | i(l8(sK1(X20,sK2(alength)))) = X20 | alength = alength | a(main_end,X20) = a(l11(sK2(alength)),X20) | alength = i(l8(sK2(alength))) (3:29) [duplicate literal removal 3648]
[SA] new: 3678. $less(alength,0) | i(l8(sK1(X16,sK2(alength)))) = X16 | alength = alength | a(main_end,X16) = a(l11(sK2(alength)),X16) | a(main_end,X17) = a(l11(sK2(alength)),X17) | i(l8(sK1(X17,sK2(alength)))) = X17 (3:40) [duplicate literal removal 3646]
[SA] new: 3679. $less(X14,0) | i(l8(sK1(X15,sK2(X14)))) = X15 | alength = X14 | a(main_end,X15) = a(l11(sK2(X14)),X15) | i(l8(sK2(X14))) = X14 | $less(alength,X14) | $less(alength,0) (3:35) [duplicate literal removal 3645]
[SA] new: 3680. $less(X10,0) | i(l8(sK1(X11,sK2(X10)))) = X11 | alength = X10 | a(main_end,X11) = a(l11(sK2(X10)),X11) | $less(sK0,X10) | $less(alength,alength) (4:29) [duplicate literal removal 3643]
[SA] new: 3681. $less(sK0,0) | i(l8(sK1(X0,sK2(sK0)))) = X0 | alength = sK0 | a(main_end,X0) = a(l11(sK2(sK0)),X0) (6:23) [duplicate literal removal 3638]
[SA] passive: 3659. $less(alength,X1) | i(l8(sK1(X2,sK2(X1)))) = X2 | alength = X1 | a(main_end,X2) = a(l11(sK2(X1)),X2) | sK0 = X3 | 0 = sK0 | $less(X3,sK0) | $less(X1,X3) (4:35) [resolution 2664,324]
[SA] passive: 3662. $less(alength,X8) | i(l8(sK1(X9,sK2(X8)))) = X9 | alength = X8 | a(main_end,X9) = a(l11(sK2(X8)),X9) | 0 = sK0 | $less(X8,sK0) (3:29) [resolution 2664,114]
[SA] passive: 3664. $less(alength,X11) | i(l8(sK1(X12,sK2(X11)))) = X12 | alength = X11 | a(main_end,X12) = a(l11(sK2(X11)),X12) | ~$less(X13,X11) | $less(X13,0) (3:29) [resolution 2664,14]
[SA] passive: 3665. $less(alength,X14) | i(l8(sK1(X15,sK2(X14)))) = X15 | alength = X14 | a(main_end,X15) = a(l11(sK2(X14)),X15) | 0 = X16 | $less(X16,0) | $less(X14,X16) (3:32) [resolution 2664,99]
[SA] active: 3630. 'Sub'(zero,nl8) | alength = 0 (7:6:1) [subsumption resolution 3628,244]
[SA] new: 3682. alength = 0 | a(main_end,X0) = a(l11(zero),X0) | i(l8(sK1(X0,zero))) = X0 (8:18) [resolution 3630,150]
[SA] new: 3683. alength = 0 | $less(i(l8(zero)),alength) (8:8) [resolution 3630,56]
[SA] passive: 3682. alength = 0 | a(main_end,X0) = a(l11(zero),X0) | i(l8(sK1(X0,zero))) = X0 (8:18) [resolution 3630,150]
[SA] new: 3684. $less(0,alength) | alength = 0 (8:6) [forward demodulation 3683,65]
[SA] active: 2981. $sum(alength,1) != sK0 | s(nl8) = sK2(sK0) (3:10:1) [superposition 2839,2484]
[SA] active: 2987. v = a(l11(s(sK2(sK0))),$sum(sK0,1)) (3:10:1) [superposition 70,2842]
[SA] new: 3685. v = a(l8(s(sK2(sK0))),$sum(sK0,1)) | i(l8(s(sK2(sK0)))) = $sum(sK0,1) (4:19) [superposition 2987,72]
[SA] new: 3686. v = a(l8(s(sK2(sK0))),$sum(sK0,1)) | i(l8(s(sK2(sK0)))) = $sum(sK0,1) (4:19) [superposition 72,2987]
[SA] active: 2994. alength != $sum(sK0,1) | s(sK2(sK0)) = nl8 (3:10:1) [superposition 2483,2842]
[SA] active: 2510. v = a(l8(s(s(nl8))),$sum(alength,1)) (4:10:1) [backward demodulation 2476,359]
[SA] new: 3687. v = a(l8(s(nl8)),$sum(alength,1)) | i(l8(s(nl8))) = $sum(alength,1) (5:17) [superposition 2510,440]
[SA] new: 3688. v = a(l11(s(nl8)),$sum(alength,1)) (5:9) [superposition 2510,57]
[SA] new: 3689. v = a(l8(s(nl8)),$sum(alength,1)) | i(l8(s(nl8))) = $sum(alength,1) (5:17) [superposition 440,2510]
[SA] new: 3690. v = a(l11(s(nl8)),$sum(alength,1)) (5:9) [superposition 57,2510]
[SA] active: 2995. $sum(sK0,1) != sK0 | s(sK2(sK0)) = sK2(sK0) (3:11:1) [superposition 2839,2842]
[SA] active: 2526. v = a(l11(sK2(sK0)),sK0) | alength = sK0 (4:10:1) [backward demodulation 2476,468]
[SA] new: 3691. v = a(l8(sK2(sK0)),sK0) | alength = sK0 | i(l8(sK2(sK0))) = sK0 (5:16) [superposition 2526,72]
[SA] new: 3692. v = a(l8(sK2(sK0)),sK0) | i(l8(sK2(sK0))) = sK0 | alength = sK0 (5:16) [superposition 72,2526]
[SA] active: 3051. $sum(alength,1) != $sum(sK0,1) | nl8 = sK2(sK0) (3:11:1) [term algebras injectivity 3047]
[SA] active: 2528. 0 != sK0 | alength = sK0 | zero = sK2(sK0) (4:10:1) [backward demodulation 2476,470]
[SA] active: 2985. i(l8(X0)) != $sum(sK0,1) | s(sK2(sK0)) = X0 (3:12:1) [superposition 48,2842]
[SA] new: 3693. $sum(sK0,1) != 0 | zero = s(sK2(sK0)) (4:10) [superposition 2985,65]
[SA] new: 3694. 1 != $sum(sK0,1) | s(zero) = s(sK2(sK0)) (4:11) [superposition 2985,119]
[SA] new: 3695. $sum(sK0,1) != 2 | s(s(zero)) = s(sK2(sK0)) (4:12) [superposition 2985,124]
[SA] new: 3696. $sum(sK0,1) != 3 | s(sK2(sK0)) = s(s(s(zero))) (4:13) [superposition 2985,133]
[SA] new: 3697. $sum(sK0,1) != 4 | s(sK2(sK0)) = s(s(s(s(zero)))) (5:14) [superposition 2985,186]
[SA] new: 3698. $sum(sK0,1) != 5 | s(sK2(sK0)) = s(s(s(s(s(zero))))) (6:15) [superposition 2985,312]
[SA] new: 3699. $sum(sK0,1) != 6 | s(sK2(sK0)) = s(s(s(s(s(s(zero)))))) (7:16) [superposition 2985,640]
[SA] new: 3700. $sum(sK0,1) != $sum($sum(alength,1),1) | s(s(nl8)) = s(sK2(sK0)) (4:16) [superposition 2985,2489]
[SA] new: 3701. $sum(alength,1) != $sum(sK0,1) | s(nl8) = s(sK2(sK0)) (4:13) [superposition 2985,2484]
[SA] new: 3703. alength != $sum(sK0,1) | s(sK2(sK0)) = nl8 (4:10) [superposition 2985,2477]
[SA] new: 3704. $sum(sK0,1) != X0 | s(sK2(sK0)) = sK1(X0,p(nl8)) | a(main_end,X0) = a(l11(p(nl8)),X0) | zero = nl8 (4:25) [superposition 2985,226]
[SA] new: 3705. $sum(sK0,1) != sK0 | s(sK2(sK0)) = sK2(sK0) | alength = sK0 (4:14) [superposition 2985,2495]
[SA] new: 3706. $sum(sK0,1) != sK0 | s(sK2(sK0)) = sK2(sK0) (4:11) [superposition 2985,2833]
[SA] new: 3707. $sum(sK0,1) != 0 | s(sK2(sK0)) = sK2(0) | alength = 0 (6:14) [superposition 2985,2654]
[SA] new: 3708. $sum(sK0,1) != 0 | s(sK2(sK0)) = sK2(0) | 0 = sK0 (4:14) [superposition 2985,2831]
[SA] new: 3709. $uminus(1) != sK0 | s(sK2(sK0)) = sK2(0) | 0 = sK0 (4:13) [evaluation 3708]
[SA] new: 3710. sK0 != -1 | s(sK2(sK0)) = sK2(0) | 0 = sK0 (4:12) [evaluation 3709]
[SA] new: 3711. $uminus(1) != sK0 | s(sK2(sK0)) = sK2(0) | alength = 0 (6:13) [evaluation 3707]
[SA] new: 3712. sK0 != -1 | s(sK2(sK0)) = sK2(0) | alength = 0 (6:12) [evaluation 3711]
[SA] new: 3713. $sum(alength,1) != $sum(sK0,1) | nl8 = sK2(sK0) (4:11) [term algebras injectivity 3701]
[SA] new: 3714. $sum(sK0,1) != $sum($sum(alength,1),1) | s(nl8) = sK2(sK0) (4:14) [term algebras injectivity 3700]
[SA] new: 3715. $sum(6,-1) != sK0 | s(sK2(sK0)) = s(s(s(s(s(s(zero)))))) (7:16) [evaluation 3699]
[SA] new: 3716. sK0 != 5 | s(sK2(sK0)) = s(s(s(s(s(s(zero)))))) (7:14) [evaluation 3715]
[SA] new: 3717. sK0 != 5 | s(s(s(s(s(zero))))) = sK2(sK0) (7:12) [term algebras injectivity 3716]
[SA] new: 3718. $sum(5,-1) != sK0 | s(sK2(sK0)) = s(s(s(s(s(zero))))) (6:15) [evaluation 3698]
[SA] new: 3719. sK0 != 4 | s(sK2(sK0)) = s(s(s(s(s(zero))))) (6:13) [evaluation 3718]
[SA] new: 3720. sK0 != 4 | s(s(s(s(zero)))) = sK2(sK0) (6:11) [term algebras injectivity 3719]
[SA] new: 3721. $sum(4,-1) != sK0 | s(sK2(sK0)) = s(s(s(s(zero)))) (5:14) [evaluation 3697]
[SA] new: 3722. sK0 != 3 | s(sK2(sK0)) = s(s(s(s(zero)))) (5:12) [evaluation 3721]
[SA] new: 3723. sK0 != 3 | s(s(s(zero))) = sK2(sK0) (5:10) [term algebras injectivity 3722]
[SA] new: 3724. $sum(3,-1) != sK0 | s(sK2(sK0)) = s(s(s(zero))) (4:13) [evaluation 3696]
[SA] new: 3725. sK0 != 2 | s(sK2(sK0)) = s(s(s(zero))) (4:11) [evaluation 3724]
[SA] new: 3726. sK0 != 2 | s(s(zero)) = sK2(sK0) (4:9) [term algebras injectivity 3725]
[SA] new: 3727. $sum(2,-1) != sK0 | s(s(zero)) = s(sK2(sK0)) (4:12) [evaluation 3695]
[SA] new: 3728. 1 != sK0 | s(s(zero)) = s(sK2(sK0)) (4:10) [evaluation 3727]
[SA] new: 3729. 1 != sK0 | s(zero) = sK2(sK0) (4:8) [term algebras injectivity 3728]
[SA] new: 3730. $sum(1,-1) != sK0 | s(zero) = s(sK2(sK0)) (4:11) [evaluation 3694]
[SA] new: 3731. 0 != sK0 | s(zero) = s(sK2(sK0)) (4:9) [evaluation 3730]
[SA] new: 3732. 0 != sK0 | zero = sK2(sK0) (4:7) [term algebras injectivity 3731]
[SA] new: 3733. $uminus(1) != sK0 | zero = s(sK2(sK0)) (4:9) [evaluation 3693]
[SA] new: 3734. sK0 != -1 | zero = s(sK2(sK0)) (4:8) [evaluation 3733]
[SA] new: 3735. sK0 != -1 (4:3) [term algebras distinctness 3734]
[SA] passive: 3714. $sum(sK0,1) != $sum($sum(alength,1),1) | s(nl8) = sK2(sK0) (4:14) [term algebras injectivity 3700]
[SA] passive: 3704. $sum(sK0,1) != X0 | s(sK2(sK0)) = sK1(X0,p(nl8)) | a(main_end,X0) = a(l11(p(nl8)),X0) | zero = nl8 (4:25) [superposition 2985,226]
[SA] active: 2579. 'Sub'(sK2(0),nl8) | alength = sK0 | 0 = sK0 (4:10:1) [backward demodulation 2476,848]
[SA] new: 3736. alength = sK0 | 0 = sK0 | a(main_end,X0) = a(l11(sK2(0)),X0) | i(l8(sK1(X0,sK2(0)))) = X0 (5:23) [resolution 2579,150]
[SA] new: 3737. alength = sK0 | 0 = sK0 | $less(i(l8(sK2(0))),alength) (5:12) [resolution 2579,56]
[SA] new: 3738. 'Sub'(zero,nl8) | alength = sK0 | 0 = sK0 | alength = 0 (7:12) [superposition 2579,3571]
[SA] new: 3739. 'Sub'(zero,nl8) | alength = sK0 | 0 = sK0 | 0 = sK0 (5:12) [superposition 2579,3207]
[SA] new: 3740. 'Sub'(zero,nl8) | alength = sK0 | 0 = sK0 (5:9) [duplicate literal removal 3739]
[SA] active: 2988. i(l8(s(s(sK2(sK0))))) = $sum($sum(sK0,1),1) (3:12:1) [superposition 74,2842]
[SA] new: 3741. i(l8(X0)) != $sum($sum(sK0,1),1) | s(s(sK2(sK0))) = X0 (4:15) [superposition 48,2988]
[SA] new: 3742. i(l8(X1)) != $sum($sum(sK0,1),1) | s(s(sK2(sK0))) = X1 (4:15) [superposition 48,2988]
[SA] new: 3743. v = a(l11(s(s(sK2(sK0)))),$sum($sum(sK0,1),1)) (4:13) [superposition 70,2988]
[SA] new: 3744. i(l8(s(s(s(sK2(sK0)))))) = $sum($sum($sum(sK0,1),1),1) (4:15) [superposition 74,2988]
[SA] new: 3745. $sum($sum(sK0,1),1) != 0 | zero = s(s(sK2(sK0))) (4:13) [superposition 108,2988]
[SA] new: 3746. 1 != $sum($sum(sK0,1),1) | s(zero) = s(s(sK2(sK0))) (4:14) [superposition 121,2988]
[SA] new: 3747. $sum($sum(sK0,1),1) != 2 | s(s(zero)) = s(s(sK2(sK0))) (4:15) [superposition 130,2988]
[SA] new: 3748. $sum($sum(sK0,1),1) != 3 | s(s(s(zero))) = s(s(sK2(sK0))) (5:16) [superposition 183,2988]
[SA] new: 3749. v = a(l8(s(s(s(sK2(sK0))))),$sum($sum(sK0,1),1)) (4:14) [superposition 349,2988]
[SA] new: 3750. alength != $sum($sum(sK0,1),1) | s(s(sK2(sK0))) = nl8 (4:13) [superposition 2483,2988]
[SA] new: 3751. $sum(alength,1) != $sum($sum(sK0,1),1) | s(nl8) = s(s(sK2(sK0))) (4:16) [superposition 2490,2988]
[SA] new: 3752. $sum($sum(sK0,1),1) != sK0 | s(s(sK2(sK0))) = sK2(sK0) (4:14) [superposition 2839,2988]
[SA] new: 3753. $sum(sK0,1) != $sum($sum(sK0,1),1) | s(sK2(sK0)) = s(s(sK2(sK0))) (4:17) [superposition 2985,2988]
[SA] new: 3754. $sum(sK0,1) != $sum($sum(sK0,1),1) | s(sK2(sK0)) = sK2(sK0) (4:15) [term algebras injectivity 3753]
[SA] new: 3755. $sum(alength,1) != $sum($sum(sK0,1),1) | s(sK2(sK0)) = nl8 (4:14) [term algebras injectivity 3751]
[SA] new: 3756. $sum($sum(3,-1),-1) != sK0 | s(s(s(zero))) = s(s(sK2(sK0))) (5:16) [evaluation 3748]
[SA] new: 3757. $sum(2,-1) != sK0 | s(s(s(zero))) = s(s(sK2(sK0))) (5:14) [evaluation 3756]
[SA] new: 3758. 1 != sK0 | s(s(s(zero))) = s(s(sK2(sK0))) (5:12) [evaluation 3757]
[SA] new: 3759. 1 != sK0 | s(s(zero)) = s(sK2(sK0)) (5:10) [term algebras injectivity 3758]
[SA] new: 3760. 1 != sK0 | s(zero) = sK2(sK0) (5:8) [term algebras injectivity 3759]
[SA] new: 3761. $sum($sum(2,-1),-1) != sK0 | s(s(zero)) = s(s(sK2(sK0))) (4:15) [evaluation 3747]
[SA] new: 3762. $sum(1,-1) != sK0 | s(s(zero)) = s(s(sK2(sK0))) (4:13) [evaluation 3761]
[SA] new: 3763. 0 != sK0 | s(s(zero)) = s(s(sK2(sK0))) (4:11) [evaluation 3762]
[SA] new: 3764. 0 != sK0 | s(zero) = s(sK2(sK0)) (4:9) [term algebras injectivity 3763]
[SA] new: 3765. 0 != sK0 | zero = sK2(sK0) (4:7) [term algebras injectivity 3764]
[SA] new: 3766. $sum($sum(1,-1),-1) != sK0 | s(zero) = s(s(sK2(sK0))) (4:14) [evaluation 3746]
[SA] new: 3767. $sum(0,-1) != sK0 | s(zero) = s(s(sK2(sK0))) (4:12) [evaluation 3766]
[SA] new: 3768. sK0 != -1 | s(zero) = s(s(sK2(sK0))) (4:10) [evaluation 3767]
[SA] new: 3769. sK0 != -1 | zero = s(sK2(sK0)) (4:8) [term algebras injectivity 3768]
[SA] new: 3770. sK0 != -1 (4:3) [term algebras distinctness 3769]
[SA] new: 3771. $sum($uminus(1),-1) != sK0 | zero = s(s(sK2(sK0))) (4:12) [evaluation 3745]
[SA] new: 3772. $sum(-1,-1) != sK0 | zero = s(s(sK2(sK0))) (4:11) [evaluation 3771]
[SA] new: 3773. sK0 != -2 | zero = s(s(sK2(sK0))) (4:9) [evaluation 3772]
[SA] new: 3774. sK0 != -2 (4:3) [term algebras distinctness 3773]
[SA] passive: 3741. i(l8(X0)) != $sum($sum(sK0,1),1) | s(s(sK2(sK0))) = X0 (4:15) [superposition 48,2988]
[SA] passive: 3743. v = a(l11(s(s(sK2(sK0)))),$sum($sum(sK0,1),1)) (4:13) [superposition 70,2988]
[SA] passive: 3744. i(l8(s(s(s(sK2(sK0)))))) = $sum($sum($sum(sK0,1),1),1) (4:15) [superposition 74,2988]
[SA] passive: 3774. sK0 != -2 (4:3) [term algebras distinctness 3773]
[SA] passive: 3749. v = a(l8(s(s(s(sK2(sK0))))),$sum($sum(sK0,1),1)) (4:14) [superposition 349,2988]
[SA] passive: 3750. alength != $sum($sum(sK0,1),1) | s(s(sK2(sK0))) = nl8 (4:13) [superposition 2483,2988]
[SA] passive: 3755. $sum(alength,1) != $sum($sum(sK0,1),1) | s(sK2(sK0)) = nl8 (4:14) [term algebras injectivity 3751]
[SA] passive: 3752. $sum($sum(sK0,1),1) != sK0 | s(s(sK2(sK0))) = sK2(sK0) (4:14) [superposition 2839,2988]
[SA] passive: 3754. $sum(sK0,1) != $sum($sum(sK0,1),1) | s(sK2(sK0)) = sK2(sK0) (4:15) [term algebras injectivity 3753]
[SA] active: 3774. sK0 != -2 (4:3:1) [term algebras distinctness 3773]
[SA] active: 3010. v = a(l11(s(s(nl8))),$sum($sum(alength,1),1)) (3:12:1) [superposition 70,2489]
[SA] new: 3775. v = a(l8(s(s(nl8))),$sum($sum(alength,1),1)) | i(l8(s(s(nl8)))) = $sum($sum(alength,1),1) (4:23) [superposition 3010,72]
[SA] new: 3776. v = a(l8(s(s(nl8))),$sum($sum(alength,1),1)) | i(l8(s(s(nl8)))) = $sum($sum(alength,1),1) (4:23) [superposition 72,3010]
[SA] active: 2718. 'Sub'(sK2(0),nl8) | alength = sK0 | alength = 0 (4:10:1) [forward demodulation 2583,2476]
[SA] new: 3777. alength = sK0 | alength = 0 | a(main_end,X0) = a(l11(sK2(0)),X0) | i(l8(sK1(X0,sK2(0)))) = X0 (5:23) [resolution 2718,150]
[SA] new: 3778. alength = sK0 | alength = 0 | $less(i(l8(sK2(0))),alength) (5:12) [resolution 2718,56]
[SA] new: 3779. 'Sub'(zero,nl8) | alength = sK0 | alength = 0 | alength = 0 (7:12) [superposition 2718,3571]
[SA] new: 3780. 'Sub'(zero,nl8) | alength = sK0 | alength = 0 | 0 = sK0 (5:12) [superposition 2718,3207]
[SA] new: 3781. 'Sub'(zero,nl8) | alength = sK0 | alength = 0 (7:9) [duplicate literal removal 3779]
[SA] passive: 3778. alength = sK0 | alength = 0 | $less(i(l8(sK2(0))),alength) (5:12) [resolution 2718,56]
[SA] active: 3017. alength != $sum($sum(alength,1),1) | s(s(nl8)) = nl8 (3:12:1) [superposition 2483,2489]
[SA] active: 3195. v = a(l11(sK2(0)),0) | 0 = sK0 (4:10:1) [superposition 70,2831]
[SA] new: 3782. v = a(l11(zero),0) | 0 = sK0 | alength = 0 (7:12) [superposition 3195,3571]
[SA] new: 3783. v = a(l11(zero),0) | 0 = sK0 | 0 = sK0 (5:12) [superposition 3195,3207]
[SA] new: 3784. v = a(l8(sK2(0)),0) | 0 = sK0 | i(l8(sK2(0))) = 0 (5:16) [superposition 3195,72]
[SA] new: 3785. v = a(l8(sK2(0)),0) | i(l8(sK2(0))) = 0 | 0 = sK0 (5:16) [superposition 72,3195]
[SA] new: 3786. v = a(l11(zero),0) | 0 = sK0 (5:9) [duplicate literal removal 3783]
[SA] active: 2666. 'Sub'(sK2(0),nl8) | alength = sK0 | $less(alength,sK0) | 0 = sK0 (3:13:1) [forward demodulation 2507,2476]
[SA] new: 3787. alength = sK0 | $less(alength,sK0) | 0 = sK0 | a(main_end,X0) = a(l11(sK2(0)),X0) | i(l8(sK1(X0,sK2(0)))) = X0 (4:26) [resolution 2666,150]
[SA] new: 3788. alength = sK0 | $less(alength,sK0) | 0 = sK0 | $less(i(l8(sK2(0))),alength) (4:15) [resolution 2666,56]
[SA] new: 3789. 'Sub'(zero,nl8) | alength = sK0 | $less(alength,sK0) | 0 = sK0 | alength = 0 (7:15) [superposition 2666,3571]
[SA] new: 3790. 'Sub'(zero,nl8) | alength = sK0 | $less(alength,sK0) | 0 = sK0 | 0 = sK0 (5:15) [superposition 2666,3207]
[SA] new: 3791. 'Sub'(zero,nl8) | alength = sK0 | $less(alength,sK0) | 0 = sK0 (5:12) [duplicate literal removal 3790]
[SA] active: 3202. alength != 0 | nl8 = sK2(0) | 0 = sK0 (4:10:1) [superposition 2483,2831]
[SA] active: 3018. $sum($sum(alength,1),1) != sK0 | s(s(nl8)) = sK2(sK0) (3:13:1) [superposition 2839,2489]
[SA] active: 3208. i(l8(s(sK2(0)))) = 1 | 0 = sK0 (4:10:1) [evaluation 3196]
[SA] new: 3792. i(l8(s(zero))) = 1 | 0 = sK0 | alength = 0 (7:12) [superposition 3208,3571]
[SA] new: 3793. i(l8(s(zero))) = 1 | 0 = sK0 | 0 = sK0 (5:12) [superposition 3208,3207]
[SA] new: 3794. i(l8(X0)) != 1 | s(sK2(0)) = X0 | 0 = sK0 (5:13) [superposition 48,3208]
[SA] new: 3795. i(l8(X1)) != 1 | s(sK2(0)) = X1 | 0 = sK0 (5:13) [superposition 48,3208]
[SA] new: 3796. v = a(l11(s(sK2(0))),1) | 0 = sK0 (5:11) [superposition 70,3208]
[SA] new: 3797. i(l8(s(s(sK2(0))))) = $sum(1,1) | 0 = sK0 (5:13) [superposition 74,3208]
[SA] new: 3798. 1 != 0 | zero = s(sK2(0)) | 0 = sK0 (5:11) [superposition 108,3208]
[SA] new: 3799. 1 != 1 | s(zero) = s(sK2(0)) | 0 = sK0 (5:12) [superposition 121,3208]
[SA] new: 3800. 1 != 2 | s(s(zero)) = s(sK2(0)) | 0 = sK0 (5:13) [superposition 130,3208]
[SA] new: 3801. 1 != 3 | s(sK2(0)) = s(s(s(zero))) | 0 = sK0 (5:14) [superposition 183,3208]
[SA] new: 3802. v = a(l8(s(s(sK2(0)))),1) | 0 = sK0 (5:12) [superposition 349,3208]
[SA] new: 3803. alength != 1 | s(sK2(0)) = nl8 | 0 = sK0 (5:11) [superposition 2483,3208]
[SA] new: 3804. 1 != $sum(alength,1) | s(nl8) = s(sK2(0)) | 0 = sK0 (5:14) [superposition 2490,3208]
[SA] new: 3805. 1 != sK0 | s(sK2(0)) = sK2(sK0) | 0 = sK0 (5:12) [superposition 2839,3208]
[SA] new: 3806. 1 != $sum(sK0,1) | s(sK2(0)) = s(sK2(sK0)) | 0 = sK0 (5:15) [superposition 2985,3208]
[SA] new: 3807. $sum(1,-1) != sK0 | s(sK2(0)) = s(sK2(sK0)) | 0 = sK0 (5:15) [evaluation 3806]
[SA] new: 3808. 0 != sK0 | s(sK2(0)) = s(sK2(sK0)) | 0 = sK0 (5:13) [evaluation 3807]
[SA] new: 3809. alength != $sum(1,-1) | s(nl8) = s(sK2(0)) | 0 = sK0 (5:14) [evaluation 3804]
[SA] new: 3810. alength != 0 | s(nl8) = s(sK2(0)) | 0 = sK0 (5:12) [evaluation 3809]
[SA] new: 3811. alength != 0 | nl8 = sK2(0) | 0 = sK0 (5:10) [term algebras injectivity 3810]
[SA] new: 3812. s(zero) = s(sK2(0)) | 0 = sK0 (5:9) [trivial inequality removal 3799]
[SA] new: 3813. zero = sK2(0) | 0 = sK0 (5:7) [term algebras injectivity 3812]
[SA] new: 3814. i(l8(s(s(sK2(0))))) = 2 | 0 = sK0 (5:11) [evaluation 3797]
[SA] new: 3815. i(l8(s(zero))) = 1 | 0 = sK0 (5:9) [duplicate literal removal 3793]
[SA] passive: 3794. i(l8(X0)) != 1 | s(sK2(0)) = X0 | 0 = sK0 (5:13) [superposition 48,3208]
[SA] passive: 3796. v = a(l11(s(sK2(0))),1) | 0 = sK0 (5:11) [superposition 70,3208]
[SA] passive: 3814. i(l8(s(s(sK2(0))))) = 2 | 0 = sK0 (5:11) [evaluation 3797]
[SA] passive: 3802. v = a(l8(s(s(sK2(0)))),1) | 0 = sK0 (5:12) [superposition 349,3208]
[SA] passive: 3803. alength != 1 | s(sK2(0)) = nl8 | 0 = sK0 (5:11) [superposition 2483,3208]
[SA] passive: 3805. 1 != sK0 | s(sK2(0)) = sK2(sK0) | 0 = sK0 (5:12) [superposition 2839,3208]
[SA] active: 3052. $sum(alength,1) != $sum($sum(alength,1),1) | s(nl8) = nl8 (3:13:1) [term algebras injectivity 3045]
[SA] active: 2506. alength != 4 | s(s(s(s(zero)))) = nl8 (5:10:1) [backward demodulation 2476,306]
[SA] active: 3008. i(l8(X0)) != $sum($sum(alength,1),1) | s(s(nl8)) = X0 (3:14:1) [superposition 48,2489]
[SA] new: 3816. $sum($sum(alength,1),1) != 2 | s(s(zero)) = s(s(nl8)) (4:14) [superposition 3008,124]
[SA] new: 3817. $sum($sum(alength,1),1) != 3 | s(s(nl8)) = s(s(s(zero))) (4:15) [superposition 3008,133]
[SA] new: 3818. $sum($sum(alength,1),1) != 4 | s(s(nl8)) = s(s(s(s(zero)))) (5:16) [superposition 3008,186]
[SA] new: 3819. $sum($sum(alength,1),1) != 5 | s(s(nl8)) = s(s(s(s(s(zero))))) (6:17) [superposition 3008,312]
[SA] new: 3820. $sum($sum(alength,1),1) != 6 | s(s(nl8)) = s(s(s(s(s(s(zero)))))) (7:18) [superposition 3008,640]
[SA] new: 3822. $sum($sum(alength,1),1) != $sum($sum(sK0,1),1) | s(s(nl8)) = s(s(sK2(sK0))) (4:19) [superposition 3008,2988]
[SA] new: 3823. $sum(sK0,1) != $sum($sum(alength,1),1) | s(s(nl8)) = s(sK2(sK0)) (4:16) [superposition 3008,2842]
[SA] new: 3824. 1 != $sum($sum(alength,1),1) | s(s(nl8)) = s(sK2(0)) | 0 = sK0 (5:17) [superposition 3008,3208]
[SA] new: 3825. $sum($sum(alength,1),1) != X0 | s(s(nl8)) = sK1(X0,p(nl8)) | a(main_end,X0) = a(l11(p(nl8)),X0) | zero = nl8 (4:27) [superposition 3008,226]
[SA] new: 3826. alength != $sum($sum(1,-1),-1) | s(s(nl8)) = s(sK2(0)) | 0 = sK0 (5:17) [evaluation 3824]
[SA] new: 3827. alength != $sum(0,-1) | s(s(nl8)) = s(sK2(0)) | 0 = sK0 (5:15) [evaluation 3826]
[SA] new: 3828. alength != -1 | s(s(nl8)) = s(sK2(0)) | 0 = sK0 (5:13) [evaluation 3827]
[SA] new: 3829. alength != -1 | s(nl8) = sK2(0) | 0 = sK0 (5:11) [term algebras injectivity 3828]
[SA] new: 3830. $sum(sK0,1) != $sum($sum(alength,1),1) | s(nl8) = sK2(sK0) (4:14) [term algebras injectivity 3823]
[SA] new: 3831. $sum($sum(alength,1),1) != $sum($sum(sK0,1),1) | s(nl8) = s(sK2(sK0)) (4:17) [term algebras injectivity 3822]
[SA] new: 3832. $sum($sum(alength,1),1) != $sum($sum(sK0,1),1) | nl8 = sK2(sK0) (4:15) [term algebras injectivity 3831]
[SA] new: 3833. alength != $sum($sum(6,-1),-1) | s(s(nl8)) = s(s(s(s(s(s(zero)))))) (7:18) [evaluation 3820]
[SA] new: 3834. alength != $sum(5,-1) | s(s(nl8)) = s(s(s(s(s(s(zero)))))) (7:16) [evaluation 3833]
[SA] new: 3835. alength != 4 | s(s(nl8)) = s(s(s(s(s(s(zero)))))) (7:14) [evaluation 3834]
[SA] new: 3836. alength != 4 | s(nl8) = s(s(s(s(s(zero))))) (7:12) [term algebras injectivity 3835]
[SA] new: 3837. alength != 4 | s(s(s(s(zero)))) = nl8 (7:10) [term algebras injectivity 3836]
[SA] new: 3838. alength != $sum($sum(5,-1),-1) | s(s(nl8)) = s(s(s(s(s(zero))))) (6:17) [evaluation 3819]
[SA] new: 3839. alength != $sum(4,-1) | s(s(nl8)) = s(s(s(s(s(zero))))) (6:15) [evaluation 3838]
[SA] new: 3840. alength != 3 | s(s(nl8)) = s(s(s(s(s(zero))))) (6:13) [evaluation 3839]
[SA] new: 3841. alength != 3 | s(nl8) = s(s(s(s(zero)))) (6:11) [term algebras injectivity 3840]
[SA] new: 3842. alength != 3 | s(s(s(zero))) = nl8 (6:9) [term algebras injectivity 3841]
[SA] new: 3843. alength != $sum($sum(4,-1),-1) | s(s(nl8)) = s(s(s(s(zero)))) (5:16) [evaluation 3818]
[SA] new: 3844. alength != $sum(3,-1) | s(s(nl8)) = s(s(s(s(zero)))) (5:14) [evaluation 3843]
[SA] new: 3845. alength != 2 | s(s(nl8)) = s(s(s(s(zero)))) (5:12) [evaluation 3844]
[SA] new: 3846. alength != 2 | s(nl8) = s(s(s(zero))) (5:10) [term algebras injectivity 3845]
[SA] new: 3847. alength != 2 | s(s(zero)) = nl8 (5:8) [term algebras injectivity 3846]
[SA] new: 3848. alength != $sum($sum(3,-1),-1) | s(s(nl8)) = s(s(s(zero))) (4:15) [evaluation 3817]
[SA] new: 3849. alength != $sum(2,-1) | s(s(nl8)) = s(s(s(zero))) (4:13) [evaluation 3848]
[SA] new: 3850. alength != 1 | s(s(nl8)) = s(s(s(zero))) (4:11) [evaluation 3849]
[SA] new: 3851. alength != 1 | s(nl8) = s(s(zero)) (4:9) [term algebras injectivity 3850]
[SA] new: 3852. alength != 1 | s(zero) = nl8 (4:7) [term algebras injectivity 3851]
[SA] new: 3853. alength != $sum($sum(2,-1),-1) | s(s(zero)) = s(s(nl8)) (4:14) [evaluation 3816]
[SA] new: 3854. alength != $sum(1,-1) | s(s(zero)) = s(s(nl8)) (4:12) [evaluation 3853]
[SA] new: 3855. alength != 0 | s(s(zero)) = s(s(nl8)) (4:10) [evaluation 3854]
[SA] new: 3856. alength != 0 | s(zero) = s(nl8) (4:8) [term algebras injectivity 3855]
[SA] new: 3857. alength != 0 | zero = nl8 (4:6) [term algebras injectivity 3856]
[SA] passive: 3832. $sum($sum(alength,1),1) != $sum($sum(sK0,1),1) | nl8 = sK2(sK0) (4:15) [term algebras injectivity 3831]
[SA] passive: 3825. $sum($sum(alength,1),1) != X0 | s(s(nl8)) = sK1(X0,p(nl8)) | a(main_end,X0) = a(l11(p(nl8)),X0) | zero = nl8 (4:27) [superposition 3008,226]
[SA] active: 2614. 'Sub'(sK2(0),nl8) | alength = 0 | 0 = sK0 (5:10:1) [backward demodulation 2476,1632]
[SA] new: 3858. alength = 0 | 0 = sK0 | a(main_end,X0) = a(l11(sK2(0)),X0) | i(l8(sK1(X0,sK2(0)))) = X0 (6:23) [resolution 2614,150]
[SA] new: 3859. alength = 0 | 0 = sK0 | $less(i(l8(sK2(0))),alength) (6:12) [resolution 2614,56]
[SA] new: 3860. 'Sub'(zero,nl8) | alength = 0 | 0 = sK0 | alength = 0 (7:12) [superposition 2614,3571]
[SA] new: 3861. 'Sub'(zero,nl8) | alength = 0 | 0 = sK0 | 0 = sK0 (6:12) [superposition 2614,3207]
[SA] new: 3862. 'Sub'(zero,nl8) | alength = 0 | 0 = sK0 (6:9) [duplicate literal removal 3861]
[SA] new: 3863. 'Sub'(zero,nl8) | alength = 0 | 0 = sK0 (7:9) [duplicate literal removal 3860]
[SA] active: 3011. i(l8(s(s(s(nl8))))) = $sum($sum($sum(alength,1),1),1) (3:14:1) [superposition 74,2489]
[SA] active: 2846. sK0 != 3 | s(s(s(zero))) = sK2(sK0) (5:10:1) [superposition 183,2833]
[SA] active: 2665. $less(alength,sK0) | alength = 0 | 0 = sK0 | i(l8(sK2(0))) = 0 (3:15:1) [forward demodulation 2504,2476]
[SA] new: 3864. alength = 0 | 0 = sK0 | i(l8(sK2(0))) = 0 | alength = sK0 (6:15) [resolution 2665,2704]
[SA] new: 3865. alength = 0 | 0 = sK0 | i(l8(sK2(0))) = 0 | alength = X0 | $less(X0,alength) | $less(alength,X0) (4:21) [resolution 2665,247]
[SA] new: 3866. alength = 0 | 0 = sK0 | i(l8(sK2(0))) = 0 | sK0 = X1 | $less(X1,alength) | $less(alength,X1) (4:21) [resolution 2665,199]
[SA] new: 3867. alength = 0 | 0 = sK0 | i(l8(sK2(0))) = 0 | $less(alength,alength) (4:15) [resolution 2665,96]
[SA] new: 3868. alength = 0 | 0 = sK0 | i(l8(sK2(0))) = 0 | ~$less(X2,alength) | $less(X2,sK0) (4:18) [resolution 2665,14]
[SA] new: 3869. alength = 0 | 0 = sK0 | i(l8(sK2(0))) = 0 | sK0 = X3 | $less(X3,sK0) | $less(alength,X3) (4:21) [resolution 2665,99]
[SA] new: 3870. alength = 0 | 0 = sK0 | i(l8(sK2(0))) = 0 | sK0 = sK0 | $less(sK0,sK0) | $less(alength,alength) (4:21) [resolution 2665,205]
[SA] new: 3871. alength = 0 | 0 = sK0 | i(l8(sK2(0))) = 0 | alength = sK0 | $less(sK0,sK0) | $less(alength,alength) (4:21) [resolution 2665,254]
[SA] new: 3872. alength = 0 | 0 = sK0 | i(l8(sK2(0))) = 0 | sK0 = sK0 | 0 = sK0 | $less(0,sK0) | $less(alength,sK0) (4:24) [resolution 2665,333]
[SA] new: 3873. alength = 0 | 0 = sK0 | i(l8(sK2(0))) = 0 | i(l8(sK2(sK0))) = sK0 | alength = sK0 | $less(alength,sK0) | $less(alength,0) (4:27) [resolution 2665,2950]
[SA] new: 3874. alength = 0 | 0 = sK0 | i(l8(sK2(0))) = 0 | sK0 = sK0 | $less(0,sK0) | $less(alength,sK0) (4:21) [duplicate literal removal 3872]
[SA] active: 3559. v = a(l11(sK2(0)),0) | alength = 0 (6:10:1) [superposition 70,2654]
[SA] new: 3875. v = a(l11(zero),0) | alength = 0 | alength = 0 (7:12) [superposition 3559,3571]
[SA] new: 3876. v = a(l11(zero),0) | alength = 0 | 0 = sK0 (7:12) [superposition 3559,3207]
[SA] new: 3877. v = a(l8(sK2(0)),0) | alength = 0 | i(l8(sK2(0))) = 0 (7:16) [superposition 3559,72]
[SA] new: 3878. v = a(l8(sK2(0)),0) | i(l8(sK2(0))) = 0 | alength = 0 (7:16) [superposition 72,3559]
[SA] new: 3879. v = a(l11(zero),0) | alength = 0 (7:9) [duplicate literal removal 3875]
[SA] active: 2667. $less(alength,sK0) | alength = sK0 | 0 = sK0 | i(l8(sK2(0))) = 0 (3:15:1) [forward demodulation 2508,2476]
[SA] new: 3880. alength = sK0 | 0 = sK0 | i(l8(sK2(0))) = 0 | alength = sK0 (6:15) [resolution 2667,2704]
[SA] new: 3881. alength = sK0 | 0 = sK0 | i(l8(sK2(0))) = 0 | alength = X0 | $less(X0,alength) | $less(alength,X0) (4:21) [resolution 2667,247]
[SA] new: 3882. alength = sK0 | 0 = sK0 | i(l8(sK2(0))) = 0 | sK0 = X1 | $less(X1,alength) | $less(alength,X1) (4:21) [resolution 2667,199]
[SA] new: 3883. alength = sK0 | 0 = sK0 | i(l8(sK2(0))) = 0 | $less(alength,alength) (4:15) [resolution 2667,96]
[SA] new: 3884. alength = sK0 | 0 = sK0 | i(l8(sK2(0))) = 0 | ~$less(X2,alength) | $less(X2,sK0) (4:18) [resolution 2667,14]
[SA] new: 3885. alength = sK0 | 0 = sK0 | i(l8(sK2(0))) = 0 | sK0 = X3 | $less(X3,sK0) | $less(alength,X3) (4:21) [resolution 2667,99]
[SA] new: 3886. alength = sK0 | 0 = sK0 | i(l8(sK2(0))) = 0 | sK0 = sK0 | $less(sK0,sK0) | $less(alength,alength) (4:21) [resolution 2667,205]
[SA] new: 3887. alength = sK0 | 0 = sK0 | i(l8(sK2(0))) = 0 | alength = sK0 | $less(sK0,sK0) | $less(alength,alength) (4:21) [resolution 2667,254]
[SA] new: 3888. alength = sK0 | 0 = sK0 | i(l8(sK2(0))) = 0 | sK0 = sK0 | 0 = sK0 | $less(0,sK0) | $less(alength,sK0) (4:24) [resolution 2667,333]
[SA] new: 3889. alength = sK0 | 0 = sK0 | i(l8(sK2(0))) = 0 | i(l8(sK2(sK0))) = sK0 | alength = sK0 | $less(alength,sK0) | $less(alength,0) (4:27) [resolution 2667,2950]
[SA] new: 3890. alength = sK0 | 0 = sK0 | i(l8(sK2(0))) = 0 | i(l8(sK2(sK0))) = sK0 | $less(alength,sK0) | $less(alength,0) (4:24) [duplicate literal removal 3889]
[SA] new: 3891. alength = sK0 | 0 = sK0 | i(l8(sK2(0))) = 0 | sK0 = sK0 | $less(0,sK0) | $less(alength,sK0) (4:21) [duplicate literal removal 3888]
[SA] new: 3892. alength = sK0 | 0 = sK0 | i(l8(sK2(0))) = 0 | $less(sK0,sK0) | $less(alength,alength) (4:18) [duplicate literal removal 3887]
[SA] new: 3893. alength = sK0 | 0 = sK0 | i(l8(sK2(0))) = 0 (6:12) [duplicate literal removal 3880]
[SA] active: 3572. i(l8(s(sK2(0)))) = 1 | alength = 0 (6:10:1) [evaluation 3560]
[SA] new: 3894. i(l8(s(zero))) = 1 | alength = 0 | alength = 0 (7:12) [superposition 3572,3571]
[SA] new: 3895. i(l8(s(zero))) = 1 | alength = 0 | 0 = sK0 (7:12) [superposition 3572,3207]
[SA] new: 3896. i(l8(X0)) != 1 | s(sK2(0)) = X0 | alength = 0 (7:13) [superposition 48,3572]
[SA] new: 3897. i(l8(X1)) != 1 | s(sK2(0)) = X1 | alength = 0 (7:13) [superposition 48,3572]
[SA] new: 3898. v = a(l11(s(sK2(0))),1) | alength = 0 (7:11) [superposition 70,3572]
[SA] new: 3899. i(l8(s(s(sK2(0))))) = $sum(1,1) | alength = 0 (7:13) [superposition 74,3572]
[SA] new: 3900. 1 != 0 | zero = s(sK2(0)) | alength = 0 (7:11) [superposition 108,3572]
[SA] new: 3901. 1 != 1 | s(zero) = s(sK2(0)) | alength = 0 (7:12) [superposition 121,3572]
[SA] new: 3902. 1 != 2 | s(s(zero)) = s(sK2(0)) | alength = 0 (7:13) [superposition 130,3572]
[SA] new: 3903. 1 != 3 | s(sK2(0)) = s(s(s(zero))) | alength = 0 (7:14) [superposition 183,3572]
[SA] new: 3904. v = a(l8(s(s(sK2(0)))),1) | alength = 0 (7:12) [superposition 349,3572]
[SA] new: 3905. alength != 1 | s(sK2(0)) = nl8 | alength = 0 (7:11) [superposition 2483,3572]
[SA] new: 3906. 1 != $sum(alength,1) | s(nl8) = s(sK2(0)) | alength = 0 (7:14) [superposition 2490,3572]
[SA] new: 3907. 1 != sK0 | s(sK2(0)) = sK2(sK0) | alength = 0 (7:12) [superposition 2839,3572]
[SA] new: 3908. 1 != $sum(sK0,1) | s(sK2(0)) = s(sK2(sK0)) | alength = 0 (7:15) [superposition 2985,3572]
[SA] new: 3909. 1 != $sum($sum(alength,1),1) | s(s(nl8)) = s(sK2(0)) | alength = 0 (7:17) [superposition 3008,3572]
[SA] new: 3910. alength != $sum($sum(1,-1),-1) | s(s(nl8)) = s(sK2(0)) | alength = 0 (7:17) [evaluation 3909]
[SA] new: 3911. alength != $sum(0,-1) | s(s(nl8)) = s(sK2(0)) | alength = 0 (7:15) [evaluation 3910]
[SA] new: 3912. alength != -1 | s(s(nl8)) = s(sK2(0)) | alength = 0 (7:13) [evaluation 3911]
[SA] new: 3913. alength != -1 | s(nl8) = sK2(0) | alength = 0 (7:11) [term algebras injectivity 3912]
[SA] new: 3914. $sum(1,-1) != sK0 | s(sK2(0)) = s(sK2(sK0)) | alength = 0 (7:15) [evaluation 3908]
[SA] new: 3915. 0 != sK0 | s(sK2(0)) = s(sK2(sK0)) | alength = 0 (7:13) [evaluation 3914]
[SA] new: 3916. 0 != sK0 | sK2(0) = sK2(sK0) | alength = 0 (7:11) [term algebras injectivity 3915]
[SA] new: 3917. alength != $sum(1,-1) | s(nl8) = s(sK2(0)) | alength = 0 (7:14) [evaluation 3906]
[SA] new: 3918. alength != 0 | s(nl8) = s(sK2(0)) | alength = 0 (7:12) [evaluation 3917]
[SA] new: 3919. s(zero) = s(sK2(0)) | alength = 0 (7:9) [trivial inequality removal 3901]
[SA] new: 3920. zero = sK2(0) | alength = 0 (7:7) [term algebras injectivity 3919]
[SA] new: 3921. i(l8(s(s(sK2(0))))) = 2 | alength = 0 (7:11) [evaluation 3899]
[SA] new: 3922. i(l8(s(zero))) = 1 | alength = 0 (7:9) [duplicate literal removal 3894]
[SA] passive: 3896. i(l8(X0)) != 1 | s(sK2(0)) = X0 | alength = 0 (7:13) [superposition 48,3572]
[SA] passive: 3898. v = a(l11(s(sK2(0))),1) | alength = 0 (7:11) [superposition 70,3572]
[SA] passive: 3921. i(l8(s(s(sK2(0))))) = 2 | alength = 0 (7:11) [evaluation 3899]
[SA] passive: 3904. v = a(l8(s(s(sK2(0)))),1) | alength = 0 (7:12) [superposition 349,3572]
[SA] passive: 3905. alength != 1 | s(sK2(0)) = nl8 | alength = 0 (7:11) [superposition 2483,3572]
[SA] passive: 3907. 1 != sK0 | s(sK2(0)) = sK2(sK0) | alength = 0 (7:12) [superposition 2839,3572]
[SA] active: 2772. $less(sK0,X0) | $less(X0,0) | i(l8(sK2(X0))) = X0 | sK0 = X0 (3:15:2) [resolution 2480,106]
[SA] new: 3923. $less(0,0) | i(l8(sK2(0))) = 0 | 0 = sK0 (4:12) [resolution 2772,46]
[SA] new: 3924. $less(sK0,0) | i(l8(sK2(sK0))) = sK0 | sK0 = sK0 (4:12) [resolution 2772,13]
[SA] new: 3925. $less(X0,0) | i(l8(sK2(X0))) = X0 | sK0 = X0 | ~$less(X1,sK0) | $less(X1,X0) (4:18) [resolution 2772,14]
[SA] new: 3926. $less(X2,0) | i(l8(sK2(X2))) = X2 | sK0 = X2 | X2 = X3 | $less(X3,X2) | $less(sK0,X3) (4:21) [resolution 2772,99]
[SA] new: 3927. $less(X4,0) | i(l8(sK2(X4))) = X4 | sK0 = X4 | sK0 = X4 | $less(sK0,X4) | $less(sK0,alength) (4:21) [resolution 2772,205]
[SA] new: 3928. $less(X5,0) | i(l8(sK2(X5))) = X5 | sK0 = X5 | alength = X5 | $less(sK0,X5) | $less(sK0,alength) (4:21) [resolution 2772,254]
[SA] new: 3929. $less(X6,0) | i(l8(sK2(X6))) = X6 | sK0 = X6 | sK0 = X6 | 0 = sK0 | $less(0,X6) | $less(sK0,sK0) (4:24) [resolution 2772,333]
[SA] new: 3930. $less(X7,0) | i(l8(sK2(X7))) = X7 | sK0 = X7 | i(l8(sK2(X7))) = X7 | alength = X7 | $less(alength,X7) | $less(sK0,0) (4:27) [resolution 2772,2950]
[SA] new: 3931. $less(alength,0) | alength = i(l8(sK2(alength))) | alength = sK0 | $less(sK0,0) | a(main_end,X8) = a(l11(sK2(sK0)),X8) | i(l8(sK1(X8,sK2(sK0)))) = X8 (4:32) [resolution 2772,2770]
[SA] new: 3932. $less(alength,0) | alength = i(l8(sK2(alength))) | alength = sK0 | alength = X9 | $less(X9,0) | i(l8(sK2(X9))) = X9 | $less(sK0,X9) (4:27) [resolution 2772,2660]
[SA] new: 3933. $less(alength,0) | alength = i(l8(sK2(alength))) | alength = sK0 | $less(sK0,0) | i(l8(sK2(sK0))) = sK0 (4:21) [resolution 2772,2480]
[SA] new: 3934. $less(i(l8(p(nl8))),0) | i(l8(p(nl8))) = i(l8(sK2(i(l8(p(nl8)))))) | i(l8(p(nl8))) = sK0 | zero = nl8 | $less(sK0,alength) (4:30) [resolution 2772,194]
[SA] new: 3935. $less(i(l8(sK2(X10))),0) | i(l8(sK2(X10))) = i(l8(sK2(i(l8(sK2(X10)))))) | i(l8(sK2(X10))) = sK0 | $less(X10,0) | ~$less(X10,alength) | $less(sK0,alength) (4:33) [resolution 2772,2920]
[SA] new: 3936. $less(0,0) | i(l8(sK2(0))) = 0 | 0 = sK0 | sK0 = X11 | 0 = sK0 | $less(X11,sK0) | $less(sK0,X11) (4:24) [resolution 2772,324]
[SA] new: 3937. $less(0,0) | i(l8(sK2(0))) = 0 | 0 = sK0 | alength = 0 | $less(sK0,alength) (5:18) [resolution 2772,262]
[SA] new: 3938. $less(0,0) | i(l8(sK2(0))) = 0 | 0 = sK0 | 0 = sK0 | $less(sK0,alength) (4:18) [resolution 2772,116]
[SA] new: 3939. $less(0,0) | i(l8(sK2(0))) = 0 | 0 = sK0 | 0 = sK0 | $less(sK0,sK0) (4:18) [resolution 2772,114]
[SA] new: 3940. $less(sK0,0) | i(l8(sK2(sK0))) = sK0 | sK0 = sK0 | alength = X12 | $less(X12,alength) | $less(sK0,X12) (4:21) [resolution 2772,247]
[SA] new: 3941. $less(sK0,0) | i(l8(sK2(sK0))) = sK0 | sK0 = sK0 | sK0 = X13 | $less(X13,alength) | $less(sK0,X13) (4:21) [resolution 2772,199]
[SA] new: 3942. $less(sK0,0) | i(l8(sK2(sK0))) = sK0 | sK0 = sK0 | $less(sK0,alength) (4:15) [resolution 2772,96]
[SA] new: 3943. $less(sK0,sK0) | i(l8(sK2(sK0))) = sK0 | sK0 = sK0 (4:12) [resolution 2772,46]
[SA] new: 3944. $less(sK0,X0) | i(l8(sK2(X0))) = X0 | sK0 = X0 | sK0 = X1 | 0 = sK0 | $less(X1,sK0) | $less(X0,X1) (4:24) [resolution 2772,324]
[SA] new: 3945. $less(sK0,X2) | i(l8(sK2(X2))) = X2 | sK0 = X2 | alength = 0 | $less(X2,alength) (5:18) [resolution 2772,262]
[SA] new: 3946. $less(sK0,X3) | i(l8(sK2(X3))) = X3 | sK0 = X3 | 0 = sK0 | $less(X3,alength) (4:18) [resolution 2772,116]
[SA] new: 3947. $less(sK0,X4) | i(l8(sK2(X4))) = X4 | sK0 = X4 | 0 = sK0 | $less(X4,sK0) (4:18) [resolution 2772,114]
[SA] new: 3948. $less(sK0,0) | i(l8(sK2(0))) = 0 | 0 = sK0 (4:12) [resolution 2772,13]
[SA] new: 3949. $less(sK0,X5) | i(l8(sK2(X5))) = X5 | sK0 = X5 | ~$less(X6,X5) | $less(X6,0) (4:18) [resolution 2772,14]
[SA] new: 3950. $less(sK0,X7) | i(l8(sK2(X7))) = X7 | sK0 = X7 | 0 = X8 | $less(X8,0) | $less(X7,X8) (4:21) [resolution 2772,99]
[SA] new: 3951. $less(sK0,X9) | i(l8(sK2(X9))) = X9 | sK0 = X9 | 0 = sK0 | $less(sK0,0) | $less(X9,alength) (4:21) [resolution 2772,205]
[SA] new: 3952. $less(sK0,X10) | i(l8(sK2(X10))) = X10 | sK0 = X10 | alength = 0 | $less(sK0,0) | $less(X10,alength) (4:21) [resolution 2772,254]
[SA] new: 3953. $less(sK0,X11) | i(l8(sK2(X11))) = X11 | sK0 = X11 | 0 = sK0 | 0 = sK0 | $less(0,0) | $less(X11,sK0) (4:24) [resolution 2772,333]
[SA] new: 3954. $less(sK0,X12) | i(l8(sK2(X12))) = X12 | sK0 = X12 | i(l8(sK2(0))) = 0 | alength = 0 | $less(alength,0) | $less(X12,0) (4:27) [resolution 2772,2950]
[SA] new: 3955. $less(sK0,X11) | i(l8(sK2(X11))) = X11 | sK0 = X11 | 0 = sK0 | $less(0,0) | $less(X11,sK0) (4:21) [duplicate literal removal 3953]
[SA] new: 3956. $less(sK0,X11) | i(l8(sK2(X11))) = X11 | sK0 = X11 | 0 = sK0 | $less(X11,sK0) (4:18) [evaluation 3955]
[SA] new: 3957. $less(0,0) | i(l8(sK2(0))) = 0 | 0 = sK0 | $less(sK0,sK0) (4:15) [duplicate literal removal 3939]
[SA] new: 3958. i(l8(sK2(0))) = 0 | 0 = sK0 | $less(sK0,sK0) (4:12) [evaluation 3957]
[SA] new: 3959. $less(0,0) | i(l8(sK2(0))) = 0 | 0 = sK0 | $less(sK0,alength) (4:15) [duplicate literal removal 3938]
[SA] new: 3960. i(l8(sK2(0))) = 0 | 0 = sK0 | $less(sK0,alength) (4:12) [evaluation 3959]
[SA] new: 3961. i(l8(sK2(0))) = 0 | 0 = sK0 | alength = 0 | $less(sK0,alength) (5:15) [evaluation 3937]
[SA] new: 3962. $less(0,0) | i(l8(sK2(0))) = 0 | 0 = sK0 | sK0 = X11 | $less(X11,sK0) | $less(sK0,X11) (4:21) [duplicate literal removal 3936]
[SA] new: 3963. i(l8(sK2(0))) = 0 | 0 = sK0 | sK0 = X11 | $less(X11,sK0) | $less(sK0,X11) (4:18) [evaluation 3962]
[SA] new: 3964. $less(X7,0) | i(l8(sK2(X7))) = X7 | sK0 = X7 | alength = X7 | $less(alength,X7) | $less(sK0,0) (4:21) [duplicate literal removal 3930]
[SA] new: 3965. $less(X6,0) | i(l8(sK2(X6))) = X6 | sK0 = X6 | 0 = sK0 | $less(0,X6) | $less(sK0,sK0) (4:21) [duplicate literal removal 3929]
[SA] new: 3966. $less(X4,0) | i(l8(sK2(X4))) = X4 | sK0 = X4 | $less(sK0,X4) | $less(sK0,alength) (4:18) [duplicate literal removal 3927]
[SA] new: 3967. i(l8(sK2(0))) = 0 | 0 = sK0 (4:9) [evaluation 3923]
[SA] passive: 3925. $less(X0,0) | i(l8(sK2(X0))) = X0 | sK0 = X0 | ~$less(X1,sK0) | $less(X1,X0) (4:18) [resolution 2772,14]
[SA] passive: 3949. $less(sK0,X5) | i(l8(sK2(X5))) = X5 | sK0 = X5 | ~$less(X6,X5) | $less(X6,0) (4:18) [resolution 2772,14]
[SA] passive: 3950. $less(sK0,X7) | i(l8(sK2(X7))) = X7 | sK0 = X7 | 0 = X8 | $less(X8,0) | $less(X7,X8) (4:21) [resolution 2772,99]
[SA] active: 1867. alength != 1 | s(zero) = nl8 | alength = 0 (7:10:1) [superposition 153,1844]
[SA] active: 3551. i(l8(sK1(X94,sK2(sK0)))) = X94 | a(main_end,X94) = a(l11(sK2(sK0)),X94) (3:17:1) [subsumption resolution 3529,46]
[SA] new: 3968. i(l8(X1)) != X0 | sK1(X0,sK2(sK0)) = X1 | a(main_end,X0) = a(l11(sK2(sK0)),X0) (4:20) [superposition 48,3551]
[SA] new: 3969. i(l8(X3)) != X2 | sK1(X2,sK2(sK0)) = X3 | a(main_end,X2) = a(l11(sK2(sK0)),X2) (4:20) [superposition 48,3551]
[SA] new: 3970. v = a(l11(sK1(X4,sK2(sK0))),X4) | a(main_end,X4) = a(l11(sK2(sK0)),X4) (4:18) [superposition 70,3551]
[SA] new: 3971. i(l8(s(sK1(X5,sK2(sK0))))) = $sum(X5,1) | a(main_end,X5) = a(l11(sK2(sK0)),X5) (4:20) [superposition 74,3551]
[SA] new: 3972. 0 != X6 | zero = sK1(X6,sK2(sK0)) | a(main_end,X6) = a(l11(sK2(sK0)),X6) (4:18) [superposition 108,3551]
[SA] new: 3973. 1 != X7 | s(zero) = sK1(X7,sK2(sK0)) | a(main_end,X7) = a(l11(sK2(sK0)),X7) (4:19) [superposition 121,3551]
[SA] new: 3974. 2 != X8 | s(s(zero)) = sK1(X8,sK2(sK0)) | a(main_end,X8) = a(l11(sK2(sK0)),X8) (4:20) [superposition 130,3551]
[SA] new: 3975. 3 != X9 | s(s(s(zero))) = sK1(X9,sK2(sK0)) | a(main_end,X9) = a(l11(sK2(sK0)),X9) (5:21) [superposition 183,3551]
[SA] new: 3976. v = a(l8(s(sK1(X10,sK2(sK0)))),X10) | a(main_end,X10) = a(l11(sK2(sK0)),X10) (4:19) [superposition 349,3551]
[SA] new: 3977. alength != X11 | nl8 = sK1(X11,sK2(sK0)) | a(main_end,X11) = a(l11(sK2(sK0)),X11) (4:18) [superposition 2483,3551]
[SA] new: 3978. $sum(alength,1) != X12 | s(nl8) = sK1(X12,sK2(sK0)) | a(main_end,X12) = a(l11(sK2(sK0)),X12) (4:21) [superposition 2490,3551]
[SA] new: 3979. sK0 != X13 | sK1(X13,sK2(sK0)) = sK2(sK0) | a(main_end,X13) = a(l11(sK2(sK0)),X13) (4:19) [superposition 2839,3551]
[SA] new: 3980. $sum(sK0,1) != X14 | s(sK2(sK0)) = sK1(X14,sK2(sK0)) | a(main_end,X14) = a(l11(sK2(sK0)),X14) (4:22) [superposition 2985,3551]
[SA] new: 3981. $sum($sum(alength,1),1) != X15 | s(s(nl8)) = sK1(X15,sK2(sK0)) | a(main_end,X15) = a(l11(sK2(sK0)),X15) (4:24) [superposition 3008,3551]
[SA] passive: 3968. i(l8(X1)) != X0 | sK1(X0,sK2(sK0)) = X1 | a(main_end,X0) = a(l11(sK2(sK0)),X0) (4:20) [superposition 48,3551]
[SA] passive: 3970. v = a(l11(sK1(X4,sK2(sK0))),X4) | a(main_end,X4) = a(l11(sK2(sK0)),X4) (4:18) [superposition 70,3551]
[SA] passive: 3971. i(l8(s(sK1(X5,sK2(sK0))))) = $sum(X5,1) | a(main_end,X5) = a(l11(sK2(sK0)),X5) (4:20) [superposition 74,3551]
[SA] passive: 3972. 0 != X6 | zero = sK1(X6,sK2(sK0)) | a(main_end,X6) = a(l11(sK2(sK0)),X6) (4:18) [superposition 108,3551]
[SA] passive: 3973. 1 != X7 | s(zero) = sK1(X7,sK2(sK0)) | a(main_end,X7) = a(l11(sK2(sK0)),X7) (4:19) [superposition 121,3551]
[SA] passive: 3974. 2 != X8 | s(s(zero)) = sK1(X8,sK2(sK0)) | a(main_end,X8) = a(l11(sK2(sK0)),X8) (4:20) [superposition 130,3551]
[SA] passive: 3975. 3 != X9 | s(s(s(zero))) = sK1(X9,sK2(sK0)) | a(main_end,X9) = a(l11(sK2(sK0)),X9) (5:21) [superposition 183,3551]
[SA] passive: 3976. v = a(l8(s(sK1(X10,sK2(sK0)))),X10) | a(main_end,X10) = a(l11(sK2(sK0)),X10) (4:19) [superposition 349,3551]
[SA] passive: 3977. alength != X11 | nl8 = sK1(X11,sK2(sK0)) | a(main_end,X11) = a(l11(sK2(sK0)),X11) (4:18) [superposition 2483,3551]
[SA] passive: 3978. $sum(alength,1) != X12 | s(nl8) = sK1(X12,sK2(sK0)) | a(main_end,X12) = a(l11(sK2(sK0)),X12) (4:21) [superposition 2490,3551]
[SA] passive: 3979. sK0 != X13 | sK1(X13,sK2(sK0)) = sK2(sK0) | a(main_end,X13) = a(l11(sK2(sK0)),X13) (4:19) [superposition 2839,3551]
[SA] passive: 3980. $sum(sK0,1) != X14 | s(sK2(sK0)) = sK1(X14,sK2(sK0)) | a(main_end,X14) = a(l11(sK2(sK0)),X14) (4:22) [superposition 2985,3551]
[SA] passive: 3981. $sum($sum(alength,1),1) != X15 | s(s(nl8)) = sK1(X15,sK2(sK0)) | a(main_end,X15) = a(l11(sK2(sK0)),X15) (4:24) [superposition 3008,3551]
[SA] active: 1886. v = a(l8(s(nl8)),alength) | alength = 0 (7:10:1) [superposition 354,1844]
[SA] new: 3982. v = a(l8(nl8),alength) | alength = 0 | alength = i(l8(nl8)) (8:14) [superposition 1886,440]
[SA] new: 3983. v = a(l11(nl8),alength) | alength = 0 (8:9) [superposition 1886,57]
[SA] new: 3984. v = a(l8(nl8),alength) | alength = i(l8(nl8)) | alength = 0 (8:14) [superposition 440,1886]
[SA] new: 3985. v = a(l11(nl8),alength) | alength = 0 (8:9) [superposition 57,1886]
[SA] active: 397. $less(0,X17) | $less(X18,sK0) | $less(X17,X18) | 0 = X18 | 0 = sK0 | X17 = X18 (3:18:3) [resolution 174,114]
[SA] new: 3986. $less(X0,sK0) | $less(0,X0) | 0 = X0 | 0 = sK0 | 0 = X0 (4:15) [resolution 397,13]
[SA] new: 3987. $less(X1,sK0) | $less(X2,X1) | 0 = X1 | 0 = sK0 | X1 = X2 | ~$less(X3,0) | $less(X3,X2) (4:21) [resolution 397,14]
[SA] new: 3988. $less(X4,sK0) | $less(X5,X4) | 0 = X4 | 0 = sK0 | X4 = X5 | X5 = X6 | $less(X6,X5) | $less(0,X6) (4:24) [resolution 397,99]
[SA] new: 3989. $less(X7,sK0) | $less(X8,X7) | 0 = X7 | 0 = sK0 | X7 = X8 | sK0 = X8 | $less(sK0,X8) | $less(0,alength) (4:24) [resolution 397,205]
[SA] new: 3990. $less(X9,sK0) | $less(X10,X9) | 0 = X9 | 0 = sK0 | X9 = X10 | alength = X10 | $less(sK0,X10) | $less(0,alength) (4:24) [resolution 397,254]
[SA] new: 3991. $less(X11,sK0) | $less(X12,X11) | 0 = X11 | 0 = sK0 | X11 = X12 | sK0 = X12 | 0 = sK0 | $less(0,X12) | $less(0,sK0) (4:27) [resolution 397,333]
[SA] new: 3992. $less(X13,sK0) | $less(X14,X13) | 0 = X13 | 0 = sK0 | X13 = X14 | i(l8(sK2(X14))) = X14 | alength = X14 | $less(alength,X14) | $less(0,0) (4:30) [resolution 397,2950]
[SA] new: 3993. $less(X15,sK0) | $less(alength,X15) | 0 = X15 | 0 = sK0 | alength = X15 | $less(0,0) | a(main_end,X16) = a(l11(sK2(0)),X16) | i(l8(sK1(X16,sK2(0)))) = X16 (4:35) [resolution 397,2770]
[SA] new: 3994. $less(X17,sK0) | $less(alength,X17) | 0 = X17 | 0 = sK0 | alength = X17 | alength = X18 | $less(X18,0) | i(l8(sK2(X18))) = X18 | $less(0,X18) (4:30) [resolution 397,2660]
[SA] new: 3995. $less(X19,sK0) | $less(alength,X19) | 0 = X19 | 0 = sK0 | alength = X19 | $less(0,0) | i(l8(sK2(0))) = 0 (4:24) [resolution 397,2480]
[SA] new: 3996. $less(X20,sK0) | $less(i(l8(p(nl8))),X20) | 0 = X20 | 0 = sK0 | i(l8(p(nl8))) = X20 | zero = nl8 | $less(0,alength) (4:27) [resolution 397,194]
[SA] new: 3997. $less(X21,sK0) | $less(i(l8(sK2(X22))),X21) | 0 = X21 | 0 = sK0 | i(l8(sK2(X22))) = X21 | $less(X22,0) | ~$less(X22,alength) | $less(0,alength) (4:30) [resolution 397,2920]
[SA] new: 3998. $less(X23,sK0) | $less(0,X23) | 0 = X23 | 0 = sK0 | 0 = X23 | sK0 = X24 | 0 = sK0 | $less(X24,sK0) | $less(0,X24) (4:27) [resolution 397,324]
[SA] new: 3999. $less(X25,sK0) | $less(0,X25) | 0 = X25 | 0 = sK0 | 0 = X25 | alength = 0 | $less(0,alength) (5:21) [resolution 397,262]
[SA] new: 4000. $less(X26,sK0) | $less(0,X26) | 0 = X26 | 0 = sK0 | 0 = X26 | 0 = sK0 | $less(0,alength) (4:21) [resolution 397,116]
[SA] new: 4001. $less(X27,sK0) | $less(0,X27) | 0 = X27 | 0 = sK0 | 0 = X27 | 0 = sK0 | $less(0,sK0) (4:21) [resolution 397,114]
[SA] new: 4002. $less(X28,sK0) | $less(sK0,X28) | 0 = X28 | 0 = sK0 | sK0 = X28 | alength = X29 | $less(X29,alength) | $less(0,X29) (4:24) [resolution 397,247]
[SA] new: 4003. $less(X30,sK0) | $less(sK0,X30) | 0 = X30 | 0 = sK0 | sK0 = X30 | sK0 = X31 | $less(X31,alength) | $less(0,X31) (4:24) [resolution 397,199]
[SA] new: 4004. $less(X32,sK0) | $less(sK0,X32) | 0 = X32 | 0 = sK0 | sK0 = X32 | $less(0,alength) (4:18) [resolution 397,96]
[SA] new: 4005. $less(0,X0) | $less(X0,X1) | 0 = X1 | 0 = sK0 | X0 = X1 | alength = X2 | $less(X2,alength) | $less(X1,X2) (4:24) [resolution 397,247]
[SA] new: 4006. $less(0,X3) | $less(X3,X4) | 0 = X4 | 0 = sK0 | X3 = X4 | sK0 = X5 | $less(X5,alength) | $less(X4,X5) (4:24) [resolution 397,199]
[SA] new: 4007. $less(0,X6) | $less(X6,X7) | 0 = X7 | 0 = sK0 | X6 = X7 | $less(X7,alength) (4:18) [resolution 397,96]
[SA] new: 4008. $less(0,X8) | $less(X8,sK0) | 0 = sK0 | 0 = sK0 | sK0 = X8 (4:15) [resolution 397,13]
[SA] new: 4009. $less(0,X9) | $less(X9,X10) | 0 = X10 | 0 = sK0 | X9 = X10 | ~$less(X11,X10) | $less(X11,sK0) (4:21) [resolution 397,14]
[SA] new: 4010. $less(0,X12) | $less(X12,X13) | 0 = X13 | 0 = sK0 | X12 = X13 | sK0 = X14 | $less(X14,sK0) | $less(X13,X14) (4:24) [resolution 397,99]
[SA] new: 4011. $less(0,X15) | $less(X15,X16) | 0 = X16 | 0 = sK0 | X15 = X16 | sK0 = sK0 | $less(sK0,sK0) | $less(X16,alength) (4:24) [resolution 397,205]
[SA] new: 4012. $less(0,X17) | $less(X17,X18) | 0 = X18 | 0 = sK0 | X17 = X18 | alength = sK0 | $less(sK0,sK0) | $less(X18,alength) (4:24) [resolution 397,254]
[SA] new: 4013. $less(0,X19) | $less(X19,X20) | 0 = X20 | 0 = sK0 | X19 = X20 | sK0 = sK0 | 0 = sK0 | $less(0,sK0) | $less(X20,sK0) (4:27) [resolution 397,333]
[SA] new: 4014. $less(0,X21) | $less(X21,X22) | 0 = X22 | 0 = sK0 | X21 = X22 | i(l8(sK2(sK0))) = sK0 | alength = sK0 | $less(alength,sK0) | $less(X22,0) (4:30) [resolution 397,2950]
[SA] new: 4015. $less(0,X23) | $less(X23,alength) | alength = 0 | 0 = sK0 | alength = X23 | alength = sK0 (6:18) [resolution 397,2704]
[SA] new: 4016. $less(0,sK0) | $less(0,sK0) | 0 = 0 | 0 = sK0 | 0 = sK0 (4:15) [resolution 397,46]
[SA] new: 4017. $less(0,X0) | $less(X0,sK0) | 0 = X0 | 0 = sK0 | X0 = X0 (4:15) [resolution 397,13]
[SA] new: 4018. $less(0,X1) | $less(X2,sK0) | 0 = X2 | 0 = sK0 | X1 = X2 | ~$less(X3,X1) | $less(X3,X2) (4:21) [resolution 397,14]
[SA] new: 4019. $less(0,X4) | $less(X5,sK0) | 0 = X5 | 0 = sK0 | X4 = X5 | X5 = X6 | $less(X6,X5) | $less(X4,X6) (4:24) [resolution 397,99]
[SA] new: 4020. $less(0,X7) | $less(X8,sK0) | 0 = X8 | 0 = sK0 | X7 = X8 | sK0 = X8 | $less(sK0,X8) | $less(X7,alength) (4:24) [resolution 397,205]
[SA] new: 4021. $less(0,X9) | $less(X10,sK0) | 0 = X10 | 0 = sK0 | X9 = X10 | alength = X10 | $less(sK0,X10) | $less(X9,alength) (4:24) [resolution 397,254]
[SA] new: 4022. $less(0,X11) | $less(X12,sK0) | 0 = X12 | 0 = sK0 | X11 = X12 | sK0 = X12 | 0 = sK0 | $less(0,X12) | $less(X11,sK0) (4:27) [resolution 397,333]
[SA] new: 4023. $less(0,X13) | $less(X14,sK0) | 0 = X14 | 0 = sK0 | X13 = X14 | i(l8(sK2(X14))) = X14 | alength = X14 | $less(alength,X14) | $less(X13,0) (4:30) [resolution 397,2950]
[SA] new: 4024. $less(0,X15) | $less(alength,sK0) | alength = 0 | 0 = sK0 | alength = X15 | $less(X15,0) | a(main_end,X16) = a(l11(sK2(X15)),X16) | i(l8(sK1(X16,sK2(X15)))) = X16 (4:35) [resolution 397,2770]
[SA] new: 4025. $less(0,X17) | $less(alength,sK0) | alength = 0 | 0 = sK0 | alength = X17 | alength = X18 | $less(X18,0) | i(l8(sK2(X18))) = X18 | $less(X17,X18) (4:30) [resolution 397,2660]
[SA] new: 4026. $less(0,X19) | $less(alength,sK0) | alength = 0 | 0 = sK0 | alength = X19 | $less(X19,0) | i(l8(sK2(X19))) = X19 (4:24) [resolution 397,2480]
[SA] new: 4027. $less(0,X20) | $less(i(l8(p(nl8))),sK0) | i(l8(p(nl8))) = 0 | 0 = sK0 | i(l8(p(nl8))) = X20 | zero = nl8 | $less(X20,alength) (4:30) [resolution 397,194]
[SA] new: 4028. $less(0,X21) | $less(i(l8(sK2(X22))),sK0) | i(l8(sK2(X22))) = 0 | 0 = sK0 | i(l8(sK2(X22))) = X21 | $less(X22,0) | ~$less(X22,alength) | $less(X21,alength) (4:33) [resolution 397,2920]
[SA] new: 4029. $less(0,X23) | $less(0,sK0) | 0 = 0 | 0 = sK0 | 0 = X23 | sK0 = X24 | 0 = sK0 | $less(X24,sK0) | $less(X23,X24) (4:27) [resolution 397,324]
[SA] new: 4030. $less(0,X25) | $less(0,sK0) | 0 = 0 | 0 = sK0 | 0 = X25 | alength = 0 | $less(X25,alength) (5:21) [resolution 397,262]
[SA] new: 4031. $less(0,X26) | $less(0,sK0) | 0 = 0 | 0 = sK0 | 0 = X26 | 0 = sK0 | $less(X26,alength) (4:21) [resolution 397,116]
[SA] new: 4032. $less(0,X27) | $less(0,sK0) | 0 = 0 | 0 = sK0 | 0 = X27 | 0 = sK0 | $less(X27,sK0) (4:21) [resolution 397,114]
[SA] new: 4033. $less(0,X28) | $less(sK0,sK0) | 0 = sK0 | 0 = sK0 | sK0 = X28 | alength = X29 | $less(X29,alength) | $less(X28,X29) (4:24) [resolution 397,247]
[SA] new: 4034. $less(0,X30) | $less(sK0,sK0) | 0 = sK0 | 0 = sK0 | sK0 = X30 | sK0 = X31 | $less(X31,alength) | $less(X30,X31) (4:24) [resolution 397,199]
[SA] new: 4035. $less(0,X32) | $less(sK0,sK0) | 0 = sK0 | 0 = sK0 | sK0 = X32 | $less(X32,alength) (4:18) [resolution 397,96]
[SA] new: 4036. $less(0,alength) | $less(sK0,sK0) | 0 = sK0 | 0 = sK0 | alength = sK0 | alength = sK0 (6:18) [resolution 397,2704]
[SA] new: 4037. $less(0,sK0) | $less(sK0,0) | 0 = 0 | 0 = sK0 | 0 = sK0 (4:15) [factoring 397]
[SA] new: 4038. $less(0,0) | $less(0,sK0) | 0 = 0 | 0 = sK0 | 0 = 0 (4:15) [factoring 397]
[SA] new: 4039. $less(0,sK0) | $less(sK0,sK0) | 0 = sK0 | 0 = sK0 | sK0 = sK0 (4:15) [factoring 397]
[SA] new: 4040. $less(0,sK0) | $less(sK0,sK0) | 0 = sK0 | sK0 = sK0 (4:12) [duplicate literal removal 4039]
[SA] new: 4041. $less(0,0) | $less(0,sK0) | 0 = 0 | 0 = sK0 (4:12) [duplicate literal removal 4038]
[SA] new: 4042. $less(0,sK0) | $less(sK0,0) | 0 = 0 | 0 = sK0 (4:12) [duplicate literal removal 4037]
[SA] new: 4043. $less(0,alength) | $less(sK0,sK0) | 0 = sK0 | alength = sK0 (6:12) [duplicate literal removal 4036]
[SA] new: 4044. $less(0,X32) | $less(sK0,sK0) | 0 = sK0 | sK0 = X32 | $less(X32,alength) (4:15) [duplicate literal removal 4035]
[SA] new: 4045. $less(0,X30) | $less(sK0,sK0) | 0 = sK0 | sK0 = X30 | sK0 = X31 | $less(X31,alength) | $less(X30,X31) (4:21) [duplicate literal removal 4034]
[SA] new: 4046. $less(0,X28) | $less(sK0,sK0) | 0 = sK0 | sK0 = X28 | alength = X29 | $less(X29,alength) | $less(X28,X29) (4:21) [duplicate literal removal 4033]
[SA] new: 4047. $less(0,X27) | $less(0,sK0) | 0 = 0 | 0 = sK0 | 0 = X27 | $less(X27,sK0) (4:18) [duplicate literal removal 4032]
[SA] new: 4048. $less(0,X26) | $less(0,sK0) | 0 = 0 | 0 = sK0 | 0 = X26 | $less(X26,alength) (4:18) [duplicate literal removal 4031]
[SA] new: 4049. $less(0,X23) | $less(0,sK0) | 0 = 0 | 0 = sK0 | 0 = X23 | sK0 = X24 | $less(X24,sK0) | $less(X23,X24) (4:24) [duplicate literal removal 4029]
[SA] new: 4050. $less(0,X11) | $less(X12,sK0) | 0 = X12 | 0 = sK0 | X11 = X12 | sK0 = X12 | $less(0,X12) | $less(X11,sK0) (4:24) [duplicate literal removal 4022]
[SA] new: 4051. $less(0,sK0) | 0 = 0 | 0 = sK0 (4:9) [duplicate literal removal 4016]
[SA] new: 4052. $less(0,X19) | $less(X19,X20) | 0 = X20 | 0 = sK0 | X19 = X20 | sK0 = sK0 | $less(0,sK0) | $less(X20,sK0) (4:24) [duplicate literal removal 4013]
[SA] new: 4053. $less(0,X8) | $less(X8,sK0) | 0 = sK0 | sK0 = X8 (4:12) [duplicate literal removal 4008]
[SA] new: 4054. $less(X27,sK0) | $less(0,X27) | 0 = X27 | 0 = sK0 | $less(0,sK0) (4:15) [duplicate literal removal 4001]
[SA] new: 4055. $less(X26,sK0) | $less(0,X26) | 0 = X26 | 0 = sK0 | $less(0,alength) (4:15) [duplicate literal removal 4000]
[SA] new: 4056. $less(X25,sK0) | $less(0,X25) | 0 = X25 | 0 = sK0 | alength = 0 | $less(0,alength) (5:18) [duplicate literal removal 3999]
[SA] new: 4057. $less(X23,sK0) | $less(0,X23) | 0 = X23 | 0 = sK0 | sK0 = X24 | $less(X24,sK0) | $less(0,X24) (4:21) [duplicate literal removal 3998]
[SA] new: 4058. $less(X19,sK0) | $less(alength,X19) | 0 = X19 | 0 = sK0 | alength = X19 | i(l8(sK2(0))) = 0 (4:21) [evaluation 3995]
[SA] new: 4059. $less(X15,sK0) | $less(alength,X15) | 0 = X15 | 0 = sK0 | alength = X15 | a(main_end,X16) = a(l11(sK2(0)),X16) | i(l8(sK1(X16,sK2(0)))) = X16 (4:32) [evaluation 3993]
[SA] new: 4060. $less(X13,sK0) | $less(X14,X13) | 0 = X13 | 0 = sK0 | X13 = X14 | i(l8(sK2(X14))) = X14 | alength = X14 | $less(alength,X14) (4:27) [evaluation 3992]
[SA] new: 4061. $less(X11,sK0) | $less(X12,X11) | 0 = X11 | 0 = sK0 | X11 = X12 | sK0 = X12 | $less(0,X12) | $less(0,sK0) (4:24) [duplicate literal removal 3991]
[SA] new: 4062. $less(X0,sK0) | $less(0,X0) | 0 = X0 | 0 = sK0 (4:12) [duplicate literal removal 3986]
[SA] passive: 3987. $less(X1,sK0) | $less(X2,X1) | 0 = X1 | 0 = sK0 | X1 = X2 | ~$less(X3,0) | $less(X3,X2) (4:21) [resolution 397,14]
[SA] passive: 3988. $less(X4,sK0) | $less(X5,X4) | 0 = X4 | 0 = sK0 | X4 = X5 | X5 = X6 | $less(X6,X5) | $less(0,X6) (4:24) [resolution 397,99]
[SA] passive: 4005. $less(0,X0) | $less(X0,X1) | 0 = X1 | 0 = sK0 | X0 = X1 | alength = X2 | $less(X2,alength) | $less(X1,X2) (4:24) [resolution 397,247]
[SA] passive: 4006. $less(0,X3) | $less(X3,X4) | 0 = X4 | 0 = sK0 | X3 = X4 | sK0 = X5 | $less(X5,alength) | $less(X4,X5) (4:24) [resolution 397,199]
[SA] passive: 4009. $less(0,X9) | $less(X9,X10) | 0 = X10 | 0 = sK0 | X9 = X10 | ~$less(X11,X10) | $less(X11,sK0) (4:21) [resolution 397,14]
[SA] passive: 4018. $less(0,X1) | $less(X2,sK0) | 0 = X2 | 0 = sK0 | X1 = X2 | ~$less(X3,X1) | $less(X3,X2) (4:21) [resolution 397,14]
[SA] passive: 4019. $less(0,X4) | $less(X5,sK0) | 0 = X5 | 0 = sK0 | X4 = X5 | X5 = X6 | $less(X6,X5) | $less(X4,X6) (4:24) [resolution 397,99]
[SA] passive: 4021. $less(0,X9) | $less(X10,sK0) | 0 = X10 | 0 = sK0 | X9 = X10 | alength = X10 | $less(sK0,X10) | $less(X9,alength) (4:24) [resolution 397,254]
[SA] passive: 4027. $less(0,X20) | $less(i(l8(p(nl8))),sK0) | i(l8(p(nl8))) = 0 | 0 = sK0 | i(l8(p(nl8))) = X20 | zero = nl8 | $less(X20,alength) (4:30) [resolution 397,194]
[SA] passive: 4028. $less(0,X21) | $less(i(l8(sK2(X22))),sK0) | i(l8(sK2(X22))) = 0 | 0 = sK0 | i(l8(sK2(X22))) = X21 | $less(X22,0) | ~$less(X22,alength) | $less(X21,alength) (4:33) [resolution 397,2920]
[SA] active: 1970. alength != 1 | s(zero) = nl8 | 0 = sK0 (7:10:1) [superposition 153,1845]
[SA] active: 407. $less(0,X17) | $less(X18,sK0) | $less(X17,X18) | 0 = X17 | 0 = sK0 | 0 = X18 (3:18:3) [resolution 174,114]
[SA] new: 4063. $less(X0,sK0) | $less(0,X0) | 0 = 0 | 0 = sK0 | 0 = X0 (4:15) [resolution 407,13]
[SA] new: 4064. $less(X1,sK0) | $less(X2,X1) | 0 = X2 | 0 = sK0 | 0 = X1 | ~$less(X3,0) | $less(X3,X2) (4:21) [resolution 407,14]
[SA] new: 4065. $less(X4,sK0) | $less(X5,X4) | 0 = X5 | 0 = sK0 | 0 = X4 | X5 = X6 | $less(X6,X5) | $less(0,X6) (4:24) [resolution 407,99]
[SA] new: 4066. $less(X7,sK0) | $less(X8,X7) | 0 = X8 | 0 = sK0 | 0 = X7 | sK0 = X8 | $less(sK0,X8) | $less(0,alength) (4:24) [resolution 407,205]
[SA] new: 4067. $less(X9,sK0) | $less(X10,X9) | 0 = X10 | 0 = sK0 | 0 = X9 | alength = X10 | $less(sK0,X10) | $less(0,alength) (4:24) [resolution 407,254]
[SA] new: 4068. $less(X11,sK0) | $less(X12,X11) | 0 = X12 | 0 = sK0 | 0 = X11 | sK0 = X12 | 0 = sK0 | $less(0,X12) | $less(0,sK0) (4:27) [resolution 407,333]
[SA] new: 4069. $less(X13,sK0) | $less(X14,X13) | 0 = X14 | 0 = sK0 | 0 = X13 | i(l8(sK2(X14))) = X14 | alength = X14 | $less(alength,X14) | $less(0,0) (4:30) [resolution 407,2950]
[SA] new: 4070. $less(X15,sK0) | $less(alength,X15) | alength = 0 | 0 = sK0 | 0 = X15 | $less(0,0) | a(main_end,X16) = a(l11(sK2(0)),X16) | i(l8(sK1(X16,sK2(0)))) = X16 (4:35) [resolution 407,2770]
[SA] new: 4071. $less(X17,sK0) | $less(alength,X17) | alength = 0 | 0 = sK0 | 0 = X17 | alength = X18 | $less(X18,0) | i(l8(sK2(X18))) = X18 | $less(0,X18) (4:30) [resolution 407,2660]
[SA] new: 4072. $less(X19,sK0) | $less(alength,X19) | alength = 0 | 0 = sK0 | 0 = X19 | $less(0,0) | i(l8(sK2(0))) = 0 (4:24) [resolution 407,2480]
[SA] new: 4073. $less(X20,sK0) | $less(i(l8(p(nl8))),X20) | i(l8(p(nl8))) = 0 | 0 = sK0 | 0 = X20 | zero = nl8 | $less(0,alength) (4:27) [resolution 407,194]
[SA] new: 4074. $less(X21,sK0) | $less(i(l8(sK2(X22))),X21) | i(l8(sK2(X22))) = 0 | 0 = sK0 | 0 = X21 | $less(X22,0) | ~$less(X22,alength) | $less(0,alength) (4:30) [resolution 407,2920]
[SA] new: 4075. $less(X23,sK0) | $less(0,X23) | 0 = 0 | 0 = sK0 | 0 = X23 | sK0 = X24 | 0 = sK0 | $less(X24,sK0) | $less(0,X24) (4:27) [resolution 407,324]
[SA] new: 4076. $less(X25,sK0) | $less(0,X25) | 0 = 0 | 0 = sK0 | 0 = X25 | alength = 0 | $less(0,alength) (5:21) [resolution 407,262]
[SA] new: 4077. $less(X26,sK0) | $less(0,X26) | 0 = 0 | 0 = sK0 | 0 = X26 | 0 = sK0 | $less(0,alength) (4:21) [resolution 407,116]
[SA] new: 4078. $less(X27,sK0) | $less(0,X27) | 0 = 0 | 0 = sK0 | 0 = X27 | 0 = sK0 | $less(0,sK0) (4:21) [resolution 407,114]
[SA] new: 4079. $less(X28,sK0) | $less(sK0,X28) | 0 = sK0 | 0 = sK0 | 0 = X28 | alength = X29 | $less(X29,alength) | $less(0,X29) (4:24) [resolution 407,247]
[SA] new: 4080. $less(X30,sK0) | $less(sK0,X30) | 0 = sK0 | 0 = sK0 | 0 = X30 | sK0 = X31 | $less(X31,alength) | $less(0,X31) (4:24) [resolution 407,199]
[SA] new: 4081. $less(X32,sK0) | $less(sK0,X32) | 0 = sK0 | 0 = sK0 | 0 = X32 | $less(0,alength) (4:18) [resolution 407,96]
[SA] new: 4082. $less(0,X0) | $less(X0,X1) | 0 = X0 | 0 = sK0 | 0 = X1 | alength = X2 | $less(X2,alength) | $less(X1,X2) (4:24) [resolution 407,247]
[SA] new: 4083. $less(0,X3) | $less(X3,X4) | 0 = X3 | 0 = sK0 | 0 = X4 | sK0 = X5 | $less(X5,alength) | $less(X4,X5) (4:24) [resolution 407,199]
[SA] new: 4084. $less(0,X6) | $less(X6,X7) | 0 = X6 | 0 = sK0 | 0 = X7 | $less(X7,alength) (4:18) [resolution 407,96]
[SA] new: 4085. $less(0,X8) | $less(X8,sK0) | 0 = X8 | 0 = sK0 | 0 = sK0 (4:15) [resolution 407,13]
[SA] new: 4086. $less(0,X9) | $less(X9,X10) | 0 = X9 | 0 = sK0 | 0 = X10 | ~$less(X11,X10) | $less(X11,sK0) (4:21) [resolution 407,14]
[SA] new: 4087. $less(0,X12) | $less(X12,X13) | 0 = X12 | 0 = sK0 | 0 = X13 | sK0 = X14 | $less(X14,sK0) | $less(X13,X14) (4:24) [resolution 407,99]
[SA] new: 4088. $less(0,X15) | $less(X15,X16) | 0 = X15 | 0 = sK0 | 0 = X16 | sK0 = sK0 | $less(sK0,sK0) | $less(X16,alength) (4:24) [resolution 407,205]
[SA] new: 4089. $less(0,X17) | $less(X17,X18) | 0 = X17 | 0 = sK0 | 0 = X18 | alength = sK0 | $less(sK0,sK0) | $less(X18,alength) (4:24) [resolution 407,254]
[SA] new: 4090. $less(0,X19) | $less(X19,X20) | 0 = X19 | 0 = sK0 | 0 = X20 | sK0 = sK0 | 0 = sK0 | $less(0,sK0) | $less(X20,sK0) (4:27) [resolution 407,333]
[SA] new: 4091. $less(0,X21) | $less(X21,X22) | 0 = X21 | 0 = sK0 | 0 = X22 | i(l8(sK2(sK0))) = sK0 | alength = sK0 | $less(alength,sK0) | $less(X22,0) (4:30) [resolution 407,2950]
[SA] new: 4092. $less(0,X23) | $less(X23,alength) | 0 = X23 | 0 = sK0 | alength = 0 | alength = sK0 (6:18) [resolution 407,2704]
[SA] new: 4093. $less(0,sK0) | $less(0,sK0) | 0 = sK0 | 0 = sK0 | 0 = 0 (4:15) [resolution 407,46]
[SA] new: 4094. $less(0,X0) | $less(X0,sK0) | 0 = X0 | 0 = sK0 | 0 = X0 (4:15) [resolution 407,13]
[SA] new: 4095. $less(0,X1) | $less(X2,sK0) | 0 = X1 | 0 = sK0 | 0 = X2 | ~$less(X3,X1) | $less(X3,X2) (4:21) [resolution 407,14]
[SA] new: 4096. $less(0,X4) | $less(X5,sK0) | 0 = X4 | 0 = sK0 | 0 = X5 | X5 = X6 | $less(X6,X5) | $less(X4,X6) (4:24) [resolution 407,99]
[SA] new: 4097. $less(0,X7) | $less(X8,sK0) | 0 = X7 | 0 = sK0 | 0 = X8 | sK0 = X8 | $less(sK0,X8) | $less(X7,alength) (4:24) [resolution 407,205]
[SA] new: 4098. $less(0,X9) | $less(X10,sK0) | 0 = X9 | 0 = sK0 | 0 = X10 | alength = X10 | $less(sK0,X10) | $less(X9,alength) (4:24) [resolution 407,254]
[SA] new: 4099. $less(0,X11) | $less(X12,sK0) | 0 = X11 | 0 = sK0 | 0 = X12 | sK0 = X12 | 0 = sK0 | $less(0,X12) | $less(X11,sK0) (4:27) [resolution 407,333]
[SA] new: 4100. $less(0,X13) | $less(X14,sK0) | 0 = X13 | 0 = sK0 | 0 = X14 | i(l8(sK2(X14))) = X14 | alength = X14 | $less(alength,X14) | $less(X13,0) (4:30) [resolution 407,2950]
[SA] new: 4101. $less(0,X15) | $less(alength,sK0) | 0 = X15 | 0 = sK0 | alength = 0 | $less(X15,0) | a(main_end,X16) = a(l11(sK2(X15)),X16) | i(l8(sK1(X16,sK2(X15)))) = X16 (4:35) [resolution 407,2770]
[SA] new: 4102. $less(0,X17) | $less(alength,sK0) | 0 = X17 | 0 = sK0 | alength = 0 | alength = X18 | $less(X18,0) | i(l8(sK2(X18))) = X18 | $less(X17,X18) (4:30) [resolution 407,2660]
[SA] new: 4103. $less(0,X19) | $less(alength,sK0) | 0 = X19 | 0 = sK0 | alength = 0 | $less(X19,0) | i(l8(sK2(X19))) = X19 (4:24) [resolution 407,2480]
[SA] new: 4104. $less(0,X20) | $less(i(l8(p(nl8))),sK0) | 0 = X20 | 0 = sK0 | i(l8(p(nl8))) = 0 | zero = nl8 | $less(X20,alength) (4:27) [resolution 407,194]
[SA] new: 4105. $less(0,X21) | $less(i(l8(sK2(X22))),sK0) | 0 = X21 | 0 = sK0 | i(l8(sK2(X22))) = 0 | $less(X22,0) | ~$less(X22,alength) | $less(X21,alength) (4:30) [resolution 407,2920]
[SA] new: 4106. $less(0,X23) | $less(0,sK0) | 0 = X23 | 0 = sK0 | 0 = 0 | sK0 = X24 | 0 = sK0 | $less(X24,sK0) | $less(X23,X24) (4:27) [resolution 407,324]
[SA] new: 4107. $less(0,X25) | $less(0,sK0) | 0 = X25 | 0 = sK0 | 0 = 0 | alength = 0 | $less(X25,alength) (5:21) [resolution 407,262]
[SA] new: 4108. $less(0,X26) | $less(0,sK0) | 0 = X26 | 0 = sK0 | 0 = 0 | 0 = sK0 | $less(X26,alength) (4:21) [resolution 407,116]
[SA] new: 4109. $less(0,X27) | $less(0,sK0) | 0 = X27 | 0 = sK0 | 0 = 0 | 0 = sK0 | $less(X27,sK0) (4:21) [resolution 407,114]
[SA] new: 4110. $less(0,X28) | $less(sK0,sK0) | 0 = X28 | 0 = sK0 | 0 = sK0 | alength = X29 | $less(X29,alength) | $less(X28,X29) (4:24) [resolution 407,247]
[SA] new: 4111. $less(0,X30) | $less(sK0,sK0) | 0 = X30 | 0 = sK0 | 0 = sK0 | sK0 = X31 | $less(X31,alength) | $less(X30,X31) (4:24) [resolution 407,199]
[SA] new: 4112. $less(0,X32) | $less(sK0,sK0) | 0 = X32 | 0 = sK0 | 0 = sK0 | $less(X32,alength) (4:18) [resolution 407,96]
[SA] new: 4113. $less(0,alength) | $less(sK0,sK0) | alength = 0 | 0 = sK0 | 0 = sK0 | alength = sK0 (6:18) [resolution 407,2704]
[SA] new: 4114. $less(0,sK0) | $less(sK0,0) | 0 = sK0 | 0 = sK0 | 0 = 0 (4:15) [factoring 407]
[SA] new: 4115. $less(0,0) | $less(0,sK0) | 0 = 0 | 0 = sK0 | 0 = 0 (4:15) [factoring 407]
[SA] new: 4116. $less(0,sK0) | $less(sK0,sK0) | 0 = sK0 | 0 = sK0 | 0 = sK0 (4:15) [factoring 407]
[SA] new: 4117. $less(0,sK0) | $less(sK0,sK0) | 0 = sK0 (4:9) [duplicate literal removal 4116]
[SA] new: 4118. $less(0,0) | $less(0,sK0) | 0 = 0 | 0 = sK0 (4:12) [duplicate literal removal 4115]
[SA] new: 4119. $less(0,sK0) | $less(sK0,0) | 0 = sK0 | 0 = 0 (4:12) [duplicate literal removal 4114]
[SA] new: 4120. $less(0,alength) | $less(sK0,sK0) | alength = 0 | 0 = sK0 | alength = sK0 (6:15) [duplicate literal removal 4113]
[SA] new: 4121. $less(0,X32) | $less(sK0,sK0) | 0 = X32 | 0 = sK0 | $less(X32,alength) (4:15) [duplicate literal removal 4112]
[SA] new: 4122. $less(0,X30) | $less(sK0,sK0) | 0 = X30 | 0 = sK0 | sK0 = X31 | $less(X31,alength) | $less(X30,X31) (4:21) [duplicate literal removal 4111]
[SA] new: 4123. $less(0,X28) | $less(sK0,sK0) | 0 = X28 | 0 = sK0 | alength = X29 | $less(X29,alength) | $less(X28,X29) (4:21) [duplicate literal removal 4110]
[SA] new: 4124. $less(0,X27) | $less(0,sK0) | 0 = X27 | 0 = sK0 | 0 = 0 | $less(X27,sK0) (4:18) [duplicate literal removal 4109]
[SA] new: 4125. $less(0,X26) | $less(0,sK0) | 0 = X26 | 0 = sK0 | 0 = 0 | $less(X26,alength) (4:18) [duplicate literal removal 4108]
[SA] new: 4126. $less(0,X23) | $less(0,sK0) | 0 = X23 | 0 = sK0 | 0 = 0 | sK0 = X24 | $less(X24,sK0) | $less(X23,X24) (4:24) [duplicate literal removal 4106]
[SA] new: 4127. $less(0,X11) | $less(X12,sK0) | 0 = X11 | 0 = sK0 | 0 = X12 | sK0 = X12 | $less(0,X12) | $less(X11,sK0) (4:24) [duplicate literal removal 4099]
[SA] new: 4128. $less(0,X0) | $less(X0,sK0) | 0 = X0 | 0 = sK0 (4:12) [duplicate literal removal 4094]
[SA] new: 4129. $less(0,sK0) | 0 = sK0 | 0 = 0 (4:9) [duplicate literal removal 4093]
[SA] new: 4130. $less(0,X19) | $less(X19,X20) | 0 = X19 | 0 = sK0 | 0 = X20 | sK0 = sK0 | $less(0,sK0) | $less(X20,sK0) (4:24) [duplicate literal removal 4090]
[SA] new: 4131. $less(0,X8) | $less(X8,sK0) | 0 = X8 | 0 = sK0 (4:12) [duplicate literal removal 4085]
[SA] new: 4132. $less(X32,sK0) | $less(sK0,X32) | 0 = sK0 | 0 = X32 | $less(0,alength) (4:15) [duplicate literal removal 4081]
[SA] new: 4133. $less(X30,sK0) | $less(sK0,X30) | 0 = sK0 | 0 = X30 | sK0 = X31 | $less(X31,alength) | $less(0,X31) (4:21) [duplicate literal removal 4080]
[SA] new: 4134. $less(X28,sK0) | $less(sK0,X28) | 0 = sK0 | 0 = X28 | alength = X29 | $less(X29,alength) | $less(0,X29) (4:21) [duplicate literal removal 4079]
[SA] new: 4135. $less(X27,sK0) | $less(0,X27) | 0 = 0 | 0 = sK0 | 0 = X27 | $less(0,sK0) (4:18) [duplicate literal removal 4078]
[SA] new: 4136. $less(X26,sK0) | $less(0,X26) | 0 = 0 | 0 = sK0 | 0 = X26 | $less(0,alength) (4:18) [duplicate literal removal 4077]
[SA] new: 4137. $less(X23,sK0) | $less(0,X23) | 0 = 0 | 0 = sK0 | 0 = X23 | sK0 = X24 | $less(X24,sK0) | $less(0,X24) (4:24) [duplicate literal removal 4075]
[SA] new: 4138. $less(X19,sK0) | $less(alength,X19) | alength = 0 | 0 = sK0 | 0 = X19 | i(l8(sK2(0))) = 0 (4:21) [evaluation 4072]
[SA] new: 4139. $less(X15,sK0) | $less(alength,X15) | alength = 0 | 0 = sK0 | 0 = X15 | a(main_end,X16) = a(l11(sK2(0)),X16) | i(l8(sK1(X16,sK2(0)))) = X16 (4:32) [evaluation 4070]
[SA] new: 4140. $less(X13,sK0) | $less(X14,X13) | 0 = X14 | 0 = sK0 | 0 = X13 | i(l8(sK2(X14))) = X14 | alength = X14 | $less(alength,X14) (4:27) [evaluation 4069]
[SA] new: 4141. $less(X11,sK0) | $less(X12,X11) | 0 = X12 | 0 = sK0 | 0 = X11 | sK0 = X12 | $less(0,X12) | $less(0,sK0) (4:24) [duplicate literal removal 4068]
[SA] passive: 4064. $less(X1,sK0) | $less(X2,X1) | 0 = X2 | 0 = sK0 | 0 = X1 | ~$less(X3,0) | $less(X3,X2) (4:21) [resolution 407,14]
[SA] passive: 4065. $less(X4,sK0) | $less(X5,X4) | 0 = X5 | 0 = sK0 | 0 = X4 | X5 = X6 | $less(X6,X5) | $less(0,X6) (4:24) [resolution 407,99]
[SA] passive: 4082. $less(0,X0) | $less(X0,X1) | 0 = X0 | 0 = sK0 | 0 = X1 | alength = X2 | $less(X2,alength) | $less(X1,X2) (4:24) [resolution 407,247]
[SA] passive: 4083. $less(0,X3) | $less(X3,X4) | 0 = X3 | 0 = sK0 | 0 = X4 | sK0 = X5 | $less(X5,alength) | $less(X4,X5) (4:24) [resolution 407,199]
[SA] passive: 4086. $less(0,X9) | $less(X9,X10) | 0 = X9 | 0 = sK0 | 0 = X10 | ~$less(X11,X10) | $less(X11,sK0) (4:21) [resolution 407,14]
[SA] passive: 4095. $less(0,X1) | $less(X2,sK0) | 0 = X1 | 0 = sK0 | 0 = X2 | ~$less(X3,X1) | $less(X3,X2) (4:21) [resolution 407,14]
[SA] passive: 4096. $less(0,X4) | $less(X5,sK0) | 0 = X4 | 0 = sK0 | 0 = X5 | X5 = X6 | $less(X6,X5) | $less(X4,X6) (4:24) [resolution 407,99]
[SA] active: 1989. v = a(l8(s(nl8)),alength) | 0 = sK0 (7:10:1) [superposition 354,1845]
[SA] new: 4142. v = a(l8(nl8),alength) | 0 = sK0 | alength = i(l8(nl8)) (8:14) [superposition 1989,440]
[SA] new: 4143. v = a(l11(nl8),alength) | 0 = sK0 (8:9) [superposition 1989,57]
[SA] new: 4144. v = a(l8(nl8),alength) | alength = i(l8(nl8)) | 0 = sK0 (8:14) [superposition 440,1989]
[SA] new: 4145. v = a(l11(nl8),alength) | 0 = sK0 (8:9) [superposition 57,1989]
[SA] active: 2547. $less(sK0,X15) | $less(X15,0) | alength = X15 | sK0 = X15 | i(l8(sK2(X15))) = X15 (3:18:2) [backward demodulation 2476,541]
[SA] new: 4146. $less(0,0) | alength = 0 | 0 = sK0 | i(l8(sK2(0))) = 0 (4:15) [resolution 2547,46]
[SA] new: 4147. $less(sK0,0) | alength = sK0 | sK0 = sK0 | i(l8(sK2(sK0))) = sK0 (4:15) [resolution 2547,13]
[SA] new: 4148. $less(X0,0) | alength = X0 | sK0 = X0 | i(l8(sK2(X0))) = X0 | ~$less(X1,sK0) | $less(X1,X0) (4:21) [resolution 2547,14]
[SA] new: 4149. $less(X2,0) | alength = X2 | sK0 = X2 | i(l8(sK2(X2))) = X2 | X2 = X3 | $less(X3,X2) | $less(sK0,X3) (4:24) [resolution 2547,99]
[SA] new: 4150. $less(X4,0) | alength = X4 | sK0 = X4 | i(l8(sK2(X4))) = X4 | sK0 = X4 | $less(sK0,X4) | $less(sK0,alength) (4:24) [resolution 2547,205]
[SA] new: 4151. $less(X5,0) | alength = X5 | sK0 = X5 | i(l8(sK2(X5))) = X5 | alength = X5 | $less(sK0,X5) | $less(sK0,alength) (4:24) [resolution 2547,254]
[SA] new: 4152. $less(X6,0) | alength = X6 | sK0 = X6 | i(l8(sK2(X6))) = X6 | sK0 = X6 | 0 = sK0 | $less(0,X6) | $less(sK0,sK0) (4:27) [resolution 2547,333]
[SA] new: 4153. $less(X7,0) | alength = X7 | sK0 = X7 | i(l8(sK2(X7))) = X7 | i(l8(sK2(X7))) = X7 | alength = X7 | $less(alength,X7) | $less(sK0,0) (4:30) [resolution 2547,2950]
[SA] new: 4154. $less(alength,0) | alength = alength | alength = sK0 | alength = i(l8(sK2(alength))) | $less(sK0,0) | a(main_end,X8) = a(l11(sK2(sK0)),X8) | i(l8(sK1(X8,sK2(sK0)))) = X8 (4:35) [resolution 2547,2770]
[SA] new: 4155. $less(alength,0) | alength = alength | alength = sK0 | alength = i(l8(sK2(alength))) | alength = X9 | $less(X9,0) | i(l8(sK2(X9))) = X9 | $less(sK0,X9) (4:30) [resolution 2547,2660]
[SA] new: 4156. $less(alength,0) | alength = alength | alength = sK0 | alength = i(l8(sK2(alength))) | $less(sK0,0) | i(l8(sK2(sK0))) = sK0 (4:24) [resolution 2547,2480]
[SA] new: 4157. $less(i(l8(p(nl8))),0) | alength = i(l8(p(nl8))) | i(l8(p(nl8))) = sK0 | i(l8(p(nl8))) = i(l8(sK2(i(l8(p(nl8)))))) | zero = nl8 | $less(sK0,alength) (4:36) [resolution 2547,194]
[SA] new: 4158. $less(i(l8(sK2(X10))),0) | alength = i(l8(sK2(X10))) | i(l8(sK2(X10))) = sK0 | i(l8(sK2(X10))) = i(l8(sK2(i(l8(sK2(X10)))))) | $less(X10,0) | ~$less(X10,alength) | $less(sK0,alength) (4:39) [resolution 2547,2920]
[SA] new: 4159. $less(0,0) | alength = 0 | 0 = sK0 | i(l8(sK2(0))) = 0 | sK0 = X11 | 0 = sK0 | $less(X11,sK0) | $less(sK0,X11) (4:27) [resolution 2547,324]
[SA] new: 4160. $less(0,0) | alength = 0 | 0 = sK0 | i(l8(sK2(0))) = 0 | alength = 0 | $less(sK0,alength) (5:21) [resolution 2547,262]
[SA] new: 4161. $less(0,0) | alength = 0 | 0 = sK0 | i(l8(sK2(0))) = 0 | 0 = sK0 | $less(sK0,alength) (4:21) [resolution 2547,116]
[SA] new: 4162. $less(0,0) | alength = 0 | 0 = sK0 | i(l8(sK2(0))) = 0 | 0 = sK0 | $less(sK0,sK0) (4:21) [resolution 2547,114]
[SA] new: 4163. $less(sK0,0) | alength = sK0 | sK0 = sK0 | i(l8(sK2(sK0))) = sK0 | alength = X12 | $less(X12,alength) | $less(sK0,X12) (4:24) [resolution 2547,247]
[SA] new: 4164. $less(sK0,0) | alength = sK0 | sK0 = sK0 | i(l8(sK2(sK0))) = sK0 | sK0 = X13 | $less(X13,alength) | $less(sK0,X13) (4:24) [resolution 2547,199]
[SA] new: 4165. $less(sK0,0) | alength = sK0 | sK0 = sK0 | i(l8(sK2(sK0))) = sK0 | $less(sK0,alength) (4:18) [resolution 2547,96]
[SA] new: 4166. $less(sK0,sK0) | alength = sK0 | sK0 = sK0 | i(l8(sK2(sK0))) = sK0 (4:15) [resolution 2547,46]
[SA] new: 4167. $less(sK0,X0) | alength = X0 | sK0 = X0 | i(l8(sK2(X0))) = X0 | sK0 = X1 | 0 = sK0 | $less(X1,sK0) | $less(X0,X1) (4:27) [resolution 2547,324]
[SA] new: 4168. $less(sK0,X2) | alength = X2 | sK0 = X2 | i(l8(sK2(X2))) = X2 | alength = 0 | $less(X2,alength) (5:21) [resolution 2547,262]
[SA] new: 4169. $less(sK0,X3) | alength = X3 | sK0 = X3 | i(l8(sK2(X3))) = X3 | 0 = sK0 | $less(X3,alength) (4:21) [resolution 2547,116]
[SA] new: 4170. $less(sK0,X4) | alength = X4 | sK0 = X4 | i(l8(sK2(X4))) = X4 | 0 = sK0 | $less(X4,sK0) (4:21) [resolution 2547,114]
[SA] new: 4171. $less(sK0,0) | alength = 0 | 0 = sK0 | i(l8(sK2(0))) = 0 (4:15) [resolution 2547,13]
[SA] new: 4172. $less(sK0,X5) | alength = X5 | sK0 = X5 | i(l8(sK2(X5))) = X5 | ~$less(X6,X5) | $less(X6,0) (4:21) [resolution 2547,14]
[SA] new: 4173. $less(sK0,X7) | alength = X7 | sK0 = X7 | i(l8(sK2(X7))) = X7 | 0 = X8 | $less(X8,0) | $less(X7,X8) (4:24) [resolution 2547,99]
[SA] new: 4174. $less(sK0,X9) | alength = X9 | sK0 = X9 | i(l8(sK2(X9))) = X9 | 0 = sK0 | $less(sK0,0) | $less(X9,alength) (4:24) [resolution 2547,205]
[SA] new: 4175. $less(sK0,X10) | alength = X10 | sK0 = X10 | i(l8(sK2(X10))) = X10 | alength = 0 | $less(sK0,0) | $less(X10,alength) (4:24) [resolution 2547,254]
[SA] new: 4176. $less(sK0,X11) | alength = X11 | sK0 = X11 | i(l8(sK2(X11))) = X11 | 0 = sK0 | 0 = sK0 | $less(0,0) | $less(X11,sK0) (4:27) [resolution 2547,333]
[SA] new: 4177. $less(sK0,X12) | alength = X12 | sK0 = X12 | i(l8(sK2(X12))) = X12 | i(l8(sK2(0))) = 0 | alength = 0 | $less(alength,0) | $less(X12,0) (4:30) [resolution 2547,2950]
[SA] new: 4178. $less(sK0,X11) | alength = X11 | sK0 = X11 | i(l8(sK2(X11))) = X11 | 0 = sK0 | $less(0,0) | $less(X11,sK0) (4:24) [duplicate literal removal 4176]
[SA] new: 4179. $less(sK0,X11) | alength = X11 | sK0 = X11 | i(l8(sK2(X11))) = X11 | 0 = sK0 | $less(X11,sK0) (4:21) [evaluation 4178]
[SA] new: 4180. $less(0,0) | alength = 0 | 0 = sK0 | i(l8(sK2(0))) = 0 | $less(sK0,sK0) (4:18) [duplicate literal removal 4162]
[SA] new: 4181. alength = 0 | 0 = sK0 | i(l8(sK2(0))) = 0 | $less(sK0,sK0) (4:15) [evaluation 4180]
[SA] new: 4182. $less(0,0) | alength = 0 | 0 = sK0 | i(l8(sK2(0))) = 0 | $less(sK0,alength) (4:18) [duplicate literal removal 4161]
[SA] new: 4183. alength = 0 | 0 = sK0 | i(l8(sK2(0))) = 0 | $less(sK0,alength) (4:15) [evaluation 4182]
[SA] new: 4184. $less(0,0) | alength = 0 | 0 = sK0 | i(l8(sK2(0))) = 0 | $less(sK0,alength) (5:18) [duplicate literal removal 4160]
[SA] new: 4185. alength = 0 | 0 = sK0 | i(l8(sK2(0))) = 0 | $less(sK0,alength) (5:15) [evaluation 4184]
[SA] new: 4186. $less(0,0) | alength = 0 | 0 = sK0 | i(l8(sK2(0))) = 0 | sK0 = X11 | $less(X11,sK0) | $less(sK0,X11) (4:24) [duplicate literal removal 4159]
[SA] new: 4187. alength = 0 | 0 = sK0 | i(l8(sK2(0))) = 0 | sK0 = X11 | $less(X11,sK0) | $less(sK0,X11) (4:21) [evaluation 4186]
[SA] new: 4188. $less(X7,0) | alength = X7 | sK0 = X7 | i(l8(sK2(X7))) = X7 | $less(alength,X7) | $less(sK0,0) (4:21) [duplicate literal removal 4153]
[SA] new: 4189. $less(X6,0) | alength = X6 | sK0 = X6 | i(l8(sK2(X6))) = X6 | 0 = sK0 | $less(0,X6) | $less(sK0,sK0) (4:24) [duplicate literal removal 4152]
[SA] new: 4190. $less(X5,0) | alength = X5 | sK0 = X5 | i(l8(sK2(X5))) = X5 | $less(sK0,X5) | $less(sK0,alength) (4:21) [duplicate literal removal 4151]
[SA] new: 4191. $less(X4,0) | alength = X4 | sK0 = X4 | i(l8(sK2(X4))) = X4 | $less(sK0,X4) | $less(sK0,alength) (4:21) [duplicate literal removal 4150]
[SA] new: 4192. alength = 0 | 0 = sK0 | i(l8(sK2(0))) = 0 (4:12) [evaluation 4146]
[SA] active: 2529. 1 != sK0 | alength = sK0 | s(zero) = sK2(sK0) (4:11:1) [backward demodulation 2476,472]
[SA] active: 2676. ~$less(X3,i(l8(sK2(X2)))) | alength = X2 | $less(X2,0) | $less(alength,X2) | $less(X3,alength) (3:18:1) [forward demodulation 2519,2476]
[SA] new: 4193. alength = X0 | $less(X0,0) | $less(alength,X0) | $less(X1,alength) | $less(i(l8(sK2(X0))),X1) | i(l8(sK2(X0))) = X1 (4:24) [resolution 2676,15]
[SA] new: 4194. alength = X2 | $less(X2,0) | $less(alength,X2) | $less(X3,alength) | $less(i(l8(sK2(X2))),X3) | i(l8(sK2(X2))) = X3 (4:24) [resolution 2676,15]
[SA] new: 4195. alength = X4 | $less(X4,0) | $less(alength,X4) | $less(X5,alength) | $less(X6,X5) | $less(i(l8(sK2(X4))),X6) | i(l8(sK2(X4))) = X5 | i(l8(sK2(X4))) = X6 (4:33) [resolution 2676,174]
[SA] new: 4196. alength = X7 | $less(X7,0) | $less(alength,X7) | $less(X8,alength) | $less(i(l8(sK2(X7))),X9) | $less(X9,X8) | X8 = X9 | i(l8(sK2(X7))) = X8 (4:30) [resolution 2676,174]
[SA] new: 4197. alength = X10 | $less(X10,0) | $less(alength,X10) | $less(X11,alength) | $less(X12,X11) | $less(i(l8(sK2(X10))),X12) | i(l8(sK2(X10))) = X12 | X11 = X12 (4:30) [resolution 2676,174]
[SA] new: 4198. alength = X13 | $less(X13,0) | $less(alength,X13) | $less(X14,alength) | $less(sK0,X14) | $less(i(l8(sK2(X13))),alength) | i(l8(sK2(X13))) = X14 | i(l8(sK2(X13))) = sK0 (4:33) [resolution 2676,198]
[SA] new: 4199. alength = X15 | $less(X15,0) | $less(alength,X15) | $less(X16,alength) | $less(sK0,X16) | $less(i(l8(sK2(X15))),alength) | alength = i(l8(sK2(X15))) | sK0 = X16 (4:30) [resolution 2676,204]
[SA] new: 4200. alength = X17 | $less(X17,0) | $less(alength,X17) | $less(X18,alength) | $less(sK0,X18) | $less(i(l8(sK2(X17))),alength) | i(l8(sK2(X17))) = X18 | alength = i(l8(sK2(X17))) (4:33) [resolution 2676,246]
[SA] new: 4201. alength = X19 | $less(X19,0) | $less(alength,X19) | $less(X20,alength) | $less(sK0,X20) | $less(i(l8(sK2(X19))),alength) | alength = i(l8(sK2(X19))) | alength = X20 (4:30) [resolution 2676,253]
[SA] new: 4202. alength = X21 | $less(X21,0) | $less(alength,X21) | $less(X22,alength) | $less(0,X22) | $less(i(l8(sK2(X21))),sK0) | i(l8(sK2(X21))) = X22 | 0 = sK0 | i(l8(sK2(X21))) = sK0 (4:36) [resolution 2676,323]
[SA] new: 4203. alength = X23 | $less(X23,0) | $less(alength,X23) | $less(X24,alength) | $less(0,X24) | $less(i(l8(sK2(X23))),sK0) | i(l8(sK2(X23))) = sK0 | 0 = sK0 | sK0 = X24 (4:33) [resolution 2676,332]
[SA] new: 4204. alength = X25 | $less(X25,0) | $less(alength,X25) | $less(X26,alength) | $less(0,X26) | $less(i(l8(sK2(X25))),sK0) | i(l8(sK2(X25))) = X26 | 0 = sK0 | 0 = X26 (4:33) [resolution 2676,387]
[SA] new: 4205. alength = X27 | $less(X27,0) | $less(alength,X27) | $less(X28,alength) | $less(sK0,X28) | $less(i(l8(sK2(X27))),alength) | i(l8(sK2(X27))) = X28 | sK0 = X28 (4:30) [resolution 2676,388]
[SA] new: 4206. alength = X29 | $less(X29,0) | $less(alength,X29) | $less(X30,alength) | $less(0,X30) | $less(i(l8(sK2(X29))),sK0) | i(l8(sK2(X29))) = 0 | 0 = sK0 | i(l8(sK2(X29))) = X30 (4:36) [resolution 2676,397]
[SA] new: 4207. alength = X31 | $less(X31,0) | $less(alength,X31) | $less(X32,alength) | $less(0,X32) | $less(i(l8(sK2(X31))),sK0) | 0 = X32 | 0 = sK0 | i(l8(sK2(X31))) = 0 (4:33) [resolution 2676,407]
[SA] new: 4208. alength = X33 | $less(X33,0) | $less(alength,X33) | $less(X34,alength) | $less(sK0,X34) | $less(i(l8(sK2(X33))),alength) | sK0 = X34 | i(l8(sK2(X33))) = sK0 (4:30) [resolution 2676,408]
[SA] new: 4209. alength = X35 | $less(X35,0) | $less(alength,X35) | $less(X36,alength) | $less(alength,X36) | $less(i(l8(sK2(X35))),0) | alength = i(l8(sK2(X35))) | i(l8(sK2(X35))) = X36 | i(l8(sK2(X35))) = i(l8(sK2(i(l8(sK2(X35)))))) (4:45) [resolution 2676,2659]
[SA] new: 4210. alength = X37 | $less(X37,0) | $less(alength,X37) | $less(X38,alength) | $less(alength,X38) | $less(i(l8(sK2(X37))),0) | i(l8(sK2(X37))) = 0 | alength = X38 | i(l8(sK2(X38))) = X38 (4:36) [resolution 2676,2951]
[SA] new: 4211. alength = X39 | $less(X39,0) | $less(alength,X39) | $less(alength,alength) | $less(i(l8(sK2(X39))),0) | i(l8(sK1(X40,sK2(i(l8(sK2(X39))))))) = X40 | alength = i(l8(sK2(X39))) | a(main_end,X40) = a(l11(sK2(i(l8(sK2(X39))))),X40) (4:47) [resolution 2676,2664]
[SA] new: 4212. alength = X41 | $less(X41,0) | $less(alength,X41) | $less(alength,alength) | $less(X42,0) | $less(i(l8(sK2(X41))),X42) | 0 = X42 | alength = i(l8(sK2(X41))) | i(l8(sK2(X41))) = i(l8(sK2(i(l8(sK2(X41)))))) (4:42) [resolution 2676,2951]
[SA] new: 4213. alength = X43 | $less(X43,0) | $less(alength,X43) | $less(alength,alength) | $less(i(l8(sK2(X43))),X44) | $less(X44,0) | alength = X44 | i(l8(sK2(X43))) = X44 | i(l8(sK2(X44))) = X44 (4:36) [resolution 2676,2659]
[SA] new: 4214. alength = X45 | $less(X45,0) | $less(alength,X45) | $less(alength,alength) | $less(i(l8(sK2(X45))),0) | i(l8(sK2(X45))) = i(l8(sK2(i(l8(sK2(X45)))))) | alength = i(l8(sK2(X45))) (4:36) [resolution 2676,2656]
[SA] new: 4215. alength = X46 | $less(X46,0) | $less(alength,X46) | $less(i(l8(p(nl8))),alength) | alength = i(l8(sK2(X46))) | $less(i(l8(sK2(X46))),alength) | zero = nl8 (4:30) [resolution 2676,193]
[SA] new: 4216. alength = X47 | $less(X47,0) | $less(alength,X47) | $less(i(l8(sK2(X48))),alength) | $less(X48,0) | alength = i(l8(sK2(X47))) | $less(i(l8(sK2(X47))),alength) | ~$less(X48,alength) (4:33) [resolution 2676,2921]
[SA] new: 4217. alength = X49 | $less(X49,0) | $less(alength,X49) | $less(0,alength) | $less(i(l8(sK2(X49))),sK0) | i(l8(sK2(X49))) = sK0 | 0 = sK0 (4:27) [resolution 2676,177]
[SA] new: 4218. alength = X50 | $less(X50,0) | $less(alength,X50) | $less(0,alength) | $less(i(l8(sK2(X50))),alength) | alength = i(l8(sK2(X50))) | 0 = sK0 (4:27) [resolution 2676,178]
[SA] new: 4219. alength = X51 | $less(X51,0) | $less(alength,X51) | $less(0,alength) | $less(i(l8(sK2(X51))),sK0) | 0 = sK0 | i(l8(sK2(X51))) = 0 (4:27) [resolution 2676,229]
[SA] new: 4220. alength = X52 | $less(X52,0) | $less(alength,X52) | $less(0,alength) | $less(i(l8(sK2(X52))),alength) | alength = i(l8(sK2(X52))) | alength = 0 (5:27) [resolution 2676,261]
[SA] new: 4221. alength = X53 | $less(X53,0) | $less(alength,X53) | $less(0,alength) | $less(i(l8(sK2(X53))),alength) | 0 = sK0 | i(l8(sK2(X53))) = 0 (5:27) [resolution 2676,278]
[SA] new: 4222. alength = X54 | $less(X54,0) | $less(alength,X54) | $less(0,alength) | $less(i(l8(sK2(X54))),X55) | $less(X55,sK0) | i(l8(sK2(X54))) = X55 | 0 = sK0 | sK0 = X55 (4:33) [resolution 2676,323]
[SA] new: 4223. alength = X56 | $less(X56,0) | $less(alength,X56) | $less(0,alength) | $less(i(l8(sK2(X56))),alength) | 0 = sK0 | i(l8(sK2(X56))) = sK0 (4:27) [resolution 2676,330]
[SA] new: 4224. alength = X57 | $less(X57,0) | $less(alength,X57) | $less(0,alength) | $less(X58,sK0) | $less(i(l8(sK2(X57))),X58) | sK0 = X58 | 0 = sK0 | i(l8(sK2(X57))) = sK0 (4:33) [resolution 2676,332]
[SA] new: 4225. alength = X59 | $less(X59,0) | $less(alength,X59) | $less(0,alength) | $less(i(l8(sK2(X59))),sK0) | i(l8(sK2(X59))) = 0 | i(l8(sK2(X59))) = sK0 (4:30) [resolution 2676,380]
[SA] new: 4226. alength = X60 | $less(X60,0) | $less(alength,X60) | $less(0,alength) | $less(X61,sK0) | $less(i(l8(sK2(X60))),X61) | i(l8(sK2(X60))) = X61 | 0 = sK0 | i(l8(sK2(X60))) = 0 (4:36) [resolution 2676,387]
[SA] new: 4227. alength = X62 | $less(X62,0) | $less(alength,X62) | $less(0,alength) | $less(X63,sK0) | $less(i(l8(sK2(X62))),X63) | 0 = X63 | 0 = sK0 | i(l8(sK2(X62))) = X63 (4:33) [resolution 2676,397]
[SA] new: 4228. alength = X64 | $less(X64,0) | $less(alength,X64) | $less(0,alength) | $less(X65,sK0) | $less(i(l8(sK2(X64))),X65) | i(l8(sK2(X64))) = 0 | 0 = sK0 | 0 = X65 (4:33) [resolution 2676,407]
[SA] new: 4229. alength = X66 | $less(X66,0) | $less(alength,X66) | $less(sK0,alength) | $less(i(l8(sK2(X66))),alength) | i(l8(sK2(X66))) = sK0 (4:24) [resolution 2676,106]
[SA] new: 4230. alength = X67 | $less(X67,0) | $less(alength,X67) | $less(sK0,alength) | $less(i(l8(sK2(X67))),alength) | alength = i(l8(sK2(X67))) (4:24) [resolution 2676,179]
[SA] new: 4231. alength = X68 | $less(X68,0) | $less(alength,X68) | $less(sK0,alength) | $less(i(l8(sK2(X68))),X69) | $less(X69,alength) | i(l8(sK2(X68))) = X69 | sK0 = X69 (4:30) [resolution 2676,198]
[SA] new: 4232. alength = X70 | $less(X70,0) | $less(alength,X70) | $less(sK0,alength) | $less(X71,alength) | $less(i(l8(sK2(X70))),X71) | alength = X71 | i(l8(sK2(X70))) = sK0 (4:30) [resolution 2676,204]
[SA] new: 4233. alength = X72 | $less(X72,0) | $less(alength,X72) | $less(sK0,alength) | $less(i(l8(sK2(X72))),X73) | $less(X73,alength) | i(l8(sK2(X72))) = X73 | alength = X73 (4:30) [resolution 2676,246]
[SA] new: 4234. alength = X74 | $less(X74,0) | $less(alength,X74) | $less(sK0,alength) | $less(X75,alength) | $less(i(l8(sK2(X74))),X75) | alength = X75 | alength = i(l8(sK2(X74))) (4:30) [resolution 2676,253]
[SA] new: 4235. alength = X76 | $less(X76,0) | $less(alength,X76) | $less(sK0,alength) | $less(X77,alength) | $less(i(l8(sK2(X76))),X77) | i(l8(sK2(X76))) = X77 | i(l8(sK2(X76))) = sK0 (4:33) [resolution 2676,388]
[SA] new: 4236. alength = X78 | $less(X78,0) | $less(alength,X78) | $less(sK0,alength) | $less(X79,alength) | $less(i(l8(sK2(X78))),X79) | i(l8(sK2(X78))) = sK0 | sK0 = X79 (4:30) [resolution 2676,408]
[SA] new: 4237. alength = X80 | $less(X80,0) | $less(alength,X80) | $less(sK0,alength) | $less(i(l8(sK2(X80))),0) | alength = i(l8(sK2(X80))) | i(l8(sK2(X80))) = sK0 | i(l8(sK2(X80))) = i(l8(sK2(i(l8(sK2(X80)))))) (4:42) [resolution 2676,2547]
[SA] new: 4238. alength = X81 | $less(X81,0) | $less(alength,X81) | $less(sK0,alength) | $less(i(l8(sK2(X81))),0) | i(l8(sK2(X81))) = i(l8(sK2(i(l8(sK2(X81)))))) | i(l8(sK2(X81))) = sK0 (4:36) [resolution 2676,2772]
[SA] new: 4239. ~$less(X0,i(l8(zero))) | alength = 0 | $less(0,0) | $less(alength,0) | $less(X0,alength) | alength = 0 (7:20) [superposition 2676,3571]
[SA] new: 4240. ~$less(X1,i(l8(zero))) | alength = 0 | $less(0,0) | $less(alength,0) | $less(X1,alength) | 0 = sK0 (5:20) [superposition 2676,3207]
[SA] new: 4241. ~$less(X0,sK0) | alength = sK0 | $less(sK0,0) | $less(alength,sK0) | $less(X0,alength) | alength = sK0 (4:18) [superposition 2676,2495]
[SA] new: 4242. ~$less(X1,sK0) | alength = sK0 | $less(sK0,0) | $less(alength,sK0) | $less(X1,alength) (4:15) [superposition 2676,2833]
[SA] new: 4243. ~$less(X2,0) | alength = 0 | $less(0,0) | $less(alength,0) | $less(X2,alength) | alength = 0 (6:18) [superposition 2676,2654]
[SA] new: 4244. ~$less(X3,0) | alength = 0 | $less(0,0) | $less(alength,0) | $less(X3,alength) | 0 = sK0 (4:18) [superposition 2676,2831]
[SA] new: 4245. ~$less(X3,0) | alength = 0 | $less(alength,0) | $less(X3,alength) | 0 = sK0 (4:15) [evaluation 4244]
[SA] new: 4246. ~$less(X2,0) | alength = 0 | $less(0,0) | $less(alength,0) | $less(X2,alength) (6:15) [duplicate literal removal 4243]
[SA] new: 4247. ~$less(X2,0) | alength = 0 | $less(alength,0) | $less(X2,alength) (6:12) [evaluation 4246]
[SA] new: 4248. ~$less(X0,sK0) | alength = sK0 | $less(sK0,0) | $less(alength,sK0) | $less(X0,alength) (4:15) [duplicate literal removal 4241]
[SA] new: 4249. ~$less(X1,i(l8(zero))) | alength = 0 | $less(alength,0) | $less(X1,alength) | 0 = sK0 (5:17) [evaluation 4240]
[SA] new: 4250. ~$less(X0,i(l8(zero))) | alength = 0 | $less(0,0) | $less(alength,0) | $less(X0,alength) (7:17) [duplicate literal removal 4239]
[SA] new: 4251. ~$less(X0,i(l8(zero))) | alength = 0 | $less(alength,0) | $less(X0,alength) (7:14) [evaluation 4250]
[SA] passive: 4193. alength = X0 | $less(X0,0) | $less(alength,X0) | $less(X1,alength) | $less(i(l8(sK2(X0))),X1) | i(l8(sK2(X0))) = X1 (4:24) [resolution 2676,15]
[SA] passive: 4195. alength = X4 | $less(X4,0) | $less(alength,X4) | $less(X5,alength) | $less(X6,X5) | $less(i(l8(sK2(X4))),X6) | i(l8(sK2(X4))) = X5 | i(l8(sK2(X4))) = X6 (4:33) [resolution 2676,174]
[SA] passive: 4196. alength = X7 | $less(X7,0) | $less(alength,X7) | $less(X8,alength) | $less(i(l8(sK2(X7))),X9) | $less(X9,X8) | X8 = X9 | i(l8(sK2(X7))) = X8 (4:30) [resolution 2676,174]
[SA] passive: 4197. alength = X10 | $less(X10,0) | $less(alength,X10) | $less(X11,alength) | $less(X12,X11) | $less(i(l8(sK2(X10))),X12) | i(l8(sK2(X10))) = X12 | X11 = X12 (4:30) [resolution 2676,174]
[SA] passive: 4202. alength = X21 | $less(X21,0) | $less(alength,X21) | $less(X22,alength) | $less(0,X22) | $less(i(l8(sK2(X21))),sK0) | i(l8(sK2(X21))) = X22 | 0 = sK0 | i(l8(sK2(X21))) = sK0 (4:36) [resolution 2676,323]
[SA] passive: 4206. alength = X29 | $less(X29,0) | $less(alength,X29) | $less(X30,alength) | $less(0,X30) | $less(i(l8(sK2(X29))),sK0) | i(l8(sK2(X29))) = 0 | 0 = sK0 | i(l8(sK2(X29))) = X30 (4:36) [resolution 2676,397]
[SA] new: 4252. alength = X41 | $less(X41,0) | $less(alength,X41) | $less(X42,0) | $less(i(l8(sK2(X41))),X42) | 0 = X42 | alength = i(l8(sK2(X41))) | i(l8(sK2(X41))) = i(l8(sK2(i(l8(sK2(X41)))))) (4:39) [subsumption resolution 4212,13]
[SA] passive: 4252. alength = X41 | $less(X41,0) | $less(alength,X41) | $less(X42,0) | $less(i(l8(sK2(X41))),X42) | 0 = X42 | alength = i(l8(sK2(X41))) | i(l8(sK2(X41))) = i(l8(sK2(i(l8(sK2(X41)))))) (4:39) [subsumption resolution 4212,13]
[SA] passive: 4225. alength = X59 | $less(X59,0) | $less(alength,X59) | $less(0,alength) | $less(i(l8(sK2(X59))),sK0) | i(l8(sK2(X59))) = 0 | i(l8(sK2(X59))) = sK0 (4:30) [resolution 2676,380]
[SA] new: 4253. ~$less(X0,0) | alength = 0 | $less(alength,0) | $less(X0,alength) (7:12) [forward demodulation 4251,65]
[SA] new: 4254. ~$less(X1,0) | alength = 0 | $less(alength,0) | $less(X1,alength) | 0 = sK0 (5:15) [forward demodulation 4249,65]
[SA] active: 2532. v = a(l8(s(sK2(sK0))),sK0) | alength = sK0 (4:11:1) [backward demodulation 2476,475]
[SA] new: 4255. v = a(l8(sK2(sK0)),sK0) | alength = sK0 | i(l8(sK2(sK0))) = sK0 (5:16) [superposition 2532,440]
[SA] new: 4256. v = a(l11(sK2(sK0)),sK0) | alength = sK0 (5:10) [superposition 2532,57]
[SA] new: 4257. v = a(l8(sK2(sK0)),sK0) | i(l8(sK2(sK0))) = sK0 | alength = sK0 (5:16) [superposition 440,2532]
[SA] new: 4258. v = a(l11(sK2(sK0)),sK0) | alength = sK0 (5:10) [superposition 57,2532]
[SA] active: 2948. $less(alength,X4) | $less(X4,sK0) | alength = X4 | 0 = sK0 | i(l8(sK2(X4))) = X4 (3:18:2) [resolution 2656,114]
[SA] new: 4259. $less(sK0,sK0) | alength = sK0 | 0 = sK0 | i(l8(sK2(sK0))) = sK0 | alength = sK0 (6:18) [resolution 2948,2704]
[SA] new: 4260. $less(alength,sK0) | alength = alength | 0 = sK0 | alength = i(l8(sK2(alength))) (4:15) [resolution 2948,13]
[SA] new: 4261. $less(X0,sK0) | alength = X0 | 0 = sK0 | i(l8(sK2(X0))) = X0 | ~$less(X1,alength) | $less(X1,X0) (4:21) [resolution 2948,14]
[SA] new: 4262. $less(X2,sK0) | alength = X2 | 0 = sK0 | i(l8(sK2(X2))) = X2 | X2 = X3 | $less(X3,X2) | $less(alength,X3) (4:24) [resolution 2948,99]
[SA] new: 4263. $less(X4,sK0) | alength = X4 | 0 = sK0 | i(l8(sK2(X4))) = X4 | sK0 = X4 | $less(sK0,X4) | $less(alength,alength) (4:24) [resolution 2948,205]
[SA] new: 4264. $less(X5,sK0) | alength = X5 | 0 = sK0 | i(l8(sK2(X5))) = X5 | alength = X5 | $less(sK0,X5) | $less(alength,alength) (4:24) [resolution 2948,254]
[SA] new: 4265. $less(X6,sK0) | alength = X6 | 0 = sK0 | i(l8(sK2(X6))) = X6 | sK0 = X6 | 0 = sK0 | $less(0,X6) | $less(alength,sK0) (4:27) [resolution 2948,333]
[SA] new: 4266. $less(X7,sK0) | alength = X7 | 0 = sK0 | i(l8(sK2(X7))) = X7 | i(l8(sK2(X7))) = X7 | alength = X7 | $less(alength,X7) | $less(alength,0) (4:30) [resolution 2948,2950]
[SA] new: 4267. $less(alength,sK0) | alength = alength | 0 = sK0 | alength = i(l8(sK2(alength))) | $less(alength,0) | a(main_end,X8) = a(l11(sK2(alength)),X8) | i(l8(sK1(X8,sK2(alength)))) = X8 (4:35) [resolution 2948,2770]
[SA] new: 4268. $less(alength,sK0) | alength = alength | 0 = sK0 | alength = i(l8(sK2(alength))) | alength = X9 | $less(X9,0) | i(l8(sK2(X9))) = X9 | $less(alength,X9) (4:30) [resolution 2948,2660]
[SA] new: 4269. $less(alength,sK0) | alength = alength | 0 = sK0 | alength = i(l8(sK2(alength))) | $less(alength,0) | alength = i(l8(sK2(alength))) (4:24) [resolution 2948,2480]
[SA] new: 4270. $less(i(l8(p(nl8))),sK0) | alength = i(l8(p(nl8))) | 0 = sK0 | i(l8(p(nl8))) = i(l8(sK2(i(l8(p(nl8)))))) | zero = nl8 | $less(alength,alength) (4:33) [resolution 2948,194]
[SA] new: 4271. $less(i(l8(sK2(X10))),sK0) | alength = i(l8(sK2(X10))) | 0 = sK0 | i(l8(sK2(X10))) = i(l8(sK2(i(l8(sK2(X10)))))) | alength = X10 | $less(X10,0) | $less(alength,X10) | $less(alength,alength) (4:39) [resolution 2948,2676]
[SA] new: 4272. $less(i(l8(sK2(X11))),sK0) | alength = i(l8(sK2(X11))) | 0 = sK0 | i(l8(sK2(X11))) = i(l8(sK2(i(l8(sK2(X11)))))) | $less(X11,0) | ~$less(X11,alength) | $less(alength,alength) (4:36) [resolution 2948,2920]
[SA] new: 4273. $less(0,sK0) | alength = 0 | 0 = sK0 | i(l8(sK2(0))) = 0 | sK0 = X12 | 0 = sK0 | $less(X12,sK0) | $less(alength,X12) (4:27) [resolution 2948,324]
[SA] new: 4274. $less(0,sK0) | alength = 0 | 0 = sK0 | i(l8(sK2(0))) = 0 | alength = 0 | $less(alength,alength) (5:21) [resolution 2948,262]
[SA] new: 4275. $less(0,sK0) | alength = 0 | 0 = sK0 | i(l8(sK2(0))) = 0 | 0 = sK0 | $less(alength,alength) (4:21) [resolution 2948,116]
[SA] new: 4276. $less(0,sK0) | alength = 0 | 0 = sK0 | i(l8(sK2(0))) = 0 | 0 = sK0 | $less(alength,sK0) (4:21) [resolution 2948,114]
[SA] new: 4277. $less(sK0,sK0) | alength = sK0 | 0 = sK0 | i(l8(sK2(sK0))) = sK0 | alength = X13 | $less(X13,alength) | $less(alength,X13) (4:24) [resolution 2948,247]
[SA] new: 4278. $less(sK0,sK0) | alength = sK0 | 0 = sK0 | i(l8(sK2(sK0))) = sK0 | sK0 = X14 | $less(X14,alength) | $less(alength,X14) (4:24) [resolution 2948,199]
[SA] new: 4279. $less(sK0,sK0) | alength = sK0 | 0 = sK0 | i(l8(sK2(sK0))) = sK0 | $less(alength,alength) (4:18) [resolution 2948,96]
[SA] new: 4280. $less(alength,X0) | alength = X0 | 0 = sK0 | i(l8(sK2(X0))) = X0 | alength = X1 | $less(X1,alength) | $less(X0,X1) (4:24) [resolution 2948,247]
[SA] new: 4281. $less(alength,X2) | alength = X2 | 0 = sK0 | i(l8(sK2(X2))) = X2 | sK0 = X3 | $less(X3,alength) | $less(X2,X3) (4:24) [resolution 2948,199]
[SA] new: 4282. $less(alength,X4) | alength = X4 | 0 = sK0 | i(l8(sK2(X4))) = X4 | $less(X4,alength) (4:18) [resolution 2948,96]
[SA] new: 4283. $less(alength,sK0) | alength = sK0 | 0 = sK0 | i(l8(sK2(sK0))) = sK0 (4:15) [resolution 2948,13]
[SA] new: 4284. $less(alength,X5) | alength = X5 | 0 = sK0 | i(l8(sK2(X5))) = X5 | ~$less(X6,X5) | $less(X6,sK0) (4:21) [resolution 2948,14]
[SA] new: 4285. $less(alength,X7) | alength = X7 | 0 = sK0 | i(l8(sK2(X7))) = X7 | sK0 = X8 | $less(X8,sK0) | $less(X7,X8) (4:24) [resolution 2948,99]
[SA] new: 4286. $less(alength,X9) | alength = X9 | 0 = sK0 | i(l8(sK2(X9))) = X9 | sK0 = sK0 | $less(sK0,sK0) | $less(X9,alength) (4:24) [resolution 2948,205]
[SA] new: 4287. $less(alength,X10) | alength = X10 | 0 = sK0 | i(l8(sK2(X10))) = X10 | alength = sK0 | $less(sK0,sK0) | $less(X10,alength) (4:24) [resolution 2948,254]
[SA] new: 4288. $less(alength,X11) | alength = X11 | 0 = sK0 | i(l8(sK2(X11))) = X11 | sK0 = sK0 | 0 = sK0 | $less(0,sK0) | $less(X11,sK0) (4:27) [resolution 2948,333]
[SA] new: 4289. $less(alength,X12) | alength = X12 | 0 = sK0 | i(l8(sK2(X12))) = X12 | i(l8(sK2(sK0))) = sK0 | alength = sK0 | $less(alength,sK0) | $less(X12,0) (4:30) [resolution 2948,2950]
[SA] new: 4290. $less(alength,alength) | alength = alength | 0 = sK0 | alength = i(l8(sK2(alength))) | alength = sK0 (6:18) [resolution 2948,2704]
[SA] new: 4291. $less(alength,X11) | alength = X11 | 0 = sK0 | i(l8(sK2(X11))) = X11 | sK0 = sK0 | $less(0,sK0) | $less(X11,sK0) (4:24) [duplicate literal removal 4288]
[SA] new: 4292. $less(0,sK0) | alength = 0 | 0 = sK0 | i(l8(sK2(0))) = 0 | $less(alength,sK0) (4:18) [duplicate literal removal 4276]
[SA] new: 4293. $less(0,sK0) | alength = 0 | 0 = sK0 | i(l8(sK2(0))) = 0 | $less(alength,alength) (4:18) [duplicate literal removal 4275]
[SA] new: 4294. $less(0,sK0) | alength = 0 | 0 = sK0 | i(l8(sK2(0))) = 0 | $less(alength,alength) (5:18) [duplicate literal removal 4274]
[SA] new: 4295. $less(0,sK0) | alength = 0 | 0 = sK0 | i(l8(sK2(0))) = 0 | sK0 = X12 | $less(X12,sK0) | $less(alength,X12) (4:24) [duplicate literal removal 4273]
[SA] new: 4296. $less(alength,sK0) | alength = alength | 0 = sK0 | alength = i(l8(sK2(alength))) | $less(alength,0) (4:18) [duplicate literal removal 4269]
[SA] new: 4297. $less(X7,sK0) | alength = X7 | 0 = sK0 | i(l8(sK2(X7))) = X7 | $less(alength,X7) | $less(alength,0) (4:21) [duplicate literal removal 4266]
[SA] new: 4298. $less(X6,sK0) | alength = X6 | 0 = sK0 | i(l8(sK2(X6))) = X6 | sK0 = X6 | $less(0,X6) | $less(alength,sK0) (4:24) [duplicate literal removal 4265]
[SA] new: 4299. $less(X5,sK0) | alength = X5 | 0 = sK0 | i(l8(sK2(X5))) = X5 | $less(sK0,X5) | $less(alength,alength) (4:21) [duplicate literal removal 4264]
[SA] new: 4300. $less(sK0,sK0) | alength = sK0 | 0 = sK0 | i(l8(sK2(sK0))) = sK0 (6:15) [duplicate literal removal 4259]
[SA] passive: 4261. $less(X0,sK0) | alength = X0 | 0 = sK0 | i(l8(sK2(X0))) = X0 | ~$less(X1,alength) | $less(X1,X0) (4:21) [resolution 2948,14]
[SA] new: 4301. $less(X5,sK0) | alength = X5 | 0 = sK0 | i(l8(sK2(X5))) = X5 | $less(sK0,X5) (4:18) [subsumption resolution 4299,13]
[SA] passive: 4301. $less(X5,sK0) | alength = X5 | 0 = sK0 | i(l8(sK2(X5))) = X5 | $less(sK0,X5) (4:18) [subsumption resolution 4299,13]
[SA] new: 4302. $less(i(l8(p(nl8))),sK0) | alength = i(l8(p(nl8))) | 0 = sK0 | i(l8(p(nl8))) = i(l8(sK2(i(l8(p(nl8)))))) | zero = nl8 (4:30) [subsumption resolution 4270,13]
[SA] passive: 4302. $less(i(l8(p(nl8))),sK0) | alength = i(l8(p(nl8))) | 0 = sK0 | i(l8(p(nl8))) = i(l8(sK2(i(l8(p(nl8)))))) | zero = nl8 (4:30) [subsumption resolution 4270,13]
[SA] new: 4303. $less(i(l8(sK2(X10))),sK0) | alength = i(l8(sK2(X10))) | 0 = sK0 | i(l8(sK2(X10))) = i(l8(sK2(i(l8(sK2(X10)))))) | alength = X10 | $less(X10,0) | $less(alength,X10) (4:36) [subsumption resolution 4271,13]
[SA] passive: 4303. $less(i(l8(sK2(X10))),sK0) | alength = i(l8(sK2(X10))) | 0 = sK0 | i(l8(sK2(X10))) = i(l8(sK2(i(l8(sK2(X10)))))) | alength = X10 | $less(X10,0) | $less(alength,X10) (4:36) [subsumption resolution 4271,13]
[SA] passive: 4284. $less(alength,X5) | alength = X5 | 0 = sK0 | i(l8(sK2(X5))) = X5 | ~$less(X6,X5) | $less(X6,sK0) (4:21) [resolution 2948,14]
[SA] active: 2993. v = a(l8(s(s(sK2(sK0)))),$sum(sK0,1)) (4:11:1) [superposition 349,2842]
[SA] new: 4304. v = a(l8(s(sK2(sK0))),$sum(sK0,1)) | i(l8(s(sK2(sK0)))) = $sum(sK0,1) (5:19) [superposition 2993,440]
[SA] new: 4305. v = a(l11(s(sK2(sK0))),$sum(sK0,1)) (5:10) [superposition 2993,57]
[SA] new: 4306. v = a(l8(s(sK2(sK0))),$sum(sK0,1)) | i(l8(s(sK2(sK0)))) = $sum(sK0,1) (5:19) [superposition 440,2993]
[SA] new: 4307. v = a(l11(s(sK2(sK0))),$sum(sK0,1)) (5:10) [superposition 57,2993]
[SA] active: 383. ~$less(X10,X8) | $less(X9,X7) | X7 = X9 | X7 = X8 | $less(X7,X8) | $less(X10,X9) (3:18:1) T [resolution 174,14]
[SA] new: 4308. $less(X0,X1) | X0 = X1 | X1 = X2 | $less(X1,X2) | $less(X3,X0) | $less(X2,X3) | X2 = X3 (4:21) T [resolution 383,15]
[SA] new: 4309. $less(X4,X5) | X4 = X5 | X5 = X6 | $less(X5,X6) | $less(X7,X4) | $less(X6,X7) | X6 = X7 (4:21) T [resolution 383,15]
[SA] new: 4310. $less(X8,X9) | X8 = X9 | X9 = X10 | $less(X9,X10) | $less(X11,X8) | $less(X12,X11) | $less(X10,X12) | X10 = X11 | X10 = X12 (4:27) T [resolution 383,174]
[SA] new: 4311. $less(X13,X14) | X13 = X14 | X14 = X15 | $less(X14,X15) | $less(X16,X13) | $less(X15,X17) | $less(X17,X16) | X16 = X17 | X15 = X16 (4:27) T [resolution 383,174]
[SA] new: 4312. $less(X18,X19) | X18 = X19 | X19 = X20 | $less(X19,X20) | $less(X21,X18) | $less(X22,X21) | $less(X20,X22) | X20 = X22 | X21 = X22 (4:27) T [resolution 383,174]
[SA] new: 4313. $less(X23,X24) | X23 = X24 | X24 = X25 | $less(X24,X25) | $less(X26,X23) | $less(sK0,X26) | $less(X25,alength) | X25 = X26 | sK0 = X25 (4:27) [resolution 383,198]
[SA] new: 4314. $less(X27,X28) | X27 = X28 | X28 = X29 | $less(X28,X29) | $less(X30,X27) | $less(sK0,X30) | $less(X29,alength) | alength = X29 | sK0 = X30 (4:27) [resolution 383,204]
[SA] new: 4315. $less(X31,X32) | X31 = X32 | X32 = X33 | $less(X32,X33) | $less(X34,X31) | $less(sK0,X34) | $less(X33,alength) | X33 = X34 | alength = X33 (4:27) [resolution 383,246]
[SA] new: 4316. $less(X35,X36) | X35 = X36 | X36 = X37 | $less(X36,X37) | $less(X38,X35) | $less(sK0,X38) | $less(X37,alength) | alength = X37 | alength = X38 (4:27) [resolution 383,253]
[SA] new: 4317. $less(X39,X40) | X39 = X40 | X40 = X41 | $less(X40,X41) | $less(X42,X39) | $less(0,X42) | $less(X41,sK0) | X41 = X42 | 0 = sK0 | sK0 = X41 (4:30) [resolution 383,323]
[SA] new: 4318. $less(X43,X44) | X43 = X44 | X44 = X45 | $less(X44,X45) | $less(X46,X43) | $less(0,X46) | $less(X45,sK0) | sK0 = X45 | 0 = sK0 | sK0 = X46 (4:30) [resolution 383,332]
[SA] new: 4319. $less(X47,X48) | X47 = X48 | X48 = X49 | $less(X48,X49) | $less(X50,X47) | $less(0,X50) | $less(X49,sK0) | X49 = X50 | 0 = sK0 | 0 = X50 (4:30) [resolution 383,387]
[SA] new: 4320. $less(X51,X52) | X51 = X52 | X52 = X53 | $less(X52,X53) | $less(X54,X51) | $less(sK0,X54) | $less(X53,alength) | X53 = X54 | sK0 = X54 (4:27) [resolution 383,388]
[SA] new: 4321. $less(X55,X56) | X55 = X56 | X56 = X57 | $less(X56,X57) | $less(X58,X55) | $less(0,X58) | $less(X57,sK0) | 0 = X57 | 0 = sK0 | X57 = X58 (4:30) [resolution 383,397]
[SA] new: 4322. $less(X59,X60) | X59 = X60 | X60 = X61 | $less(X60,X61) | $less(X62,X59) | $less(0,X62) | $less(X61,sK0) | 0 = X62 | 0 = sK0 | 0 = X61 (4:30) [resolution 383,407]
[SA] new: 4323. $less(X63,X64) | X63 = X64 | X64 = X65 | $less(X64,X65) | $less(X66,X63) | $less(sK0,X66) | $less(X65,alength) | sK0 = X66 | sK0 = X65 (4:27) [resolution 383,408]
[SA] new: 4324. $less(X67,X68) | X67 = X68 | X68 = X69 | $less(X68,X69) | $less(X70,X67) | $less(alength,X70) | $less(X69,0) | alength = X69 | X69 = X70 | i(l8(sK2(X69))) = X69 (4:33) [resolution 383,2659]
[SA] new: 4325. $less(X71,X72) | X71 = X72 | X72 = X73 | $less(X72,X73) | $less(X74,X71) | $less(alength,X74) | $less(X73,0) | 0 = X73 | alength = X74 | i(l8(sK2(X74))) = X74 (4:33) [resolution 383,2951]
[SA] new: 4326. $less(X75,X76) | X75 = X76 | alength = X76 | $less(X76,alength) | $less(X77,X75) | $less(sK0,X77) | sK0 = X77 (4:21) [resolution 383,106]
[SA] new: 4327. $less(X78,X79) | X78 = X79 | alength = X79 | $less(X79,alength) | $less(X80,X78) | $less(0,X80) | alength = X80 | 0 = sK0 (4:24) [resolution 383,178]
[SA] new: 4328. $less(X81,X82) | X81 = X82 | alength = X82 | $less(X82,alength) | $less(X83,X81) | $less(sK0,X83) | alength = X83 (4:21) [resolution 383,179]
[SA] new: 4329. $less(X84,X85) | X84 = X85 | alength = X85 | $less(X85,alength) | $less(X86,X84) | $less(sK0,X87) | $less(X87,X86) | X86 = X87 | sK0 = X86 (4:27) [resolution 383,198]
[SA] new: 4330. $less(X88,X89) | X88 = X89 | alength = X89 | $less(X89,alength) | $less(X90,X88) | $less(sK0,X91) | $less(X91,X90) | alength = X90 | sK0 = X91 (4:27) [resolution 383,204]
[SA] new: 4331. $less(X92,X93) | X92 = X93 | alength = X93 | $less(X93,alength) | $less(X94,X92) | $less(sK0,X95) | $less(X95,X94) | X94 = X95 | alength = X94 (4:27) [resolution 383,246]
[SA] new: 4332. $less(X96,X97) | X96 = X97 | alength = X97 | $less(X97,alength) | $less(X98,X96) | $less(sK0,X99) | $less(X99,X98) | alength = X98 | alength = X99 (4:27) [resolution 383,253]
[SA] new: 4333. $less(X100,X101) | X100 = X101 | alength = X101 | $less(X101,alength) | $less(X102,X100) | $less(0,X102) | alength = X102 | alength = 0 (5:24) [resolution 383,261]
[SA] new: 4334. $less(X103,X104) | X103 = X104 | alength = X104 | $less(X104,alength) | $less(X105,X103) | $less(0,X105) | 0 = sK0 | 0 = X105 (5:24) [resolution 383,278]
[SA] new: 4335. $less(X106,X107) | X106 = X107 | alength = X107 | $less(X107,alength) | $less(X108,X106) | $less(0,X108) | 0 = sK0 | sK0 = X108 (4:24) [resolution 383,330]
[SA] new: 4336. $less(X109,X110) | X109 = X110 | alength = X110 | $less(X110,alength) | $less(X111,X109) | $less(sK0,X112) | $less(X112,X111) | X111 = X112 | sK0 = X112 (4:27) [resolution 383,388]
[SA] new: 4337. $less(X113,X114) | X113 = X114 | alength = X114 | $less(X114,alength) | $less(X115,X113) | $less(sK0,X116) | $less(X116,X115) | sK0 = X116 | sK0 = X115 (4:27) [resolution 383,408]
[SA] new: 4338. $less(X117,X118) | X117 = X118 | 0 = X118 | $less(X118,0) | $less(X119,X117) | $less(sK0,X119) | alength = X119 | sK0 = X119 | i(l8(sK2(X119))) = X119 (4:30) [resolution 383,2547]
[SA] new: 4339. $less(X120,X121) | X120 = X121 | 0 = X121 | $less(X121,0) | $less(X122,X120) | $less(sK0,X122) | i(l8(sK2(X122))) = X122 | sK0 = X122 (4:27) [resolution 383,2772]
[SA] new: 4340. $less(X123,X124) | X123 = X124 | 0 = X124 | $less(X124,0) | $less(X125,X123) | $less(alength,X125) | i(l8(sK1(X126,sK2(X125)))) = X126 | alength = X125 | a(main_end,X126) = a(l11(sK2(X125)),X126) (4:38) [resolution 383,2664]
[SA] new: 4341. $less(X127,X128) | X127 = X128 | 0 = X128 | $less(X128,0) | $less(X129,X127) | $less(alength,X130) | $less(X130,X129) | 0 = X129 | alength = X130 | i(l8(sK2(X130))) = X130 (4:33) [resolution 383,2951]
[SA] new: 4342. $less(X131,X132) | X131 = X132 | 0 = X132 | $less(X132,0) | $less(X133,X131) | $less(alength,X134) | $less(X134,X133) | alength = X133 | X133 = X134 | i(l8(sK2(X133))) = X133 (4:33) [resolution 383,2659]
[SA] new: 4343. $less(X135,X136) | X135 = X136 | 0 = X136 | $less(X136,0) | $less(X137,X135) | $less(alength,X137) | i(l8(sK2(X137))) = X137 | alength = X137 (4:27) [resolution 383,2656]
[SA] new: 4344. $less(X138,X139) | X138 = X139 | sK0 = X139 | $less(X139,sK0) | $less(X140,X138) | $less(0,X140) | sK0 = X140 | 0 = sK0 (4:24) [resolution 383,177]
[SA] new: 4345. $less(X141,X142) | X141 = X142 | sK0 = X142 | $less(X142,sK0) | $less(X143,X141) | $less(0,X143) | 0 = sK0 | 0 = X143 (4:24) [resolution 383,229]
[SA] new: 4346. $less(X144,X145) | X144 = X145 | sK0 = X145 | $less(X145,sK0) | $less(X146,X144) | $less(0,X147) | $less(X147,X146) | X146 = X147 | 0 = sK0 | sK0 = X146 (4:30) [resolution 383,323]
[SA] new: 4347. $less(X148,X149) | X148 = X149 | sK0 = X149 | $less(X149,sK0) | $less(X150,X148) | $less(0,X151) | $less(X151,X150) | sK0 = X150 | 0 = sK0 | sK0 = X151 (4:30) [resolution 383,332]
[SA] new: 4348. $less(X152,X153) | X152 = X153 | sK0 = X153 | $less(X153,sK0) | $less(X154,X152) | $less(0,X154) | 0 = X154 | sK0 = X154 (4:24) [resolution 383,380]
[SA] new: 4349. $less(X155,X156) | X155 = X156 | sK0 = X156 | $less(X156,sK0) | $less(X157,X155) | $less(0,X158) | $less(X158,X157) | X157 = X158 | 0 = sK0 | 0 = X158 (4:30) [resolution 383,387]
[SA] new: 4350. $less(X159,X160) | X159 = X160 | sK0 = X160 | $less(X160,sK0) | $less(X161,X159) | $less(0,X162) | $less(X162,X161) | 0 = X161 | 0 = sK0 | X161 = X162 (4:30) [resolution 383,397]
[SA] new: 4351. $less(X163,X164) | X163 = X164 | sK0 = X164 | $less(X164,sK0) | $less(X165,X163) | $less(0,X166) | $less(X166,X165) | 0 = X166 | 0 = sK0 | 0 = X165 (4:30) [resolution 383,407]
[SA] new: 4352. $less(X167,X168) | X167 = X168 | sK0 = X168 | $less(X168,sK0) | $less(X169,X167) | $less(alength,X169) | alength = X169 | 0 = sK0 | i(l8(sK2(X169))) = X169 (4:30) [resolution 383,2948]
[SA] new: 4353. $less(X170,X171) | X170 = X171 | X171 = X172 | $less(X171,X172) | $less(alength,X170) | $less(X172,sK0) | alength = X172 | 0 = sK0 | i(l8(sK2(X172))) = X172 (4:30) [resolution 383,2948]
[SA] new: 4354. $less(X173,X174) | X173 = X174 | X174 = X175 | $less(X174,X175) | $less(alength,X173) | $less(X175,0) | i(l8(sK1(X176,sK2(X175)))) = X176 | alength = X175 | a(main_end,X176) = a(l11(sK2(X175)),X176) (4:38) [resolution 383,2664]
[SA] new: 4355. $less(X177,X178) | X177 = X178 | X178 = X179 | $less(X178,X179) | $less(alength,X177) | $less(X180,0) | $less(X179,X180) | 0 = X180 | alength = X179 | i(l8(sK2(X179))) = X179 (4:33) [resolution 383,2951]
[SA] new: 4356. $less(X181,X182) | X181 = X182 | X182 = X183 | $less(X182,X183) | $less(alength,X181) | $less(X183,X184) | $less(X184,0) | alength = X184 | X183 = X184 | i(l8(sK2(X184))) = X184 (4:33) [resolution 383,2659]
[SA] new: 4357. $less(X185,X186) | X185 = X186 | X186 = X187 | $less(X186,X187) | $less(alength,X185) | $less(X187,0) | i(l8(sK2(X187))) = X187 | alength = X187 (4:27) [resolution 383,2656]
[SA] new: 4358. $less(X188,X189) | X188 = X189 | sK0 = X189 | $less(X189,sK0) | $less(alength,X188) | alength = sK0 | 0 = sK0 | i(l8(sK2(0))) = 0 (4:27) [resolution 383,2667]
[SA] new: 4359. $less(X190,X191) | X190 = X191 | sK0 = X191 | $less(X191,sK0) | $less(alength,X190) | alength = 0 | 0 = sK0 | i(l8(sK2(0))) = 0 (4:27) [resolution 383,2665]
[SA] new: 4360. $less(X192,X193) | X192 = X193 | X193 = X194 | $less(X193,X194) | $less(i(l8(p(nl8))),X192) | alength = X194 | $less(X194,alength) | zero = nl8 (4:27) [resolution 383,193]
[SA] new: 4361. $less(X195,X196) | X195 = X196 | X196 = X197 | $less(X196,X197) | $less(i(l8(sK2(X198))),X195) | $less(X198,0) | alength = X197 | $less(X197,alength) | ~$less(X198,alength) (4:30) [resolution 383,2921]
[SA] new: 4362. $less(X199,X200) | X199 = X200 | alength = X200 | $less(X200,alength) | $less(i(l8(p(nl8))),X199) | zero = nl8 (4:21) [resolution 383,95]
[SA] new: 4363. $less(X201,X202) | X201 = X202 | alength = X202 | $less(X202,alength) | $less(i(l8(p(p(nl8)))),X201) | zero = p(nl8) | zero = nl8 (4:26) [resolution 383,314]
[SA] new: 4364. $less(X203,X204) | X203 = X204 | alength = X204 | $less(X204,alength) | $less(i(l8(sK2(X205))),X203) | $less(alength,X205) | alength = X205 | $less(X205,0) (4:27) [resolution 383,2657]
[SA] new: 4365. $less(X206,X207) | X206 = X207 | alength = X207 | $less(X207,alength) | $less(i(l8(sK2(X208))),X206) | ~$less(X208,alength) | $less(X208,0) (4:24) [resolution 383,2771]
[SA] new: 4366. $less(X209,X210) | X209 = X210 | 0 = X210 | $less(X210,0) | $less(i(l8(sK2(X211))),X209) | $less(X211,0) | ~$less(X211,alength) | i(l8(sK2(X211))) = i(l8(sK2(i(l8(sK2(X211)))))) (4:36) [resolution 383,2919]
[SA] new: 4367. $less(X212,X213) | X212 = X213 | X213 = X214 | $less(X213,X214) | $less(0,X212) | $less(X214,sK0) | sK0 = X214 | 0 = sK0 (4:24) [resolution 383,177]
[SA] new: 4368. $less(X215,X216) | X215 = X216 | X216 = X217 | $less(X216,X217) | $less(0,X215) | $less(X217,alength) | alength = X217 | 0 = sK0 (4:24) [resolution 383,178]
[SA] new: 4369. $less(X218,X219) | X218 = X219 | X219 = X220 | $less(X219,X220) | $less(0,X218) | $less(X220,sK0) | 0 = sK0 | 0 = X220 (4:24) [resolution 383,229]
[SA] new: 4370. $less(X221,X222) | X221 = X222 | X222 = X223 | $less(X222,X223) | $less(0,X221) | $less(X223,alength) | alength = X223 | alength = 0 (5:24) [resolution 383,261]
[SA] new: 4371. $less(X224,X225) | X224 = X225 | X225 = X226 | $less(X225,X226) | $less(0,X224) | $less(X226,alength) | 0 = sK0 | 0 = X226 (5:24) [resolution 383,278]
[SA] new: 4372. $less(X227,X228) | X227 = X228 | X228 = X229 | $less(X228,X229) | $less(0,X227) | $less(X229,X230) | $less(X230,sK0) | X229 = X230 | 0 = sK0 | sK0 = X230 (4:30) [resolution 383,323]
[SA] new: 4373. $less(X231,X232) | X231 = X232 | X232 = X233 | $less(X232,X233) | $less(0,X231) | $less(X233,alength) | 0 = sK0 | sK0 = X233 (4:24) [resolution 383,330]
[SA] new: 4374. $less(X234,X235) | X234 = X235 | X235 = X236 | $less(X235,X236) | $less(0,X234) | $less(X237,sK0) | $less(X236,X237) | sK0 = X237 | 0 = sK0 | sK0 = X236 (4:30) [resolution 383,332]
[SA] new: 4375. $less(X238,X239) | X238 = X239 | X239 = X240 | $less(X239,X240) | $less(0,X238) | $less(X240,sK0) | 0 = X240 | sK0 = X240 (4:24) [resolution 383,380]
[SA] new: 4376. $less(X241,X242) | X241 = X242 | X242 = X243 | $less(X242,X243) | $less(0,X241) | $less(X244,sK0) | $less(X243,X244) | X243 = X244 | 0 = sK0 | 0 = X243 (4:30) [resolution 383,387]
[SA] new: 4377. $less(X245,X246) | X245 = X246 | X246 = X247 | $less(X246,X247) | $less(0,X245) | $less(X248,sK0) | $less(X247,X248) | 0 = X248 | 0 = sK0 | X247 = X248 (4:30) [resolution 383,397]
[SA] new: 4378. $less(X249,X250) | X249 = X250 | X250 = X251 | $less(X250,X251) | $less(0,X249) | $less(X252,sK0) | $less(X251,X252) | 0 = X251 | 0 = sK0 | 0 = X252 (4:30) [resolution 383,407]
[SA] new: 4379. $less(X253,X254) | X253 = X254 | alength = X254 | $less(X254,alength) | $less(0,X253) | alength = 0 (4:18) [resolution 383,244]
[SA] new: 4380. $less(X255,X256) | X255 = X256 | alength = X256 | $less(X256,alength) | $less(0,X255) | 0 = sK0 (4:18) [resolution 383,113]
[SA] new: 4381. $less(X257,X258) | X257 = X258 | sK0 = X258 | $less(X258,sK0) | $less(0,X257) | 0 = sK0 (4:18) [resolution 383,97]
[SA] new: 4382. $less(X259,X260) | X259 = X260 | alength = X260 | $less(X260,alength) | $less(sK0,X259) (4:15) [resolution 383,45]
[SA] new: 4383. $less(X261,X262) | X261 = X262 | X262 = X263 | $less(X262,X263) | $less(sK0,X261) | $less(X263,alength) | sK0 = X263 (4:21) [resolution 383,106]
[SA] new: 4384. $less(X264,X265) | X264 = X265 | X265 = X266 | $less(X265,X266) | $less(sK0,X264) | $less(X266,alength) | alength = X266 (4:21) [resolution 383,179]
[SA] new: 4385. $less(X267,X268) | X267 = X268 | X268 = X269 | $less(X268,X269) | $less(sK0,X267) | $less(X269,X270) | $less(X270,alength) | X269 = X270 | sK0 = X270 (4:27) [resolution 383,198]
[SA] new: 4386. $less(X271,X272) | X271 = X272 | X272 = X273 | $less(X272,X273) | $less(sK0,X271) | $less(X274,alength) | $less(X273,X274) | alength = X274 | sK0 = X273 (4:27) [resolution 383,204]
[SA] new: 4387. $less(X275,X276) | X275 = X276 | X276 = X277 | $less(X276,X277) | $less(sK0,X275) | $less(X277,X278) | $less(X278,alength) | X277 = X278 | alength = X278 (4:27) [resolution 383,246]
[SA] new: 4388. $less(X279,X280) | X279 = X280 | X280 = X281 | $less(X280,X281) | $less(sK0,X279) | $less(X282,alength) | $less(X281,X282) | alength = X282 | alength = X281 (4:27) [resolution 383,253]
[SA] new: 4389. $less(X283,X284) | X283 = X284 | X284 = X285 | $less(X284,X285) | $less(sK0,X283) | $less(X286,alength) | $less(X285,X286) | X285 = X286 | sK0 = X285 (4:27) [resolution 383,388]
[SA] new: 4390. $less(X287,X288) | X287 = X288 | X288 = X289 | $less(X288,X289) | $less(sK0,X287) | $less(X290,alength) | $less(X289,X290) | sK0 = X289 | sK0 = X290 (4:27) [resolution 383,408]
[SA] new: 4391. $less(X291,X292) | X291 = X292 | X292 = X293 | $less(X292,X293) | $less(sK0,X291) | $less(X293,0) | alength = X293 | sK0 = X293 | i(l8(sK2(X293))) = X293 (4:30) [resolution 383,2547]
[SA] new: 4392. $less(X294,X295) | X294 = X295 | X295 = X296 | $less(X295,X296) | $less(sK0,X294) | $less(X296,0) | i(l8(sK2(X296))) = X296 | sK0 = X296 (4:27) [resolution 383,2772]
[SA] passive: 4310. $less(X8,X9) | X8 = X9 | X9 = X10 | $less(X9,X10) | $less(X11,X8) | $less(X12,X11) | $less(X10,X12) | X10 = X11 | X10 = X12 (4:27) T [resolution 383,174]
[SA] passive: 4311. $less(X13,X14) | X13 = X14 | X14 = X15 | $less(X14,X15) | $less(X16,X13) | $less(X15,X17) | $less(X17,X16) | X16 = X17 | X15 = X16 (4:27) T [resolution 383,174]
[SA] passive: 4312. $less(X18,X19) | X18 = X19 | X19 = X20 | $less(X19,X20) | $less(X21,X18) | $less(X22,X21) | $less(X20,X22) | X20 = X22 | X21 = X22 (4:27) T [resolution 383,174]
[SA] passive: 4313. $less(X23,X24) | X23 = X24 | X24 = X25 | $less(X24,X25) | $less(X26,X23) | $less(sK0,X26) | $less(X25,alength) | X25 = X26 | sK0 = X25 (4:27) [resolution 383,198]
[SA] passive: 4314. $less(X27,X28) | X27 = X28 | X28 = X29 | $less(X28,X29) | $less(X30,X27) | $less(sK0,X30) | $less(X29,alength) | alength = X29 | sK0 = X30 (4:27) [resolution 383,204]
[SA] passive: 4315. $less(X31,X32) | X31 = X32 | X32 = X33 | $less(X32,X33) | $less(X34,X31) | $less(sK0,X34) | $less(X33,alength) | X33 = X34 | alength = X33 (4:27) [resolution 383,246]
[SA] passive: 4316. $less(X35,X36) | X35 = X36 | X36 = X37 | $less(X36,X37) | $less(X38,X35) | $less(sK0,X38) | $less(X37,alength) | alength = X37 | alength = X38 (4:27) [resolution 383,253]
[SA] passive: 4317. $less(X39,X40) | X39 = X40 | X40 = X41 | $less(X40,X41) | $less(X42,X39) | $less(0,X42) | $less(X41,sK0) | X41 = X42 | 0 = sK0 | sK0 = X41 (4:30) [resolution 383,323]
[SA] passive: 4318. $less(X43,X44) | X43 = X44 | X44 = X45 | $less(X44,X45) | $less(X46,X43) | $less(0,X46) | $less(X45,sK0) | sK0 = X45 | 0 = sK0 | sK0 = X46 (4:30) [resolution 383,332]
[SA] passive: 4319. $less(X47,X48) | X47 = X48 | X48 = X49 | $less(X48,X49) | $less(X50,X47) | $less(0,X50) | $less(X49,sK0) | X49 = X50 | 0 = sK0 | 0 = X50 (4:30) [resolution 383,387]
[SA] passive: 4320. $less(X51,X52) | X51 = X52 | X52 = X53 | $less(X52,X53) | $less(X54,X51) | $less(sK0,X54) | $less(X53,alength) | X53 = X54 | sK0 = X54 (4:27) [resolution 383,388]
[SA] passive: 4321. $less(X55,X56) | X55 = X56 | X56 = X57 | $less(X56,X57) | $less(X58,X55) | $less(0,X58) | $less(X57,sK0) | 0 = X57 | 0 = sK0 | X57 = X58 (4:30) [resolution 383,397]
[SA] passive: 4322. $less(X59,X60) | X59 = X60 | X60 = X61 | $less(X60,X61) | $less(X62,X59) | $less(0,X62) | $less(X61,sK0) | 0 = X62 | 0 = sK0 | 0 = X61 (4:30) [resolution 383,407]
[SA] passive: 4323. $less(X63,X64) | X63 = X64 | X64 = X65 | $less(X64,X65) | $less(X66,X63) | $less(sK0,X66) | $less(X65,alength) | sK0 = X66 | sK0 = X65 (4:27) [resolution 383,408]
[SA] passive: 4324. $less(X67,X68) | X67 = X68 | X68 = X69 | $less(X68,X69) | $less(X70,X67) | $less(alength,X70) | $less(X69,0) | alength = X69 | X69 = X70 | i(l8(sK2(X69))) = X69 (4:33) [resolution 383,2659]
[SA] passive: 4325. $less(X71,X72) | X71 = X72 | X72 = X73 | $less(X72,X73) | $less(X74,X71) | $less(alength,X74) | $less(X73,0) | 0 = X73 | alength = X74 | i(l8(sK2(X74))) = X74 (4:33) [resolution 383,2951]
[SA] passive: 4327. $less(X78,X79) | X78 = X79 | alength = X79 | $less(X79,alength) | $less(X80,X78) | $less(0,X80) | alength = X80 | 0 = sK0 (4:24) [resolution 383,178]
[SA] passive: 4329. $less(X84,X85) | X84 = X85 | alength = X85 | $less(X85,alength) | $less(X86,X84) | $less(sK0,X87) | $less(X87,X86) | X86 = X87 | sK0 = X86 (4:27) [resolution 383,198]
[SA] passive: 4330. $less(X88,X89) | X88 = X89 | alength = X89 | $less(X89,alength) | $less(X90,X88) | $less(sK0,X91) | $less(X91,X90) | alength = X90 | sK0 = X91 (4:27) [resolution 383,204]
[SA] passive: 4331. $less(X92,X93) | X92 = X93 | alength = X93 | $less(X93,alength) | $less(X94,X92) | $less(sK0,X95) | $less(X95,X94) | X94 = X95 | alength = X94 (4:27) [resolution 383,246]
[SA] passive: 4332. $less(X96,X97) | X96 = X97 | alength = X97 | $less(X97,alength) | $less(X98,X96) | $less(sK0,X99) | $less(X99,X98) | alength = X98 | alength = X99 (4:27) [resolution 383,253]
[SA] passive: 4333. $less(X100,X101) | X100 = X101 | alength = X101 | $less(X101,alength) | $less(X102,X100) | $less(0,X102) | alength = X102 | alength = 0 (5:24) [resolution 383,261]
[SA] passive: 4334. $less(X103,X104) | X103 = X104 | alength = X104 | $less(X104,alength) | $less(X105,X103) | $less(0,X105) | 0 = sK0 | 0 = X105 (5:24) [resolution 383,278]
[SA] passive: 4335. $less(X106,X107) | X106 = X107 | alength = X107 | $less(X107,alength) | $less(X108,X106) | $less(0,X108) | 0 = sK0 | sK0 = X108 (4:24) [resolution 383,330]
[SA] passive: 4336. $less(X109,X110) | X109 = X110 | alength = X110 | $less(X110,alength) | $less(X111,X109) | $less(sK0,X112) | $less(X112,X111) | X111 = X112 | sK0 = X112 (4:27) [resolution 383,388]
[SA] passive: 4337. $less(X113,X114) | X113 = X114 | alength = X114 | $less(X114,alength) | $less(X115,X113) | $less(sK0,X116) | $less(X116,X115) | sK0 = X116 | sK0 = X115 (4:27) [resolution 383,408]
[SA] passive: 4338. $less(X117,X118) | X117 = X118 | 0 = X118 | $less(X118,0) | $less(X119,X117) | $less(sK0,X119) | alength = X119 | sK0 = X119 | i(l8(sK2(X119))) = X119 (4:30) [resolution 383,2547]
[SA] passive: 4339. $less(X120,X121) | X120 = X121 | 0 = X121 | $less(X121,0) | $less(X122,X120) | $less(sK0,X122) | i(l8(sK2(X122))) = X122 | sK0 = X122 (4:27) [resolution 383,2772]
[SA] passive: 4340. $less(X123,X124) | X123 = X124 | 0 = X124 | $less(X124,0) | $less(X125,X123) | $less(alength,X125) | i(l8(sK1(X126,sK2(X125)))) = X126 | alength = X125 | a(main_end,X126) = a(l11(sK2(X125)),X126) (4:38) [resolution 383,2664]
[SA] passive: 4341. $less(X127,X128) | X127 = X128 | 0 = X128 | $less(X128,0) | $less(X129,X127) | $less(alength,X130) | $less(X130,X129) | 0 = X129 | alength = X130 | i(l8(sK2(X130))) = X130 (4:33) [resolution 383,2951]
[SA] passive: 4342. $less(X131,X132) | X131 = X132 | 0 = X132 | $less(X132,0) | $less(X133,X131) | $less(alength,X134) | $less(X134,X133) | alength = X133 | X133 = X134 | i(l8(sK2(X133))) = X133 (4:33) [resolution 383,2659]
[SA] passive: 4345. $less(X141,X142) | X141 = X142 | sK0 = X142 | $less(X142,sK0) | $less(X143,X141) | $less(0,X143) | 0 = sK0 | 0 = X143 (4:24) [resolution 383,229]
[SA] passive: 4346. $less(X144,X145) | X144 = X145 | sK0 = X145 | $less(X145,sK0) | $less(X146,X144) | $less(0,X147) | $less(X147,X146) | X146 = X147 | 0 = sK0 | sK0 = X146 (4:30) [resolution 383,323]
[SA] passive: 4347. $less(X148,X149) | X148 = X149 | sK0 = X149 | $less(X149,sK0) | $less(X150,X148) | $less(0,X151) | $less(X151,X150) | sK0 = X150 | 0 = sK0 | sK0 = X151 (4:30) [resolution 383,332]
[SA] passive: 4348. $less(X152,X153) | X152 = X153 | sK0 = X153 | $less(X153,sK0) | $less(X154,X152) | $less(0,X154) | 0 = X154 | sK0 = X154 (4:24) [resolution 383,380]
[SA] passive: 4349. $less(X155,X156) | X155 = X156 | sK0 = X156 | $less(X156,sK0) | $less(X157,X155) | $less(0,X158) | $less(X158,X157) | X157 = X158 | 0 = sK0 | 0 = X158 (4:30) [resolution 383,387]
[SA] passive: 4350. $less(X159,X160) | X159 = X160 | sK0 = X160 | $less(X160,sK0) | $less(X161,X159) | $less(0,X162) | $less(X162,X161) | 0 = X161 | 0 = sK0 | X161 = X162 (4:30) [resolution 383,397]
[SA] passive: 4351. $less(X163,X164) | X163 = X164 | sK0 = X164 | $less(X164,sK0) | $less(X165,X163) | $less(0,X166) | $less(X166,X165) | 0 = X166 | 0 = sK0 | 0 = X165 (4:30) [resolution 383,407]
[SA] passive: 4352. $less(X167,X168) | X167 = X168 | sK0 = X168 | $less(X168,sK0) | $less(X169,X167) | $less(alength,X169) | alength = X169 | 0 = sK0 | i(l8(sK2(X169))) = X169 (4:30) [resolution 383,2948]
[SA] passive: 4353. $less(X170,X171) | X170 = X171 | X171 = X172 | $less(X171,X172) | $less(alength,X170) | $less(X172,sK0) | alength = X172 | 0 = sK0 | i(l8(sK2(X172))) = X172 (4:30) [resolution 383,2948]
[SA] passive: 4354. $less(X173,X174) | X173 = X174 | X174 = X175 | $less(X174,X175) | $less(alength,X173) | $less(X175,0) | i(l8(sK1(X176,sK2(X175)))) = X176 | alength = X175 | a(main_end,X176) = a(l11(sK2(X175)),X176) (4:38) [resolution 383,2664]
[SA] passive: 4355. $less(X177,X178) | X177 = X178 | X178 = X179 | $less(X178,X179) | $less(alength,X177) | $less(X180,0) | $less(X179,X180) | 0 = X180 | alength = X179 | i(l8(sK2(X179))) = X179 (4:33) [resolution 383,2951]
[SA] passive: 4356. $less(X181,X182) | X181 = X182 | X182 = X183 | $less(X182,X183) | $less(alength,X181) | $less(X183,X184) | $less(X184,0) | alength = X184 | X183 = X184 | i(l8(sK2(X184))) = X184 (4:33) [resolution 383,2659]
[SA] passive: 4360. $less(X192,X193) | X192 = X193 | X193 = X194 | $less(X193,X194) | $less(i(l8(p(nl8))),X192) | alength = X194 | $less(X194,alength) | zero = nl8 (4:27) [resolution 383,193]
[SA] passive: 4361. $less(X195,X196) | X195 = X196 | X196 = X197 | $less(X196,X197) | $less(i(l8(sK2(X198))),X195) | $less(X198,0) | alength = X197 | $less(X197,alength) | ~$less(X198,alength) (4:30) [resolution 383,2921]
[SA] passive: 4363. $less(X201,X202) | X201 = X202 | alength = X202 | $less(X202,alength) | $less(i(l8(p(p(nl8)))),X201) | zero = p(nl8) | zero = nl8 (4:26) [resolution 383,314]
[SA] passive: 4364. $less(X203,X204) | X203 = X204 | alength = X204 | $less(X204,alength) | $less(i(l8(sK2(X205))),X203) | $less(alength,X205) | alength = X205 | $less(X205,0) (4:27) [resolution 383,2657]
[SA] passive: 4366. $less(X209,X210) | X209 = X210 | 0 = X210 | $less(X210,0) | $less(i(l8(sK2(X211))),X209) | $less(X211,0) | ~$less(X211,alength) | i(l8(sK2(X211))) = i(l8(sK2(i(l8(sK2(X211)))))) (4:36) [resolution 383,2919]
[SA] passive: 4368. $less(X215,X216) | X215 = X216 | X216 = X217 | $less(X216,X217) | $less(0,X215) | $less(X217,alength) | alength = X217 | 0 = sK0 (4:24) [resolution 383,178]
[SA] passive: 4370. $less(X221,X222) | X221 = X222 | X222 = X223 | $less(X222,X223) | $less(0,X221) | $less(X223,alength) | alength = X223 | alength = 0 (5:24) [resolution 383,261]
[SA] passive: 4371. $less(X224,X225) | X224 = X225 | X225 = X226 | $less(X225,X226) | $less(0,X224) | $less(X226,alength) | 0 = sK0 | 0 = X226 (5:24) [resolution 383,278]
[SA] passive: 4372. $less(X227,X228) | X227 = X228 | X228 = X229 | $less(X228,X229) | $less(0,X227) | $less(X229,X230) | $less(X230,sK0) | X229 = X230 | 0 = sK0 | sK0 = X230 (4:30) [resolution 383,323]
[SA] passive: 4373. $less(X231,X232) | X231 = X232 | X232 = X233 | $less(X232,X233) | $less(0,X231) | $less(X233,alength) | 0 = sK0 | sK0 = X233 (4:24) [resolution 383,330]
[SA] passive: 4374. $less(X234,X235) | X234 = X235 | X235 = X236 | $less(X235,X236) | $less(0,X234) | $less(X237,sK0) | $less(X236,X237) | sK0 = X237 | 0 = sK0 | sK0 = X236 (4:30) [resolution 383,332]
[SA] passive: 4375. $less(X238,X239) | X238 = X239 | X239 = X240 | $less(X239,X240) | $less(0,X238) | $less(X240,sK0) | 0 = X240 | sK0 = X240 (4:24) [resolution 383,380]
[SA] passive: 4376. $less(X241,X242) | X241 = X242 | X242 = X243 | $less(X242,X243) | $less(0,X241) | $less(X244,sK0) | $less(X243,X244) | X243 = X244 | 0 = sK0 | 0 = X243 (4:30) [resolution 383,387]
[SA] passive: 4377. $less(X245,X246) | X245 = X246 | X246 = X247 | $less(X246,X247) | $less(0,X245) | $less(X248,sK0) | $less(X247,X248) | 0 = X248 | 0 = sK0 | X247 = X248 (4:30) [resolution 383,397]
[SA] passive: 4378. $less(X249,X250) | X249 = X250 | X250 = X251 | $less(X250,X251) | $less(0,X249) | $less(X252,sK0) | $less(X251,X252) | 0 = X251 | 0 = sK0 | 0 = X252 (4:30) [resolution 383,407]
[SA] passive: 4385. $less(X267,X268) | X267 = X268 | X268 = X269 | $less(X268,X269) | $less(sK0,X267) | $less(X269,X270) | $less(X270,alength) | X269 = X270 | sK0 = X270 (4:27) [resolution 383,198]
[SA] passive: 4386. $less(X271,X272) | X271 = X272 | X272 = X273 | $less(X272,X273) | $less(sK0,X271) | $less(X274,alength) | $less(X273,X274) | alength = X274 | sK0 = X273 (4:27) [resolution 383,204]
[SA] passive: 4387. $less(X275,X276) | X275 = X276 | X276 = X277 | $less(X276,X277) | $less(sK0,X275) | $less(X277,X278) | $less(X278,alength) | X277 = X278 | alength = X278 (4:27) [resolution 383,246]
[SA] passive: 4388. $less(X279,X280) | X279 = X280 | X280 = X281 | $less(X280,X281) | $less(sK0,X279) | $less(X282,alength) | $less(X281,X282) | alength = X282 | alength = X281 (4:27) [resolution 383,253]
[SA] passive: 4389. $less(X283,X284) | X283 = X284 | X284 = X285 | $less(X284,X285) | $less(sK0,X283) | $less(X286,alength) | $less(X285,X286) | X285 = X286 | sK0 = X285 (4:27) [resolution 383,388]
[SA] passive: 4390. $less(X287,X288) | X287 = X288 | X288 = X289 | $less(X288,X289) | $less(sK0,X287) | $less(X290,alength) | $less(X289,X290) | sK0 = X289 | sK0 = X290 (4:27) [resolution 383,408]
[SA] passive: 4391. $less(X291,X292) | X291 = X292 | X292 = X293 | $less(X292,X293) | $less(sK0,X291) | $less(X293,0) | alength = X293 | sK0 = X293 | i(l8(sK2(X293))) = X293 (4:30) [resolution 383,2547]
[SA] passive: 4392. $less(X294,X295) | X294 = X295 | X295 = X296 | $less(X295,X296) | $less(sK0,X294) | $less(X296,0) | i(l8(sK2(X296))) = X296 | sK0 = X296 (4:27) [resolution 383,2772]
[SA] active: 3201. v = a(l8(s(sK2(0))),0) | 0 = sK0 (4:11:1) [superposition 349,2831]
[SA] new: 4393. v = a(l8(s(zero)),0) | 0 = sK0 | alength = 0 (7:13) [superposition 3201,3571]
[SA] new: 4394. v = a(l8(s(zero)),0) | 0 = sK0 | 0 = sK0 (5:13) [superposition 3201,3207]
[SA] new: 4395. v = a(l8(sK2(0)),0) | 0 = sK0 | i(l8(sK2(0))) = 0 (5:16) [superposition 3201,440]
[SA] new: 4396. v = a(l11(sK2(0)),0) | 0 = sK0 (5:10) [superposition 3201,57]
[SA] new: 4397. v = a(l8(sK2(0)),0) | i(l8(sK2(0))) = 0 | 0 = sK0 (5:16) [superposition 440,3201]
[SA] new: 4398. v = a(l11(sK2(0)),0) | 0 = sK0 (5:10) [superposition 57,3201]
[SA] new: 4399. v = a(l8(s(zero)),0) | 0 = sK0 (5:10) [duplicate literal removal 4394]
[SA] active: 393. ~$less(X10,X9) | $less(X8,X9) | X8 = X9 | X7 = X9 | $less(X7,X8) | $less(X10,X7) (3:18:1) T [resolution 174,14]
[SA] new: 4400. $less(X0,X1) | X0 = X1 | X1 = X2 | $less(X2,X0) | $less(X3,X2) | $less(X1,X3) | X1 = X3 (4:21) T [resolution 393,15]
[SA] new: 4401. $less(X4,X5) | X4 = X5 | X5 = X6 | $less(X6,X4) | $less(X7,X6) | $less(X5,X7) | X5 = X7 (4:21) T [resolution 393,15]
[SA] new: 4402. $less(X8,X9) | X8 = X9 | X9 = X10 | $less(X10,X8) | $less(X11,X10) | $less(X12,X11) | $less(X9,X12) | X9 = X11 | X9 = X12 (4:27) T [resolution 393,174]
[SA] new: 4403. $less(X13,X14) | X13 = X14 | X14 = X15 | $less(X15,X13) | $less(X16,X15) | $less(X14,X17) | $less(X17,X16) | X16 = X17 | X14 = X16 (4:27) T [resolution 393,174]
[SA] new: 4404. $less(X18,X19) | X18 = X19 | X19 = X20 | $less(X20,X18) | $less(X21,X20) | $less(X22,X21) | $less(X19,X22) | X19 = X22 | X21 = X22 (4:27) T [resolution 393,174]
[SA] new: 4405. $less(X23,X24) | X23 = X24 | X24 = X25 | $less(X25,X23) | $less(X26,X25) | $less(sK0,X26) | $less(X24,alength) | X24 = X26 | sK0 = X24 (4:27) [resolution 393,198]
[SA] new: 4406. $less(X27,X28) | X27 = X28 | X28 = X29 | $less(X29,X27) | $less(X30,X29) | $less(sK0,X30) | $less(X28,alength) | alength = X28 | sK0 = X30 (4:27) [resolution 393,204]
[SA] new: 4407. $less(X31,X32) | X31 = X32 | X32 = X33 | $less(X33,X31) | $less(X34,X33) | $less(sK0,X34) | $less(X32,alength) | X32 = X34 | alength = X32 (4:27) [resolution 393,246]
[SA] new: 4408. $less(X35,X36) | X35 = X36 | X36 = X37 | $less(X37,X35) | $less(X38,X37) | $less(sK0,X38) | $less(X36,alength) | alength = X36 | alength = X38 (4:27) [resolution 393,253]
[SA] new: 4409. $less(X39,X40) | X39 = X40 | X40 = X41 | $less(X41,X39) | $less(X42,X41) | $less(0,X42) | $less(X40,sK0) | X40 = X42 | 0 = sK0 | sK0 = X40 (4:30) [resolution 393,323]
[SA] new: 4410. $less(X43,X44) | X43 = X44 | X44 = X45 | $less(X45,X43) | $less(X46,X45) | $less(0,X46) | $less(X44,sK0) | sK0 = X44 | 0 = sK0 | sK0 = X46 (4:30) [resolution 393,332]
[SA] new: 4411. $less(X47,X48) | X47 = X48 | X48 = X49 | $less(X49,X47) | $less(X50,X49) | $less(0,X50) | $less(X48,sK0) | X48 = X50 | 0 = sK0 | 0 = X50 (4:30) [resolution 393,387]
[SA] new: 4412. $less(X51,X52) | X51 = X52 | X52 = X53 | $less(X53,X51) | $less(X54,X53) | $less(sK0,X54) | $less(X52,alength) | X52 = X54 | sK0 = X54 (4:27) [resolution 393,388]
[SA] new: 4413. $less(X55,X56) | X55 = X56 | X56 = X57 | $less(X57,X55) | $less(X58,X57) | $less(0,X58) | $less(X56,sK0) | 0 = X56 | 0 = sK0 | X56 = X58 (4:30) [resolution 393,397]
[SA] new: 4414. $less(X59,X60) | X59 = X60 | X60 = X61 | $less(X61,X59) | $less(X62,X61) | $less(0,X62) | $less(X60,sK0) | 0 = X62 | 0 = sK0 | 0 = X60 (4:30) [resolution 393,407]
[SA] new: 4415. $less(X63,X64) | X63 = X64 | X64 = X65 | $less(X65,X63) | $less(X66,X65) | $less(sK0,X66) | $less(X64,alength) | sK0 = X66 | sK0 = X64 (4:27) [resolution 393,408]
[SA] new: 4416. $less(X67,X68) | X67 = X68 | X68 = X69 | $less(X69,X67) | $less(X70,X69) | $less(alength,X70) | $less(X68,0) | alength = X68 | X68 = X70 | i(l8(sK2(X68))) = X68 (4:33) [resolution 393,2659]
[SA] new: 4417. $less(X71,X72) | X71 = X72 | X72 = X73 | $less(X73,X71) | $less(X74,X73) | $less(alength,X74) | $less(X72,0) | 0 = X72 | alength = X74 | i(l8(sK2(X74))) = X74 (4:33) [resolution 393,2951]
[SA] new: 4418. $less(X75,alength) | alength = X75 | alength = X76 | $less(X76,X75) | $less(X77,X76) | $less(sK0,X77) | sK0 = X77 (4:21) [resolution 393,106]
[SA] new: 4419. $less(X78,alength) | alength = X78 | alength = X79 | $less(X79,X78) | $less(X80,X79) | $less(0,X80) | alength = X80 | 0 = sK0 (4:24) [resolution 393,178]
[SA] new: 4420. $less(X81,alength) | alength = X81 | alength = X82 | $less(X82,X81) | $less(X83,X82) | $less(sK0,X83) | alength = X83 (4:21) [resolution 393,179]
[SA] new: 4421. $less(X84,alength) | alength = X84 | alength = X85 | $less(X85,X84) | $less(X86,X85) | $less(sK0,X87) | $less(X87,X86) | X86 = X87 | sK0 = X86 (4:27) [resolution 393,198]
[SA] new: 4422. $less(X88,alength) | alength = X88 | alength = X89 | $less(X89,X88) | $less(X90,X89) | $less(sK0,X91) | $less(X91,X90) | alength = X90 | sK0 = X91 (4:27) [resolution 393,204]
[SA] new: 4423. $less(X92,alength) | alength = X92 | alength = X93 | $less(X93,X92) | $less(X94,X93) | $less(sK0,X95) | $less(X95,X94) | X94 = X95 | alength = X94 (4:27) [resolution 393,246]
[SA] new: 4424. $less(X96,alength) | alength = X96 | alength = X97 | $less(X97,X96) | $less(X98,X97) | $less(sK0,X99) | $less(X99,X98) | alength = X98 | alength = X99 (4:27) [resolution 393,253]
[SA] new: 4425. $less(X100,alength) | alength = X100 | alength = X101 | $less(X101,X100) | $less(X102,X101) | $less(0,X102) | alength = X102 | alength = 0 (5:24) [resolution 393,261]
[SA] new: 4426. $less(X103,alength) | alength = X103 | alength = X104 | $less(X104,X103) | $less(X105,X104) | $less(0,X105) | 0 = sK0 | 0 = X105 (5:24) [resolution 393,278]
[SA] new: 4427. $less(X106,alength) | alength = X106 | alength = X107 | $less(X107,X106) | $less(X108,X107) | $less(0,X108) | 0 = sK0 | sK0 = X108 (4:24) [resolution 393,330]
[SA] new: 4428. $less(X109,alength) | alength = X109 | alength = X110 | $less(X110,X109) | $less(X111,X110) | $less(sK0,X112) | $less(X112,X111) | X111 = X112 | sK0 = X112 (4:27) [resolution 393,388]
[SA] new: 4429. $less(X113,alength) | alength = X113 | alength = X114 | $less(X114,X113) | $less(X115,X114) | $less(sK0,X116) | $less(X116,X115) | sK0 = X116 | sK0 = X115 (4:27) [resolution 393,408]
[SA] new: 4430. $less(X117,0) | 0 = X117 | 0 = X118 | $less(X118,X117) | $less(X119,X118) | $less(sK0,X119) | alength = X119 | sK0 = X119 | i(l8(sK2(X119))) = X119 (4:30) [resolution 393,2547]
[SA] new: 4431. $less(X120,0) | 0 = X120 | 0 = X121 | $less(X121,X120) | $less(X122,X121) | $less(sK0,X122) | i(l8(sK2(X122))) = X122 | sK0 = X122 (4:27) [resolution 393,2772]
[SA] new: 4432. $less(X123,0) | 0 = X123 | 0 = X124 | $less(X124,X123) | $less(X125,X124) | $less(alength,X125) | i(l8(sK1(X126,sK2(X125)))) = X126 | alength = X125 | a(main_end,X126) = a(l11(sK2(X125)),X126) (4:38) [resolution 393,2664]
[SA] new: 4433. $less(X127,0) | 0 = X127 | 0 = X128 | $less(X128,X127) | $less(X129,X128) | $less(alength,X130) | $less(X130,X129) | 0 = X129 | alength = X130 | i(l8(sK2(X130))) = X130 (4:33) [resolution 393,2951]
[SA] new: 4434. $less(X131,0) | 0 = X131 | 0 = X132 | $less(X132,X131) | $less(X133,X132) | $less(alength,X134) | $less(X134,X133) | alength = X133 | X133 = X134 | i(l8(sK2(X133))) = X133 (4:33) [resolution 393,2659]
[SA] new: 4435. $less(X135,0) | 0 = X135 | 0 = X136 | $less(X136,X135) | $less(X137,X136) | $less(alength,X137) | i(l8(sK2(X137))) = X137 | alength = X137 (4:27) [resolution 393,2656]
[SA] new: 4436. $less(X138,sK0) | sK0 = X138 | sK0 = X139 | $less(X139,X138) | $less(X140,X139) | $less(0,X140) | sK0 = X140 | 0 = sK0 (4:24) [resolution 393,177]
[SA] new: 4437. $less(X141,sK0) | sK0 = X141 | sK0 = X142 | $less(X142,X141) | $less(X143,X142) | $less(0,X143) | 0 = sK0 | 0 = X143 (4:24) [resolution 393,229]
[SA] new: 4438. $less(X144,sK0) | sK0 = X144 | sK0 = X145 | $less(X145,X144) | $less(X146,X145) | $less(0,X147) | $less(X147,X146) | X146 = X147 | 0 = sK0 | sK0 = X146 (4:30) [resolution 393,323]
[SA] new: 4439. $less(X148,sK0) | sK0 = X148 | sK0 = X149 | $less(X149,X148) | $less(X150,X149) | $less(0,X151) | $less(X151,X150) | sK0 = X150 | 0 = sK0 | sK0 = X151 (4:30) [resolution 393,332]
[SA] new: 4440. $less(X152,sK0) | sK0 = X152 | sK0 = X153 | $less(X153,X152) | $less(X154,X153) | $less(0,X154) | 0 = X154 | sK0 = X154 (4:24) [resolution 393,380]
[SA] new: 4441. $less(X155,sK0) | sK0 = X155 | sK0 = X156 | $less(X156,X155) | $less(X157,X156) | $less(0,X158) | $less(X158,X157) | X157 = X158 | 0 = sK0 | 0 = X158 (4:30) [resolution 393,387]
[SA] new: 4442. $less(X159,sK0) | sK0 = X159 | sK0 = X160 | $less(X160,X159) | $less(X161,X160) | $less(0,X162) | $less(X162,X161) | 0 = X161 | 0 = sK0 | X161 = X162 (4:30) [resolution 393,397]
[SA] new: 4443. $less(X163,sK0) | sK0 = X163 | sK0 = X164 | $less(X164,X163) | $less(X165,X164) | $less(0,X166) | $less(X166,X165) | 0 = X166 | 0 = sK0 | 0 = X165 (4:30) [resolution 393,407]
[SA] new: 4444. $less(X167,sK0) | sK0 = X167 | sK0 = X168 | $less(X168,X167) | $less(X169,X168) | $less(alength,X169) | alength = X169 | 0 = sK0 | i(l8(sK2(X169))) = X169 (4:30) [resolution 393,2948]
[SA] new: 4445. $less(X170,X171) | X170 = X171 | X171 = X172 | $less(X172,X170) | $less(alength,X172) | $less(X171,sK0) | alength = X171 | 0 = sK0 | i(l8(sK2(X171))) = X171 (4:30) [resolution 393,2948]
[SA] new: 4446. $less(X173,X174) | X173 = X174 | X174 = X175 | $less(X175,X173) | $less(alength,X175) | $less(X174,0) | i(l8(sK1(X176,sK2(X174)))) = X176 | alength = X174 | a(main_end,X176) = a(l11(sK2(X174)),X176) (4:38) [resolution 393,2664]
[SA] new: 4447. $less(X177,X178) | X177 = X178 | X178 = X179 | $less(X179,X177) | $less(alength,X179) | $less(X180,0) | $less(X178,X180) | 0 = X180 | alength = X178 | i(l8(sK2(X178))) = X178 (4:33) [resolution 393,2951]
[SA] new: 4448. $less(X181,X182) | X181 = X182 | X182 = X183 | $less(X183,X181) | $less(alength,X183) | $less(X182,X184) | $less(X184,0) | alength = X184 | X182 = X184 | i(l8(sK2(X184))) = X184 (4:33) [resolution 393,2659]
[SA] new: 4449. $less(X185,X186) | X185 = X186 | X186 = X187 | $less(X187,X185) | $less(alength,X187) | $less(X186,0) | i(l8(sK2(X186))) = X186 | alength = X186 (4:27) [resolution 393,2656]
[SA] new: 4450. $less(X188,sK0) | sK0 = X188 | sK0 = X189 | $less(X189,X188) | $less(alength,X189) | alength = sK0 | 0 = sK0 | i(l8(sK2(0))) = 0 (4:27) [resolution 393,2667]
[SA] new: 4451. $less(X190,sK0) | sK0 = X190 | sK0 = X191 | $less(X191,X190) | $less(alength,X191) | alength = 0 | 0 = sK0 | i(l8(sK2(0))) = 0 (4:27) [resolution 393,2665]
[SA] new: 4452. $less(X192,X193) | X192 = X193 | X193 = X194 | $less(X194,X192) | $less(i(l8(p(nl8))),X194) | alength = X193 | $less(X193,alength) | zero = nl8 (4:27) [resolution 393,193]
[SA] new: 4453. $less(X195,X196) | X195 = X196 | X196 = X197 | $less(X197,X195) | $less(i(l8(sK2(X198))),X197) | $less(X198,0) | alength = X196 | $less(X196,alength) | ~$less(X198,alength) (4:30) [resolution 393,2921]
[SA] new: 4454. $less(X199,alength) | alength = X199 | alength = X200 | $less(X200,X199) | $less(i(l8(p(nl8))),X200) | zero = nl8 (4:21) [resolution 393,95]
[SA] new: 4455. $less(X201,alength) | alength = X201 | alength = X202 | $less(X202,X201) | $less(i(l8(p(p(nl8)))),X202) | zero = p(nl8) | zero = nl8 (4:26) [resolution 393,314]
[SA] new: 4456. $less(X203,alength) | alength = X203 | alength = X204 | $less(X204,X203) | $less(i(l8(sK2(X205))),X204) | $less(alength,X205) | alength = X205 | $less(X205,0) (4:27) [resolution 393,2657]
[SA] new: 4457. $less(X206,alength) | alength = X206 | alength = X207 | $less(X207,X206) | $less(i(l8(sK2(X208))),X207) | ~$less(X208,alength) | $less(X208,0) (4:24) [resolution 393,2771]
[SA] new: 4458. $less(X209,0) | 0 = X209 | 0 = X210 | $less(X210,X209) | $less(i(l8(sK2(X211))),X210) | $less(X211,0) | ~$less(X211,alength) | i(l8(sK2(X211))) = i(l8(sK2(i(l8(sK2(X211)))))) (4:36) [resolution 393,2919]
[SA] new: 4459. $less(X212,X213) | X212 = X213 | X213 = X214 | $less(X214,X212) | $less(0,X214) | $less(X213,sK0) | sK0 = X213 | 0 = sK0 (4:24) [resolution 393,177]
[SA] new: 4460. $less(X215,X216) | X215 = X216 | X216 = X217 | $less(X217,X215) | $less(0,X217) | $less(X216,alength) | alength = X216 | 0 = sK0 (4:24) [resolution 393,178]
[SA] new: 4461. $less(X218,X219) | X218 = X219 | X219 = X220 | $less(X220,X218) | $less(0,X220) | $less(X219,sK0) | 0 = sK0 | 0 = X219 (4:24) [resolution 393,229]
[SA] new: 4462. $less(X221,X222) | X221 = X222 | X222 = X223 | $less(X223,X221) | $less(0,X223) | $less(X222,alength) | alength = X222 | alength = 0 (5:24) [resolution 393,261]
[SA] new: 4463. $less(X224,X225) | X224 = X225 | X225 = X226 | $less(X226,X224) | $less(0,X226) | $less(X225,alength) | 0 = sK0 | 0 = X225 (5:24) [resolution 393,278]
[SA] new: 4464. $less(X227,X228) | X227 = X228 | X228 = X229 | $less(X229,X227) | $less(0,X229) | $less(X228,X230) | $less(X230,sK0) | X228 = X230 | 0 = sK0 | sK0 = X230 (4:30) [resolution 393,323]
[SA] new: 4465. $less(X231,X232) | X231 = X232 | X232 = X233 | $less(X233,X231) | $less(0,X233) | $less(X232,alength) | 0 = sK0 | sK0 = X232 (4:24) [resolution 393,330]
[SA] new: 4466. $less(X234,X235) | X234 = X235 | X235 = X236 | $less(X236,X234) | $less(0,X236) | $less(X237,sK0) | $less(X235,X237) | sK0 = X237 | 0 = sK0 | sK0 = X235 (4:30) [resolution 393,332]
[SA] new: 4467. $less(X238,X239) | X238 = X239 | X239 = X240 | $less(X240,X238) | $less(0,X240) | $less(X239,sK0) | 0 = X239 | sK0 = X239 (4:24) [resolution 393,380]
[SA] new: 4468. $less(X241,X242) | X241 = X242 | X242 = X243 | $less(X243,X241) | $less(0,X243) | $less(X244,sK0) | $less(X242,X244) | X242 = X244 | 0 = sK0 | 0 = X242 (4:30) [resolution 393,387]
[SA] new: 4469. $less(X245,X246) | X245 = X246 | X246 = X247 | $less(X247,X245) | $less(0,X247) | $less(X248,sK0) | $less(X246,X248) | 0 = X248 | 0 = sK0 | X246 = X248 (4:30) [resolution 393,397]
[SA] new: 4470. $less(X249,X250) | X249 = X250 | X250 = X251 | $less(X251,X249) | $less(0,X251) | $less(X252,sK0) | $less(X250,X252) | 0 = X250 | 0 = sK0 | 0 = X252 (4:30) [resolution 393,407]
[SA] new: 4471. $less(X253,alength) | alength = X253 | alength = X254 | $less(X254,X253) | $less(0,X254) | alength = 0 (4:18) [resolution 393,244]
[SA] new: 4472. $less(X255,alength) | alength = X255 | alength = X256 | $less(X256,X255) | $less(0,X256) | 0 = sK0 (4:18) [resolution 393,113]
[SA] new: 4473. $less(X257,sK0) | sK0 = X257 | sK0 = X258 | $less(X258,X257) | $less(0,X258) | 0 = sK0 (4:18) [resolution 393,97]
[SA] new: 4474. $less(X259,alength) | alength = X259 | alength = X260 | $less(X260,X259) | $less(sK0,X260) (4:15) [resolution 393,45]
[SA] new: 4475. $less(X261,X262) | X261 = X262 | X262 = X263 | $less(X263,X261) | $less(sK0,X263) | $less(X262,alength) | sK0 = X262 (4:21) [resolution 393,106]
[SA] new: 4476. $less(X264,X265) | X264 = X265 | X265 = X266 | $less(X266,X264) | $less(sK0,X266) | $less(X265,alength) | alength = X265 (4:21) [resolution 393,179]
[SA] new: 4477. $less(X267,X268) | X267 = X268 | X268 = X269 | $less(X269,X267) | $less(sK0,X269) | $less(X268,X270) | $less(X270,alength) | X268 = X270 | sK0 = X270 (4:27) [resolution 393,198]
[SA] new: 4478. $less(X271,X272) | X271 = X272 | X272 = X273 | $less(X273,X271) | $less(sK0,X273) | $less(X274,alength) | $less(X272,X274) | alength = X274 | sK0 = X272 (4:27) [resolution 393,204]
[SA] new: 4479. $less(X275,X276) | X275 = X276 | X276 = X277 | $less(X277,X275) | $less(sK0,X277) | $less(X276,X278) | $less(X278,alength) | X276 = X278 | alength = X278 (4:27) [resolution 393,246]
[SA] new: 4480. $less(X279,X280) | X279 = X280 | X280 = X281 | $less(X281,X279) | $less(sK0,X281) | $less(X282,alength) | $less(X280,X282) | alength = X282 | alength = X280 (4:27) [resolution 393,253]
[SA] new: 4481. $less(X283,X284) | X283 = X284 | X284 = X285 | $less(X285,X283) | $less(sK0,X285) | $less(X286,alength) | $less(X284,X286) | X284 = X286 | sK0 = X284 (4:27) [resolution 393,388]
[SA] new: 4482. $less(X287,X288) | X287 = X288 | X288 = X289 | $less(X289,X287) | $less(sK0,X289) | $less(X290,alength) | $less(X288,X290) | sK0 = X288 | sK0 = X290 (4:27) [resolution 393,408]
[SA] new: 4483. $less(X291,X292) | X291 = X292 | X292 = X293 | $less(X293,X291) | $less(sK0,X293) | $less(X292,0) | alength = X292 | sK0 = X292 | i(l8(sK2(X292))) = X292 (4:30) [resolution 393,2547]
[SA] new: 4484. $less(X294,X295) | X294 = X295 | X295 = X296 | $less(X296,X294) | $less(sK0,X296) | $less(X295,0) | i(l8(sK2(X295))) = X295 | sK0 = X295 (4:27) [resolution 393,2772]
[SA] passive: 4402. $less(X8,X9) | X8 = X9 | X9 = X10 | $less(X10,X8) | $less(X11,X10) | $less(X12,X11) | $less(X9,X12) | X9 = X11 | X9 = X12 (4:27) T [resolution 393,174]
[SA] passive: 4403. $less(X13,X14) | X13 = X14 | X14 = X15 | $less(X15,X13) | $less(X16,X15) | $less(X14,X17) | $less(X17,X16) | X16 = X17 | X14 = X16 (4:27) T [resolution 393,174]
[SA] passive: 4404. $less(X18,X19) | X18 = X19 | X19 = X20 | $less(X20,X18) | $less(X21,X20) | $less(X22,X21) | $less(X19,X22) | X19 = X22 | X21 = X22 (4:27) T [resolution 393,174]
[SA] passive: 4405. $less(X23,X24) | X23 = X24 | X24 = X25 | $less(X25,X23) | $less(X26,X25) | $less(sK0,X26) | $less(X24,alength) | X24 = X26 | sK0 = X24 (4:27) [resolution 393,198]
[SA] passive: 4406. $less(X27,X28) | X27 = X28 | X28 = X29 | $less(X29,X27) | $less(X30,X29) | $less(sK0,X30) | $less(X28,alength) | alength = X28 | sK0 = X30 (4:27) [resolution 393,204]
[SA] passive: 4407. $less(X31,X32) | X31 = X32 | X32 = X33 | $less(X33,X31) | $less(X34,X33) | $less(sK0,X34) | $less(X32,alength) | X32 = X34 | alength = X32 (4:27) [resolution 393,246]
[SA] passive: 4408. $less(X35,X36) | X35 = X36 | X36 = X37 | $less(X37,X35) | $less(X38,X37) | $less(sK0,X38) | $less(X36,alength) | alength = X36 | alength = X38 (4:27) [resolution 393,253]
[SA] passive: 4409. $less(X39,X40) | X39 = X40 | X40 = X41 | $less(X41,X39) | $less(X42,X41) | $less(0,X42) | $less(X40,sK0) | X40 = X42 | 0 = sK0 | sK0 = X40 (4:30) [resolution 393,323]
[SA] passive: 4410. $less(X43,X44) | X43 = X44 | X44 = X45 | $less(X45,X43) | $less(X46,X45) | $less(0,X46) | $less(X44,sK0) | sK0 = X44 | 0 = sK0 | sK0 = X46 (4:30) [resolution 393,332]
[SA] passive: 4411. $less(X47,X48) | X47 = X48 | X48 = X49 | $less(X49,X47) | $less(X50,X49) | $less(0,X50) | $less(X48,sK0) | X48 = X50 | 0 = sK0 | 0 = X50 (4:30) [resolution 393,387]
[SA] passive: 4412. $less(X51,X52) | X51 = X52 | X52 = X53 | $less(X53,X51) | $less(X54,X53) | $less(sK0,X54) | $less(X52,alength) | X52 = X54 | sK0 = X54 (4:27) [resolution 393,388]
[SA] passive: 4413. $less(X55,X56) | X55 = X56 | X56 = X57 | $less(X57,X55) | $less(X58,X57) | $less(0,X58) | $less(X56,sK0) | 0 = X56 | 0 = sK0 | X56 = X58 (4:30) [resolution 393,397]
[SA] passive: 4414. $less(X59,X60) | X59 = X60 | X60 = X61 | $less(X61,X59) | $less(X62,X61) | $less(0,X62) | $less(X60,sK0) | 0 = X62 | 0 = sK0 | 0 = X60 (4:30) [resolution 393,407]
[SA] passive: 4415. $less(X63,X64) | X63 = X64 | X64 = X65 | $less(X65,X63) | $less(X66,X65) | $less(sK0,X66) | $less(X64,alength) | sK0 = X66 | sK0 = X64 (4:27) [resolution 393,408]
[SA] passive: 4416. $less(X67,X68) | X67 = X68 | X68 = X69 | $less(X69,X67) | $less(X70,X69) | $less(alength,X70) | $less(X68,0) | alength = X68 | X68 = X70 | i(l8(sK2(X68))) = X68 (4:33) [resolution 393,2659]
[SA] passive: 4417. $less(X71,X72) | X71 = X72 | X72 = X73 | $less(X73,X71) | $less(X74,X73) | $less(alength,X74) | $less(X72,0) | 0 = X72 | alength = X74 | i(l8(sK2(X74))) = X74 (4:33) [resolution 393,2951]
[SA] passive: 4419. $less(X78,alength) | alength = X78 | alength = X79 | $less(X79,X78) | $less(X80,X79) | $less(0,X80) | alength = X80 | 0 = sK0 (4:24) [resolution 393,178]
[SA] passive: 4421. $less(X84,alength) | alength = X84 | alength = X85 | $less(X85,X84) | $less(X86,X85) | $less(sK0,X87) | $less(X87,X86) | X86 = X87 | sK0 = X86 (4:27) [resolution 393,198]
[SA] passive: 4422. $less(X88,alength) | alength = X88 | alength = X89 | $less(X89,X88) | $less(X90,X89) | $less(sK0,X91) | $less(X91,X90) | alength = X90 | sK0 = X91 (4:27) [resolution 393,204]
[SA] passive: 4423. $less(X92,alength) | alength = X92 | alength = X93 | $less(X93,X92) | $less(X94,X93) | $less(sK0,X95) | $less(X95,X94) | X94 = X95 | alength = X94 (4:27) [resolution 393,246]
[SA] passive: 4424. $less(X96,alength) | alength = X96 | alength = X97 | $less(X97,X96) | $less(X98,X97) | $less(sK0,X99) | $less(X99,X98) | alength = X98 | alength = X99 (4:27) [resolution 393,253]
[SA] passive: 4425. $less(X100,alength) | alength = X100 | alength = X101 | $less(X101,X100) | $less(X102,X101) | $less(0,X102) | alength = X102 | alength = 0 (5:24) [resolution 393,261]
[SA] passive: 4426. $less(X103,alength) | alength = X103 | alength = X104 | $less(X104,X103) | $less(X105,X104) | $less(0,X105) | 0 = sK0 | 0 = X105 (5:24) [resolution 393,278]
[SA] passive: 4427. $less(X106,alength) | alength = X106 | alength = X107 | $less(X107,X106) | $less(X108,X107) | $less(0,X108) | 0 = sK0 | sK0 = X108 (4:24) [resolution 393,330]
[SA] passive: 4428. $less(X109,alength) | alength = X109 | alength = X110 | $less(X110,X109) | $less(X111,X110) | $less(sK0,X112) | $less(X112,X111) | X111 = X112 | sK0 = X112 (4:27) [resolution 393,388]
[SA] passive: 4429. $less(X113,alength) | alength = X113 | alength = X114 | $less(X114,X113) | $less(X115,X114) | $less(sK0,X116) | $less(X116,X115) | sK0 = X116 | sK0 = X115 (4:27) [resolution 393,408]
[SA] passive: 4430. $less(X117,0) | 0 = X117 | 0 = X118 | $less(X118,X117) | $less(X119,X118) | $less(sK0,X119) | alength = X119 | sK0 = X119 | i(l8(sK2(X119))) = X119 (4:30) [resolution 393,2547]
[SA] passive: 4431. $less(X120,0) | 0 = X120 | 0 = X121 | $less(X121,X120) | $less(X122,X121) | $less(sK0,X122) | i(l8(sK2(X122))) = X122 | sK0 = X122 (4:27) [resolution 393,2772]
[SA] passive: 4432. $less(X123,0) | 0 = X123 | 0 = X124 | $less(X124,X123) | $less(X125,X124) | $less(alength,X125) | i(l8(sK1(X126,sK2(X125)))) = X126 | alength = X125 | a(main_end,X126) = a(l11(sK2(X125)),X126) (4:38) [resolution 393,2664]
[SA] passive: 4433. $less(X127,0) | 0 = X127 | 0 = X128 | $less(X128,X127) | $less(X129,X128) | $less(alength,X130) | $less(X130,X129) | 0 = X129 | alength = X130 | i(l8(sK2(X130))) = X130 (4:33) [resolution 393,2951]
[SA] passive: 4434. $less(X131,0) | 0 = X131 | 0 = X132 | $less(X132,X131) | $less(X133,X132) | $less(alength,X134) | $less(X134,X133) | alength = X133 | X133 = X134 | i(l8(sK2(X133))) = X133 (4:33) [resolution 393,2659]
[SA] passive: 4437. $less(X141,sK0) | sK0 = X141 | sK0 = X142 | $less(X142,X141) | $less(X143,X142) | $less(0,X143) | 0 = sK0 | 0 = X143 (4:24) [resolution 393,229]
[SA] passive: 4438. $less(X144,sK0) | sK0 = X144 | sK0 = X145 | $less(X145,X144) | $less(X146,X145) | $less(0,X147) | $less(X147,X146) | X146 = X147 | 0 = sK0 | sK0 = X146 (4:30) [resolution 393,323]
[SA] passive: 4439. $less(X148,sK0) | sK0 = X148 | sK0 = X149 | $less(X149,X148) | $less(X150,X149) | $less(0,X151) | $less(X151,X150) | sK0 = X150 | 0 = sK0 | sK0 = X151 (4:30) [resolution 393,332]
[SA] passive: 4440. $less(X152,sK0) | sK0 = X152 | sK0 = X153 | $less(X153,X152) | $less(X154,X153) | $less(0,X154) | 0 = X154 | sK0 = X154 (4:24) [resolution 393,380]
[SA] passive: 4441. $less(X155,sK0) | sK0 = X155 | sK0 = X156 | $less(X156,X155) | $less(X157,X156) | $less(0,X158) | $less(X158,X157) | X157 = X158 | 0 = sK0 | 0 = X158 (4:30) [resolution 393,387]
[SA] passive: 4442. $less(X159,sK0) | sK0 = X159 | sK0 = X160 | $less(X160,X159) | $less(X161,X160) | $less(0,X162) | $less(X162,X161) | 0 = X161 | 0 = sK0 | X161 = X162 (4:30) [resolution 393,397]
[SA] passive: 4443. $less(X163,sK0) | sK0 = X163 | sK0 = X164 | $less(X164,X163) | $less(X165,X164) | $less(0,X166) | $less(X166,X165) | 0 = X166 | 0 = sK0 | 0 = X165 (4:30) [resolution 393,407]
[SA] passive: 4444. $less(X167,sK0) | sK0 = X167 | sK0 = X168 | $less(X168,X167) | $less(X169,X168) | $less(alength,X169) | alength = X169 | 0 = sK0 | i(l8(sK2(X169))) = X169 (4:30) [resolution 393,2948]
[SA] passive: 4445. $less(X170,X171) | X170 = X171 | X171 = X172 | $less(X172,X170) | $less(alength,X172) | $less(X171,sK0) | alength = X171 | 0 = sK0 | i(l8(sK2(X171))) = X171 (4:30) [resolution 393,2948]
[SA] passive: 4446. $less(X173,X174) | X173 = X174 | X174 = X175 | $less(X175,X173) | $less(alength,X175) | $less(X174,0) | i(l8(sK1(X176,sK2(X174)))) = X176 | alength = X174 | a(main_end,X176) = a(l11(sK2(X174)),X176) (4:38) [resolution 393,2664]
[SA] passive: 4447. $less(X177,X178) | X177 = X178 | X178 = X179 | $less(X179,X177) | $less(alength,X179) | $less(X180,0) | $less(X178,X180) | 0 = X180 | alength = X178 | i(l8(sK2(X178))) = X178 (4:33) [resolution 393,2951]
[SA] passive: 4448. $less(X181,X182) | X181 = X182 | X182 = X183 | $less(X183,X181) | $less(alength,X183) | $less(X182,X184) | $less(X184,0) | alength = X184 | X182 = X184 | i(l8(sK2(X184))) = X184 (4:33) [resolution 393,2659]
[SA] passive: 4452. $less(X192,X193) | X192 = X193 | X193 = X194 | $less(X194,X192) | $less(i(l8(p(nl8))),X194) | alength = X193 | $less(X193,alength) | zero = nl8 (4:27) [resolution 393,193]
[SA] passive: 4453. $less(X195,X196) | X195 = X196 | X196 = X197 | $less(X197,X195) | $less(i(l8(sK2(X198))),X197) | $less(X198,0) | alength = X196 | $less(X196,alength) | ~$less(X198,alength) (4:30) [resolution 393,2921]
[SA] passive: 4454. $less(X199,alength) | alength = X199 | alength = X200 | $less(X200,X199) | $less(i(l8(p(nl8))),X200) | zero = nl8 (4:21) [resolution 393,95]
[SA] passive: 4455. $less(X201,alength) | alength = X201 | alength = X202 | $less(X202,X201) | $less(i(l8(p(p(nl8)))),X202) | zero = p(nl8) | zero = nl8 (4:26) [resolution 393,314]
[SA] passive: 4456. $less(X203,alength) | alength = X203 | alength = X204 | $less(X204,X203) | $less(i(l8(sK2(X205))),X204) | $less(alength,X205) | alength = X205 | $less(X205,0) (4:27) [resolution 393,2657]
[SA] passive: 4457. $less(X206,alength) | alength = X206 | alength = X207 | $less(X207,X206) | $less(i(l8(sK2(X208))),X207) | ~$less(X208,alength) | $less(X208,0) (4:24) [resolution 393,2771]
[SA] passive: 4458. $less(X209,0) | 0 = X209 | 0 = X210 | $less(X210,X209) | $less(i(l8(sK2(X211))),X210) | $less(X211,0) | ~$less(X211,alength) | i(l8(sK2(X211))) = i(l8(sK2(i(l8(sK2(X211)))))) (4:36) [resolution 393,2919]
[SA] passive: 4460. $less(X215,X216) | X215 = X216 | X216 = X217 | $less(X217,X215) | $less(0,X217) | $less(X216,alength) | alength = X216 | 0 = sK0 (4:24) [resolution 393,178]
[SA] passive: 4462. $less(X221,X222) | X221 = X222 | X222 = X223 | $less(X223,X221) | $less(0,X223) | $less(X222,alength) | alength = X222 | alength = 0 (5:24) [resolution 393,261]
[SA] passive: 4463. $less(X224,X225) | X224 = X225 | X225 = X226 | $less(X226,X224) | $less(0,X226) | $less(X225,alength) | 0 = sK0 | 0 = X225 (5:24) [resolution 393,278]
[SA] passive: 4464. $less(X227,X228) | X227 = X228 | X228 = X229 | $less(X229,X227) | $less(0,X229) | $less(X228,X230) | $less(X230,sK0) | X228 = X230 | 0 = sK0 | sK0 = X230 (4:30) [resolution 393,323]
[SA] passive: 4465. $less(X231,X232) | X231 = X232 | X232 = X233 | $less(X233,X231) | $less(0,X233) | $less(X232,alength) | 0 = sK0 | sK0 = X232 (4:24) [resolution 393,330]
[SA] passive: 4466. $less(X234,X235) | X234 = X235 | X235 = X236 | $less(X236,X234) | $less(0,X236) | $less(X237,sK0) | $less(X235,X237) | sK0 = X237 | 0 = sK0 | sK0 = X235 (4:30) [resolution 393,332]
[SA] passive: 4467. $less(X238,X239) | X238 = X239 | X239 = X240 | $less(X240,X238) | $less(0,X240) | $less(X239,sK0) | 0 = X239 | sK0 = X239 (4:24) [resolution 393,380]
[SA] passive: 4468. $less(X241,X242) | X241 = X242 | X242 = X243 | $less(X243,X241) | $less(0,X243) | $less(X244,sK0) | $less(X242,X244) | X242 = X244 | 0 = sK0 | 0 = X242 (4:30) [resolution 393,387]
[SA] passive: 4469. $less(X245,X246) | X245 = X246 | X246 = X247 | $less(X247,X245) | $less(0,X247) | $less(X248,sK0) | $less(X246,X248) | 0 = X248 | 0 = sK0 | X246 = X248 (4:30) [resolution 393,397]
[SA] passive: 4470. $less(X249,X250) | X249 = X250 | X250 = X251 | $less(X251,X249) | $less(0,X251) | $less(X252,sK0) | $less(X250,X252) | 0 = X250 | 0 = sK0 | 0 = X252 (4:30) [resolution 393,407]
[SA] passive: 4477. $less(X267,X268) | X267 = X268 | X268 = X269 | $less(X269,X267) | $less(sK0,X269) | $less(X268,X270) | $less(X270,alength) | X268 = X270 | sK0 = X270 (4:27) [resolution 393,198]
[SA] passive: 4478. $less(X271,X272) | X271 = X272 | X272 = X273 | $less(X273,X271) | $less(sK0,X273) | $less(X274,alength) | $less(X272,X274) | alength = X274 | sK0 = X272 (4:27) [resolution 393,204]
[SA] passive: 4479. $less(X275,X276) | X275 = X276 | X276 = X277 | $less(X277,X275) | $less(sK0,X277) | $less(X276,X278) | $less(X278,alength) | X276 = X278 | alength = X278 (4:27) [resolution 393,246]
[SA] passive: 4480. $less(X279,X280) | X279 = X280 | X280 = X281 | $less(X281,X279) | $less(sK0,X281) | $less(X282,alength) | $less(X280,X282) | alength = X282 | alength = X280 (4:27) [resolution 393,253]
[SA] passive: 4481. $less(X283,X284) | X283 = X284 | X284 = X285 | $less(X285,X283) | $less(sK0,X285) | $less(X286,alength) | $less(X284,X286) | X284 = X286 | sK0 = X284 (4:27) [resolution 393,388]
[SA] passive: 4482. $less(X287,X288) | X287 = X288 | X288 = X289 | $less(X289,X287) | $less(sK0,X289) | $less(X290,alength) | $less(X288,X290) | sK0 = X288 | sK0 = X290 (4:27) [resolution 393,408]
[SA] passive: 4483. $less(X291,X292) | X291 = X292 | X292 = X293 | $less(X293,X291) | $less(sK0,X293) | $less(X292,0) | alength = X292 | sK0 = X292 | i(l8(sK2(X292))) = X292 (4:30) [resolution 393,2547]
[SA] passive: 4484. $less(X294,X295) | X294 = X295 | X295 = X296 | $less(X296,X294) | $less(sK0,X296) | $less(X295,0) | i(l8(sK2(X295))) = X295 | sK0 = X295 (4:27) [resolution 393,2772]
[SA] active: 2978. sK0 != 4 | s(s(s(s(zero)))) = sK2(sK0) (5:11:1) [superposition 2839,186]
[SA] active: 403. ~$less(X10,X8) | $less(X9,X7) | X8 = X9 | X7 = X9 | $less(X7,X8) | $less(X10,X9) (3:18:1) T [resolution 174,14]
[SA] new: 4485. $less(X0,X1) | X0 = X2 | X0 = X1 | $less(X1,X2) | $less(X3,X0) | $less(X2,X3) | X2 = X3 (4:21) T [resolution 403,15]
[SA] new: 4486. $less(X4,X5) | X4 = X6 | X4 = X5 | $less(X5,X6) | $less(X7,X4) | $less(X6,X7) | X6 = X7 (4:21) T [resolution 403,15]
[SA] new: 4487. $less(X8,X9) | X8 = X10 | X8 = X9 | $less(X9,X10) | $less(X11,X8) | $less(X12,X11) | $less(X10,X12) | X10 = X11 | X10 = X12 (4:27) T [resolution 403,174]
[SA] new: 4488. $less(X13,X14) | X13 = X15 | X13 = X14 | $less(X14,X15) | $less(X16,X13) | $less(X15,X17) | $less(X17,X16) | X16 = X17 | X15 = X16 (4:27) T [resolution 403,174]
[SA] new: 4489. $less(X18,X19) | X18 = X20 | X18 = X19 | $less(X19,X20) | $less(X21,X18) | $less(X22,X21) | $less(X20,X22) | X20 = X22 | X21 = X22 (4:27) T [resolution 403,174]
[SA] new: 4490. $less(X23,X24) | X23 = X25 | X23 = X24 | $less(X24,X25) | $less(X26,X23) | $less(sK0,X26) | $less(X25,alength) | X25 = X26 | sK0 = X25 (4:27) [resolution 403,198]
[SA] new: 4491. $less(X27,X28) | X27 = X29 | X27 = X28 | $less(X28,X29) | $less(X30,X27) | $less(sK0,X30) | $less(X29,alength) | alength = X29 | sK0 = X30 (4:27) [resolution 403,204]
[SA] new: 4492. $less(X31,X32) | X31 = X33 | X31 = X32 | $less(X32,X33) | $less(X34,X31) | $less(sK0,X34) | $less(X33,alength) | X33 = X34 | alength = X33 (4:27) [resolution 403,246]
[SA] new: 4493. $less(X35,X36) | X35 = X37 | X35 = X36 | $less(X36,X37) | $less(X38,X35) | $less(sK0,X38) | $less(X37,alength) | alength = X37 | alength = X38 (4:27) [resolution 403,253]
[SA] new: 4494. $less(X39,X40) | X39 = X41 | X39 = X40 | $less(X40,X41) | $less(X42,X39) | $less(0,X42) | $less(X41,sK0) | X41 = X42 | 0 = sK0 | sK0 = X41 (4:30) [resolution 403,323]
[SA] new: 4495. $less(X43,X44) | X43 = X45 | X43 = X44 | $less(X44,X45) | $less(X46,X43) | $less(0,X46) | $less(X45,sK0) | sK0 = X45 | 0 = sK0 | sK0 = X46 (4:30) [resolution 403,332]
[SA] new: 4496. $less(X47,X48) | X47 = X49 | X47 = X48 | $less(X48,X49) | $less(X50,X47) | $less(0,X50) | $less(X49,sK0) | X49 = X50 | 0 = sK0 | 0 = X50 (4:30) [resolution 403,387]
[SA] new: 4497. $less(X51,X52) | X51 = X53 | X51 = X52 | $less(X52,X53) | $less(X54,X51) | $less(sK0,X54) | $less(X53,alength) | X53 = X54 | sK0 = X54 (4:27) [resolution 403,388]
[SA] new: 4498. $less(X55,X56) | X55 = X57 | X55 = X56 | $less(X56,X57) | $less(X58,X55) | $less(0,X58) | $less(X57,sK0) | 0 = X57 | 0 = sK0 | X57 = X58 (4:30) [resolution 403,397]
[SA] new: 4499. $less(X59,X60) | X59 = X61 | X59 = X60 | $less(X60,X61) | $less(X62,X59) | $less(0,X62) | $less(X61,sK0) | 0 = X62 | 0 = sK0 | 0 = X61 (4:30) [resolution 403,407]
[SA] new: 4500. $less(X63,X64) | X63 = X65 | X63 = X64 | $less(X64,X65) | $less(X66,X63) | $less(sK0,X66) | $less(X65,alength) | sK0 = X66 | sK0 = X65 (4:27) [resolution 403,408]
[SA] new: 4501. $less(X67,X68) | X67 = X69 | X67 = X68 | $less(X68,X69) | $less(X70,X67) | $less(alength,X70) | $less(X69,0) | alength = X69 | X69 = X70 | i(l8(sK2(X69))) = X69 (4:33) [resolution 403,2659]
[SA] new: 4502. $less(X71,X72) | X71 = X73 | X71 = X72 | $less(X72,X73) | $less(X74,X71) | $less(alength,X74) | $less(X73,0) | 0 = X73 | alength = X74 | i(l8(sK2(X74))) = X74 (4:33) [resolution 403,2951]
[SA] new: 4503. $less(X75,X76) | alength = X75 | X75 = X76 | $less(X76,alength) | $less(X77,X75) | $less(sK0,X77) | sK0 = X77 (4:21) [resolution 403,106]
[SA] new: 4504. $less(X78,X79) | alength = X78 | X78 = X79 | $less(X79,alength) | $less(X80,X78) | $less(0,X80) | alength = X80 | 0 = sK0 (4:24) [resolution 403,178]
[SA] new: 4505. $less(X81,X82) | alength = X81 | X81 = X82 | $less(X82,alength) | $less(X83,X81) | $less(sK0,X83) | alength = X83 (4:21) [resolution 403,179]
[SA] new: 4506. $less(X84,X85) | alength = X84 | X84 = X85 | $less(X85,alength) | $less(X86,X84) | $less(sK0,X87) | $less(X87,X86) | X86 = X87 | sK0 = X86 (4:27) [resolution 403,198]
[SA] new: 4507. $less(X88,X89) | alength = X88 | X88 = X89 | $less(X89,alength) | $less(X90,X88) | $less(sK0,X91) | $less(X91,X90) | alength = X90 | sK0 = X91 (4:27) [resolution 403,204]
[SA] new: 4508. $less(X92,X93) | alength = X92 | X92 = X93 | $less(X93,alength) | $less(X94,X92) | $less(sK0,X95) | $less(X95,X94) | X94 = X95 | alength = X94 (4:27) [resolution 403,246]
[SA] new: 4509. $less(X96,X97) | alength = X96 | X96 = X97 | $less(X97,alength) | $less(X98,X96) | $less(sK0,X99) | $less(X99,X98) | alength = X98 | alength = X99 (4:27) [resolution 403,253]
[SA] new: 4510. $less(X100,X101) | alength = X100 | X100 = X101 | $less(X101,alength) | $less(X102,X100) | $less(0,X102) | alength = X102 | alength = 0 (5:24) [resolution 403,261]
[SA] new: 4511. $less(X103,X104) | alength = X103 | X103 = X104 | $less(X104,alength) | $less(X105,X103) | $less(0,X105) | 0 = sK0 | 0 = X105 (5:24) [resolution 403,278]
[SA] new: 4512. $less(X106,X107) | alength = X106 | X106 = X107 | $less(X107,alength) | $less(X108,X106) | $less(0,X108) | 0 = sK0 | sK0 = X108 (4:24) [resolution 403,330]
[SA] new: 4513. $less(X109,X110) | alength = X109 | X109 = X110 | $less(X110,alength) | $less(X111,X109) | $less(sK0,X112) | $less(X112,X111) | X111 = X112 | sK0 = X112 (4:27) [resolution 403,388]
[SA] new: 4514. $less(X113,X114) | alength = X113 | X113 = X114 | $less(X114,alength) | $less(X115,X113) | $less(sK0,X116) | $less(X116,X115) | sK0 = X116 | sK0 = X115 (4:27) [resolution 403,408]
[SA] new: 4515. $less(X117,X118) | 0 = X117 | X117 = X118 | $less(X118,0) | $less(X119,X117) | $less(sK0,X119) | alength = X119 | sK0 = X119 | i(l8(sK2(X119))) = X119 (4:30) [resolution 403,2547]
[SA] new: 4516. $less(X120,X121) | 0 = X120 | X120 = X121 | $less(X121,0) | $less(X122,X120) | $less(sK0,X122) | i(l8(sK2(X122))) = X122 | sK0 = X122 (4:27) [resolution 403,2772]
[SA] new: 4517. $less(X123,X124) | 0 = X123 | X123 = X124 | $less(X124,0) | $less(X125,X123) | $less(alength,X125) | i(l8(sK1(X126,sK2(X125)))) = X126 | alength = X125 | a(main_end,X126) = a(l11(sK2(X125)),X126) (4:38) [resolution 403,2664]
[SA] new: 4518. $less(X127,X128) | 0 = X127 | X127 = X128 | $less(X128,0) | $less(X129,X127) | $less(alength,X130) | $less(X130,X129) | 0 = X129 | alength = X130 | i(l8(sK2(X130))) = X130 (4:33) [resolution 403,2951]
[SA] new: 4519. $less(X131,X132) | 0 = X131 | X131 = X132 | $less(X132,0) | $less(X133,X131) | $less(alength,X134) | $less(X134,X133) | alength = X133 | X133 = X134 | i(l8(sK2(X133))) = X133 (4:33) [resolution 403,2659]
[SA] new: 4520. $less(X135,X136) | 0 = X135 | X135 = X136 | $less(X136,0) | $less(X137,X135) | $less(alength,X137) | i(l8(sK2(X137))) = X137 | alength = X137 (4:27) [resolution 403,2656]
[SA] new: 4521. $less(X138,X139) | sK0 = X138 | X138 = X139 | $less(X139,sK0) | $less(X140,X138) | $less(0,X140) | sK0 = X140 | 0 = sK0 (4:24) [resolution 403,177]
[SA] new: 4522. $less(X141,X142) | sK0 = X141 | X141 = X142 | $less(X142,sK0) | $less(X143,X141) | $less(0,X143) | 0 = sK0 | 0 = X143 (4:24) [resolution 403,229]
[SA] new: 4523. $less(X144,X145) | sK0 = X144 | X144 = X145 | $less(X145,sK0) | $less(X146,X144) | $less(0,X147) | $less(X147,X146) | X146 = X147 | 0 = sK0 | sK0 = X146 (4:30) [resolution 403,323]
[SA] new: 4524. $less(X148,X149) | sK0 = X148 | X148 = X149 | $less(X149,sK0) | $less(X150,X148) | $less(0,X151) | $less(X151,X150) | sK0 = X150 | 0 = sK0 | sK0 = X151 (4:30) [resolution 403,332]
[SA] new: 4525. $less(X152,X153) | sK0 = X152 | X152 = X153 | $less(X153,sK0) | $less(X154,X152) | $less(0,X154) | 0 = X154 | sK0 = X154 (4:24) [resolution 403,380]
[SA] new: 4526. $less(X155,X156) | sK0 = X155 | X155 = X156 | $less(X156,sK0) | $less(X157,X155) | $less(0,X158) | $less(X158,X157) | X157 = X158 | 0 = sK0 | 0 = X158 (4:30) [resolution 403,387]
[SA] new: 4527. $less(X159,X160) | sK0 = X159 | X159 = X160 | $less(X160,sK0) | $less(X161,X159) | $less(0,X162) | $less(X162,X161) | 0 = X161 | 0 = sK0 | X161 = X162 (4:30) [resolution 403,397]
[SA] new: 4528. $less(X163,X164) | sK0 = X163 | X163 = X164 | $less(X164,sK0) | $less(X165,X163) | $less(0,X166) | $less(X166,X165) | 0 = X166 | 0 = sK0 | 0 = X165 (4:30) [resolution 403,407]
[SA] new: 4529. $less(X167,X168) | sK0 = X167 | X167 = X168 | $less(X168,sK0) | $less(X169,X167) | $less(alength,X169) | alength = X169 | 0 = sK0 | i(l8(sK2(X169))) = X169 (4:30) [resolution 403,2948]
[SA] new: 4530. $less(X170,X171) | X170 = X172 | X170 = X171 | $less(X171,X172) | $less(alength,X170) | $less(X172,sK0) | alength = X172 | 0 = sK0 | i(l8(sK2(X172))) = X172 (4:30) [resolution 403,2948]
[SA] new: 4531. $less(X173,X174) | X173 = X175 | X173 = X174 | $less(X174,X175) | $less(alength,X173) | $less(X175,0) | i(l8(sK1(X176,sK2(X175)))) = X176 | alength = X175 | a(main_end,X176) = a(l11(sK2(X175)),X176) (4:38) [resolution 403,2664]
[SA] new: 4532. $less(X177,X178) | X177 = X179 | X177 = X178 | $less(X178,X179) | $less(alength,X177) | $less(X180,0) | $less(X179,X180) | 0 = X180 | alength = X179 | i(l8(sK2(X179))) = X179 (4:33) [resolution 403,2951]
[SA] new: 4533. $less(X181,X182) | X181 = X183 | X181 = X182 | $less(X182,X183) | $less(alength,X181) | $less(X183,X184) | $less(X184,0) | alength = X184 | X183 = X184 | i(l8(sK2(X184))) = X184 (4:33) [resolution 403,2659]
[SA] new: 4534. $less(X185,X186) | X185 = X187 | X185 = X186 | $less(X186,X187) | $less(alength,X185) | $less(X187,0) | i(l8(sK2(X187))) = X187 | alength = X187 (4:27) [resolution 403,2656]
[SA] new: 4535. $less(X188,X189) | sK0 = X188 | X188 = X189 | $less(X189,sK0) | $less(alength,X188) | alength = sK0 | 0 = sK0 | i(l8(sK2(0))) = 0 (4:27) [resolution 403,2667]
[SA] new: 4536. $less(X190,X191) | sK0 = X190 | X190 = X191 | $less(X191,sK0) | $less(alength,X190) | alength = 0 | 0 = sK0 | i(l8(sK2(0))) = 0 (4:27) [resolution 403,2665]
[SA] new: 4537. $less(X192,X193) | X192 = X194 | X192 = X193 | $less(X193,X194) | $less(i(l8(p(nl8))),X192) | alength = X194 | $less(X194,alength) | zero = nl8 (4:27) [resolution 403,193]
[SA] new: 4538. $less(X195,X196) | X195 = X197 | X195 = X196 | $less(X196,X197) | $less(i(l8(sK2(X198))),X195) | $less(X198,0) | alength = X197 | $less(X197,alength) | ~$less(X198,alength) (4:30) [resolution 403,2921]
[SA] new: 4539. $less(X199,X200) | alength = X199 | X199 = X200 | $less(X200,alength) | $less(i(l8(p(nl8))),X199) | zero = nl8 (4:21) [resolution 403,95]
[SA] new: 4540. $less(X201,X202) | alength = X201 | X201 = X202 | $less(X202,alength) | $less(i(l8(p(p(nl8)))),X201) | zero = p(nl8) | zero = nl8 (4:26) [resolution 403,314]
[SA] new: 4541. $less(X203,X204) | alength = X203 | X203 = X204 | $less(X204,alength) | $less(i(l8(sK2(X205))),X203) | $less(alength,X205) | alength = X205 | $less(X205,0) (4:27) [resolution 403,2657]
[SA] new: 4542. $less(X206,X207) | alength = X206 | X206 = X207 | $less(X207,alength) | $less(i(l8(sK2(X208))),X206) | ~$less(X208,alength) | $less(X208,0) (4:24) [resolution 403,2771]
[SA] new: 4543. $less(X209,X210) | 0 = X209 | X209 = X210 | $less(X210,0) | $less(i(l8(sK2(X211))),X209) | $less(X211,0) | ~$less(X211,alength) | i(l8(sK2(X211))) = i(l8(sK2(i(l8(sK2(X211)))))) (4:36) [resolution 403,2919]
[SA] new: 4544. $less(X212,X213) | X212 = X214 | X212 = X213 | $less(X213,X214) | $less(0,X212) | $less(X214,sK0) | sK0 = X214 | 0 = sK0 (4:24) [resolution 403,177]
[SA] new: 4545. $less(X215,X216) | X215 = X217 | X215 = X216 | $less(X216,X217) | $less(0,X215) | $less(X217,alength) | alength = X217 | 0 = sK0 (4:24) [resolution 403,178]
[SA] new: 4546. $less(X218,X219) | X218 = X220 | X218 = X219 | $less(X219,X220) | $less(0,X218) | $less(X220,sK0) | 0 = sK0 | 0 = X220 (4:24) [resolution 403,229]
[SA] new: 4547. $less(X221,X222) | X221 = X223 | X221 = X222 | $less(X222,X223) | $less(0,X221) | $less(X223,alength) | alength = X223 | alength = 0 (5:24) [resolution 403,261]
[SA] new: 4548. $less(X224,X225) | X224 = X226 | X224 = X225 | $less(X225,X226) | $less(0,X224) | $less(X226,alength) | 0 = sK0 | 0 = X226 (5:24) [resolution 403,278]
[SA] new: 4549. $less(X227,X228) | X227 = X229 | X227 = X228 | $less(X228,X229) | $less(0,X227) | $less(X229,X230) | $less(X230,sK0) | X229 = X230 | 0 = sK0 | sK0 = X230 (4:30) [resolution 403,323]
[SA] new: 4550. $less(X231,X232) | X231 = X233 | X231 = X232 | $less(X232,X233) | $less(0,X231) | $less(X233,alength) | 0 = sK0 | sK0 = X233 (4:24) [resolution 403,330]
[SA] new: 4551. $less(X234,X235) | X234 = X236 | X234 = X235 | $less(X235,X236) | $less(0,X234) | $less(X237,sK0) | $less(X236,X237) | sK0 = X237 | 0 = sK0 | sK0 = X236 (4:30) [resolution 403,332]
[SA] new: 4552. $less(X238,X239) | X238 = X240 | X238 = X239 | $less(X239,X240) | $less(0,X238) | $less(X240,sK0) | 0 = X240 | sK0 = X240 (4:24) [resolution 403,380]
[SA] new: 4553. $less(X241,X242) | X241 = X243 | X241 = X242 | $less(X242,X243) | $less(0,X241) | $less(X244,sK0) | $less(X243,X244) | X243 = X244 | 0 = sK0 | 0 = X243 (4:30) [resolution 403,387]
[SA] new: 4554. $less(X245,X246) | X245 = X247 | X245 = X246 | $less(X246,X247) | $less(0,X245) | $less(X248,sK0) | $less(X247,X248) | 0 = X248 | 0 = sK0 | X247 = X248 (4:30) [resolution 403,397]
[SA] new: 4555. $less(X249,X250) | X249 = X251 | X249 = X250 | $less(X250,X251) | $less(0,X249) | $less(X252,sK0) | $less(X251,X252) | 0 = X251 | 0 = sK0 | 0 = X252 (4:30) [resolution 403,407]
[SA] new: 4556. $less(X253,X254) | alength = X253 | X253 = X254 | $less(X254,alength) | $less(0,X253) | alength = 0 (4:18) [resolution 403,244]
[SA] new: 4557. $less(X255,X256) | alength = X255 | X255 = X256 | $less(X256,alength) | $less(0,X255) | 0 = sK0 (4:18) [resolution 403,113]
[SA] new: 4558. $less(X257,X258) | sK0 = X257 | X257 = X258 | $less(X258,sK0) | $less(0,X257) | 0 = sK0 (4:18) [resolution 403,97]
[SA] new: 4559. $less(X259,X260) | alength = X259 | X259 = X260 | $less(X260,alength) | $less(sK0,X259) (4:15) [resolution 403,45]
[SA] new: 4560. $less(X261,X262) | X261 = X263 | X261 = X262 | $less(X262,X263) | $less(sK0,X261) | $less(X263,alength) | sK0 = X263 (4:21) [resolution 403,106]
[SA] new: 4561. $less(X264,X265) | X264 = X266 | X264 = X265 | $less(X265,X266) | $less(sK0,X264) | $less(X266,alength) | alength = X266 (4:21) [resolution 403,179]
[SA] new: 4562. $less(X267,X268) | X267 = X269 | X267 = X268 | $less(X268,X269) | $less(sK0,X267) | $less(X269,X270) | $less(X270,alength) | X269 = X270 | sK0 = X270 (4:27) [resolution 403,198]
[SA] new: 4563. $less(X271,X272) | X271 = X273 | X271 = X272 | $less(X272,X273) | $less(sK0,X271) | $less(X274,alength) | $less(X273,X274) | alength = X274 | sK0 = X273 (4:27) [resolution 403,204]
[SA] new: 4564. $less(X275,X276) | X275 = X277 | X275 = X276 | $less(X276,X277) | $less(sK0,X275) | $less(X277,X278) | $less(X278,alength) | X277 = X278 | alength = X278 (4:27) [resolution 403,246]
[SA] new: 4565. $less(X279,X280) | X279 = X281 | X279 = X280 | $less(X280,X281) | $less(sK0,X279) | $less(X282,alength) | $less(X281,X282) | alength = X282 | alength = X281 (4:27) [resolution 403,253]
[SA] new: 4566. $less(X283,X284) | X283 = X285 | X283 = X284 | $less(X284,X285) | $less(sK0,X283) | $less(X286,alength) | $less(X285,X286) | X285 = X286 | sK0 = X285 (4:27) [resolution 403,388]
[SA] new: 4567. $less(X287,X288) | X287 = X289 | X287 = X288 | $less(X288,X289) | $less(sK0,X287) | $less(X290,alength) | $less(X289,X290) | sK0 = X289 | sK0 = X290 (4:27) [resolution 403,408]
[SA] new: 4568. $less(X291,X292) | X291 = X293 | X291 = X292 | $less(X292,X293) | $less(sK0,X291) | $less(X293,0) | alength = X293 | sK0 = X293 | i(l8(sK2(X293))) = X293 (4:30) [resolution 403,2547]
[SA] new: 4569. $less(X294,X295) | X294 = X296 | X294 = X295 | $less(X295,X296) | $less(sK0,X294) | $less(X296,0) | i(l8(sK2(X296))) = X296 | sK0 = X296 (4:27) [resolution 403,2772]
[SA] passive: 4485. $less(X0,X1) | X0 = X2 | X0 = X1 | $less(X1,X2) | $less(X3,X0) | $less(X2,X3) | X2 = X3 (4:21) T [resolution 403,15]
[SA] passive: 4487. $less(X8,X9) | X8 = X10 | X8 = X9 | $less(X9,X10) | $less(X11,X8) | $less(X12,X11) | $less(X10,X12) | X10 = X11 | X10 = X12 (4:27) T [resolution 403,174]
[SA] passive: 4488. $less(X13,X14) | X13 = X15 | X13 = X14 | $less(X14,X15) | $less(X16,X13) | $less(X15,X17) | $less(X17,X16) | X16 = X17 | X15 = X16 (4:27) T [resolution 403,174]
[SA] passive: 4489. $less(X18,X19) | X18 = X20 | X18 = X19 | $less(X19,X20) | $less(X21,X18) | $less(X22,X21) | $less(X20,X22) | X20 = X22 | X21 = X22 (4:27) T [resolution 403,174]
[SA] passive: 4490. $less(X23,X24) | X23 = X25 | X23 = X24 | $less(X24,X25) | $less(X26,X23) | $less(sK0,X26) | $less(X25,alength) | X25 = X26 | sK0 = X25 (4:27) [resolution 403,198]
[SA] passive: 4491. $less(X27,X28) | X27 = X29 | X27 = X28 | $less(X28,X29) | $less(X30,X27) | $less(sK0,X30) | $less(X29,alength) | alength = X29 | sK0 = X30 (4:27) [resolution 403,204]
[SA] passive: 4492. $less(X31,X32) | X31 = X33 | X31 = X32 | $less(X32,X33) | $less(X34,X31) | $less(sK0,X34) | $less(X33,alength) | X33 = X34 | alength = X33 (4:27) [resolution 403,246]
[SA] passive: 4493. $less(X35,X36) | X35 = X37 | X35 = X36 | $less(X36,X37) | $less(X38,X35) | $less(sK0,X38) | $less(X37,alength) | alength = X37 | alength = X38 (4:27) [resolution 403,253]
[SA] passive: 4494. $less(X39,X40) | X39 = X41 | X39 = X40 | $less(X40,X41) | $less(X42,X39) | $less(0,X42) | $less(X41,sK0) | X41 = X42 | 0 = sK0 | sK0 = X41 (4:30) [resolution 403,323]
[SA] passive: 4495. $less(X43,X44) | X43 = X45 | X43 = X44 | $less(X44,X45) | $less(X46,X43) | $less(0,X46) | $less(X45,sK0) | sK0 = X45 | 0 = sK0 | sK0 = X46 (4:30) [resolution 403,332]
[SA] passive: 4496. $less(X47,X48) | X47 = X49 | X47 = X48 | $less(X48,X49) | $less(X50,X47) | $less(0,X50) | $less(X49,sK0) | X49 = X50 | 0 = sK0 | 0 = X50 (4:30) [resolution 403,387]